# SLU12 - Linear Algebra & NumPy, Part 1: Exercise notebook

Welcome! Let's check you Linear Algebra and NumPy knowledge on:

- Vectors and basic operations;
- Matrices and basic operations;
- NumPy arrays.

```python
# run this cell before anything else

# numpy
import numpy as np

# auxiliary stuff (don't worry about it)
from utils import *

# for evaluation purposes
import hashlib
from math import isclose
def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()
```

<img src="./media/what_if.png" width="500">

## 1 - Vectors

In this section of exercises, you'll apply what you learned about vectors definition, basic properties and operations.

### Exercise 1.1

Which **two** of the following objects **do NOT** represent vectors in Python?

a) The integer `1`;

b) The list `[1, 2, 3, 4]`;

c) The NumPy array `np.array([[2,0], [1,0]])`;

d) The NumPy array `np.array([[0], [1]])`;

e) The NumPy array `np.array([[2]])`.

<!-- grade_id: cell-4d00af57c3b45c43, points: 0, grade: false, solution: true -->
```python
# Uncomment the correct answer
#answer = "a and b"
#answer = "a and e"
#answer = "a and c"
#answer = "c and d"
#answer = "a and d"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-73590ab43900dec8, points: 0, grade: true, solution: false -->
```python
assert _hash(answer) == 'a2b790bda1', "Something's not right! Check if all the quantities are vectors and if the shape of the arrays makes sense. Remember all the ways in which you can represent a vector."
```

### Exercise 1.2

Which **one** of the following sentences is **correct**?

a) The vectors $[0, 1]$ and $[2, 0]$ are not orthogonal.

b) The vectors $[0, 1]$ and $[2, 0]$ are collinear (linearly dependent).

c) The vectors $[1, 1]$ and $[2, 2]$ are non-collinear (linearly independent).

d) We can describe any 2D vector we want with some linear combination of the vectors $[0, 1]$ and $[2, 0]$.

e) We can describe any 2D vector we want with some linear combination of the vectors $[1, 1]$ and $[2, 2]$.

<!-- grade_id: cell-c14b14ba5e4db4b5, points: 0, grade: false, solution: true -->
```python
# Uncomment the correct answer
#my_answer = "a"
#my_answer = "b"
#my_answer = "c"
#my_answer = "d"
#my_answer = "e"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4834ff213fc38e55, points: 0, grade: true, solution: false -->
```python
if _hash(my_answer) == "f4e169f8ee":
    print("Review section 1.7 of Learning Notebook 1! What's the angle between these vectors?")
elif _hash(my_answer) == "9350f68d6b" or _hash(my_answer) == "2add4c06d4":
    print("Boris Johnson says no... Check sections 1.4-1.5 of Learning Notebook 1!!")
elif _hash(my_answer) == "94ca9f1720":
    print("Go check sections 1.4-1.5 of Learning Notebook 1. Being 'independent' matters.")
else:
    assert _hash(my_answer) == '838cd7f570', \
    "Did you just make up an answer?\n You're only supposed to uncomment the correct answer."
```

Oh and by the way... zero velocity kitten is still waiting for a refill.
<img src="./media/kitten_stopped.png" width="600"/>

## 2 - Vectors in NumPy

### Exercise 2.1

Look at the NumPy array below:

```Python
    np.array([[1, 0, 0, 0]])
```

Assign the shape and number of array dimensions, respectively, to variables `array_shape` and `array_ndim`.

The ideia is that you should be able to assign the correct tuple to `array_shape` and the correct integer to `array_ndim`, without using any auxiliary functions.

<!-- grade_id: cell-f841135867b9665f, points: 0, grade: false, solution: true -->
```python
# Assign the shape of array to array_shape (remember that the shape is represented by a tuple!!)
# array_shape = ...

# Assign the number of array dimensions to array_ndim
# array_ndim = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b9142cfed551dcd6, points: 0, grade: true, solution: false -->
```python
assert isinstance(array_shape, tuple), "The shape of an array should by a tuple!"
assert isinstance(array_ndim, int), "The number of array dimensions should be an integer!"
assert _hash(array_shape) != '61dcee3c2e', (
    "The shape is almost right! You  just got the order of the axes wrong!") 
assert array_ndim != 4, "Wrong dimensions! You can say that vector is 4D, but the number of array dimensions is not 4!...\n\
    Check section 2.2.2 of Learning Notebook 1 if you get stuck."
assert _hash(array_shape) == '56c0c2870d', (
    "The shape is wrong. Check section 2.2.2 of Learning Notebook 1 if you get stuck.")
assert _hash(array_ndim) == 'cf2d85ea1d', (
    "The number of array dimensions is wrong. Check section 2.2.2 of Learning Notebook 1 if you get stuck.")
```

### Exercise 2.2

**(i)** Use the method [`ndarray.reshape()`](https://numpy.org/doc/1.20/reference/generated/numpy.ndarray.reshape.html) to convert `u` to a **column vector** `u_column`, represented by a 2D array;

**(ii)** Use the transpose attribute (or use `reshape` again) to get the transpose of `u_column` and assign it to `u_row`.

```python
# run this cell first
u = np.array([0, 1, .5, .25])
```

<!-- grade_id: cell-2b9e8b2a78adb3a1, points: 0, grade: false, solution: true -->
```python
# Convert u to a column vector represented by a 2D array 
# u_column = ...

# Convert u_column to a row vector represented by a 2D array
# u_row = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ca99e0d84a924381, points: 0, grade: true, solution: false -->
```python
assert _hash(u_column) != '047849a4d8', "u_column should be a column vector and not a row vector!"
assert _hash(u_row) != '56cc2df426', "u_row should be a row vector and not a column vector!"
assert u_column.ndim == 2 and u_row.ndim == 2, "Your arrays need to be 2D!"
assert (_hash(u_row) == '047849a4d8') and (_hash(u_column) == '56cc2df426'), "Did you change the content of the arrays?"
```

### Exercise 2.3

#### 2.3.1)

Find the dot product between vectors `s` and `t` and assign the result to `scalar`:

```python
# run this cell first
s = np.array([1, -2, -2, 2])
t = np.array([-6, -3, 1, 1])
```

<!-- grade_id: cell-a0d116e8a480c5f8, points: 0, grade: false, solution: true -->
```python
# use numpy to determine the dot product between s and t
# scalar = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f20096f22bb65085, points: 0, grade: true, solution: false -->
```python
assert not isinstance(scalar, np.ndarray), "The result should be a scalar, not a numpy array!"
assert _hash(scalar) == '5b4838043f', "Wrong! :("
```

#### 2.3.2)

Based on the result, what can you conclude about the vectors `s` and `t`? (uncomment the correct answer)

a) `s` is the transpose of `t`;

b) `s` and `t` are orthogonal;

c) `s` and `t` can describe the space of all 4D vectors;

d) `s` and `t` are collinear.

<!-- grade_id: cell-19349b4d58067f1c, points: 0, grade: false, solution: true -->
```python
# Uncommment the correct answer
#correct_answer = "a"
#correct_answer = "b"
#correct_answer = "c"
#correct_answer = "d"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6c738dac851f1902, points: 0, grade: true, solution: false -->
```python
assert _hash(correct_answer) != 'f4e169f8ee', "That's not correct! What is the transpose of a vector?"
assert _hash(correct_answer) != '2add4c06d4', "Are you sure? Read section 1.5 of Learning Notebook 1..."
assert _hash(correct_answer) != '838cd7f570', "Not really. Can you find a scalar that transforms s into t, or vice-versa?"
assert _hash(correct_answer) == '9350f68d6b', "Don't write anything, you just need to uncomment the correct answer!'"
```

## 3 - Matrices

### Exercise 3.1

Which **two** of the following sentences are **not correct**?

a) A symmetric matrix is always equal to its transpose;

b) An identity matrix of size $n\times n$ is a square matrix where all entries are zero;

c) A square matrix is a matrix whose elements are square roots;

d) Given any number of matrices of the same size, we can add them in any order we want.

<!-- grade_id: cell-d47d31cfa2b6e1e1, points: 0, grade: false, solution: true -->
```python
# Uncommment the correct answer
#me_answers = "a and b"
#me_answers = "b and c"
#me_answers = "a and d"
#me_answers = "c and d"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-da76ed413334f76f, points: 0, grade: true, solution: false -->
```python
assert _hash(me_answers) != 'f9f1efd901', "That's not right. One of the chosen sentences is actually correct.\nYou need to choose 2 incorrect sentences."
assert _hash(me_answers) != '153d3d6799', "That's not right. Both sentences are actually correct. Check section 3 of Learning Notebook 2 if you don't understand why!"
assert _hash(me_answers) != 'd5ea9fb2d7', "That's not right. Remember that you can add matrices in any order you want,\n as long as they have the same size, so the last sentence is true."
assert _hash(me_answers) == '3f46356968', "Don't write any code! Just uncomment the correct answer!"
```

### Exercise 3.2

What's the result of the following operation? (don't write any code!!)

$$2 \cdot \begin{bmatrix}0 & 1\\ 1 & 0\end{bmatrix}
- 1 \cdot \begin{bmatrix}1 & 0\\ 0 & 1\end{bmatrix}$$

a) $\begin{bmatrix}1 & 0\\ 0 & 1\end{bmatrix}$

b) $\begin{bmatrix}2 & 0\\ 0 & 2\end{bmatrix}$

c) $\begin{bmatrix}-1 & 2\\ 2 & -1\end{bmatrix}$

d) $\begin{bmatrix}1 & 1\\ 1 & 1\end{bmatrix}$

<!-- grade_id: cell-7f5c602a732f3401, points: 0, grade: false, solution: true -->
```python
# Uncommment the correct answer
#correct_matrix = "a"
#correct_matrix = "b"
#correct_matrix = "c"
#correct_matrix = "d"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-96bcd1509b60dfb9, points: 0, grade: true, solution: false -->
```python
assert _hash(correct_matrix) != 'f4e169f8ee', "Wrong! Check section 3.3 of Learning Notebook 2 if you don't understand why."
assert _hash(correct_matrix) != '9350f68d6b', "Wrong! Check section 3.3 of Learning Notebook 2 if you don't understand why."
assert _hash(correct_matrix) != '838cd7f570', "Wrong! Check section 3.3 of Learning Notebook 2 if you don't understand why."
assert _hash(correct_matrix) == '2add4c06d4', "Don't write any code, you just need to uncomment the correct answer!'"
```

## 4 - NumPy arrays and matrices

*Everything* is about matrices! Images are nothing but matrices.

On the cell below we'll open a greyscale image of a cute panda, reading it into a matrix of pixels.

Because we're dealing with a greyscale image, we can represent it by a 2D ndarray of shape `(height, width)`, where each entry is a value in the range `0` to `255`, corresponding to a pixel in the image.

```python
# read cute panda image into 2D numpy array
panda = load_panda().astype(int)

# show image
plot_img(panda);

# print array shape (image size in pixels)
print("panda greyscale image array:", panda.shape)

# preview some rows
print("\nFirst 5 rows and 5 columns:", panda[:5, :5])
```

### Exercise 4.1

Let's invert the panda colours!! 🐼

Invert the image colours by performing the operation `255 - value` for each element `value` in the array `panda`.

**Hint**: NumPy will perform the addition (or subtraction) between scalars and arrays in an element-wise fashion. In Linear Algebra terms, this is the equivalent of subtracting the matrix `panda` from a matrix of the same size, where all entries are equal to `255`.

<!-- grade_id: cell-412fbe3220933365, points: 0, grade: false, solution: true -->
```python
# invert the panda!
# opposite_panda = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bdeb943297e36545, points: 0, grade: true, solution: false -->
```python
assert isinstance(panda, np.ndarray) and panda.shape == (460, 460) and panda.min() == 0 \
    and panda.max() == 255 and  _hash(panda[0, 10]) == '751b860653', \
    "OMG you changed the panda variable!! Reload by running: panda = load_panda()"
assert opposite_panda.shape == (460, 460), "The panda_inverted array should have the same shape as the panda array!"
assert _hash(int(opposite_panda[320, 400])) == '0d1d9d8d66', "Something's not right."
assert _hash(int(opposite_panda[10, 100])) == 'f7e25634c3', "Something's not right."
print("\n- You've just turned the panda into its negative!\n")
print("- But don't worry, pandas are always positive, even on the negative side!\n\n")
print("(words by a wise LDSA devops)")
plot_pandas(opposite_panda)
```

### Exercise 4.2

Let's transpose the panda!

Find the transpose of `panda` and assign it to `transposed_panda`.

<!-- grade_id: cell-e69cabcf5fd3a3d0, points: 0, grade: false, solution: true -->
```python
# find the transpose of panda
# transposed_panda = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eb2d7aea1d5af18a, points: 0, grade: true, solution: false -->
```python
assert isinstance(panda, np.ndarray) and panda.shape == (460, 460) and panda.min() == 0 \
    and panda.max() == 255 and  _hash(panda[0, 10]) == '751b860653', \
    "OMG you changed the panda variable!! Reload it with the code: panda = load_panda()"
assert _hash(panda[0, 10]) == '751b860653', "Please don't change the panda variable!! Create a copy instead!"
assert _hash(transposed_panda) == '06f31209c9', "Wrong! What did you do to the panda?"
print("\nCORRECT! Oh no look, the panda is falling!!...")
plot_pandas(transposed_panda)
```

### Exercise 4.3

Create a *binary* panda called `binary_panda` where:
- You set the value of all entries that are **greater than** `100` to `255`;
- You set the value of the remaining entries to `0`.

<!-- grade_id: cell-90690c356008f10a, points: 0, grade: false, solution: true -->
```python
# create a boolean mask to filter all entries > 100
# mask = ...

# create a COPY of panda using the method .copy() - DO NOT USE "binary_panda = panda"!!
# binary_panda = ...

# use the mask to set all entries above 100 in binary_panda to 255
# ...

# set all the other values to 0 (tip: use ~mask)
# ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-66809eeaf15f2563, points: 0, grade: true, solution: false -->
```python
assert isinstance(panda, np.ndarray) and panda.shape == (460, 460) and panda.min() == 0 \
    and panda.max() == 255 and  _hash(panda[0, 10]) == '751b860653', \
    "OMG you changed the panda variable!! Reload it with the code: panda = load_panda()"
assert _hash(int(binary_panda[0, 10])) == '6bafb3698f', "Wrong! What are you doing to the panda?"
assert _hash(int(binary_panda[0, 100])) == '5b4838043f', "Wrong! What are you doing to the panda?"
print("\nCORRECT! Panda is now literally only black and white...")
plot_pandas(binary_panda)
```

### Exercise 4.4

4 pandas is better than 1!! 🐼🐼🐼🐼

Because we like tile effect, let's create an image with 2x2 pandas, concatenating our `panda` array into an image with 4 pandas, which should look like this:

<img src="./media/tile_pandas.png">

<!-- grade_id: cell-dc5b19629feb3e59, points: 0, grade: false, solution: true -->
```python
## create a 2D array with 4 concatenated pandas

# it should have 2 pandas on the first column (first concatenation step)
# column_pandas = ...

# and now concatenate 2 columns of pandas side by side (second concatenation step)
# tile_pandas = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f3ec3088aac4e295, points: 0, grade: true, solution: false -->
```python
assert isinstance(panda, np.ndarray) and panda.shape == (460, 460) and panda.min() == 0 \
    and panda.max() == 255 and  _hash(panda[0, 10]) == '751b860653', \
    "OMG you changed the panda variable!! Reload it with the code: panda = load_panda()"
assert _hash(column_pandas[720,400]) == "9963b5511e", "Wrong. Check the concatenation step for column_pandas!"
assert tile_pandas.shape == (920, 920), "The image size is not right!"
assert tile_pandas.shape != (1840, 460), "Did you concatenate  along the correct axis?"
assert _hash(tile_pandas[300,400]) == '6c5db95765', "Wrong! What did you do to the panda?"
assert _hash(tile_pandas[260,720]) == 'f4bbfd7be2', "Wrong! What did you do to the panda?"
assert _hash(tile_pandas[600,420]) == '0d1d9d8d66', "Wrong! What did you do to the panda?"
assert _hash(tile_pandas[600,800]) == 'ede52244b6', "Wrong! What did you do to the panda?"
print("\nCORRECT! Tiled pandas!")
plot_pandas(tile_pandas)
```

---

### Dizzy panda?

It's a linear combination of 3 pandas...

```python
# plot panda matrix and a linear combination of panda matrices!!
panda_combination = 0.5*panda + 0.5*panda.T + 0.5*panda[::-1]
plot_pandas(panda_combination)
```

---

### Exercise 4.5 (boss level!!)

Do you like puzzles? 

Let's create a 100-piece puzzle together with our panda image! We'll split the tasks:

- Your task will be to slice the image into 100 pieces!

- My task will be to randomly mix the pieces you've built and display the final result.

What I need from you is:

- `pieces` - a **list** with all the 100 pieces inside it, represented as 2D array slices of our `panda` array, each of shape `(side, side)` (I don't care about the order of the pieces inside the list...).

- `side` - the length of each side of the 100 square pieces.


I know I know, your part is harder! 😁😝

Note the following:

- The end size of your array should be the same as the original `panda`;

- Each pixel is simply an element of the array.

Let's do this!

<!-- grade_id: cell-b0f72f2efb1e9e6c, points: 0, grade: false, solution: true -->
```python
# Your turn

#### what's the side length per square piece?
# convert the result to an INTEGER!!
# side = ...


#### now slice the image into 100 square pieces!
# initialize a list to store the 100 numpy array pieces
pieces = []
# this is where you get creative!
# try to find a way to iterate through the panda array
# and save all the different 100 array pieces in pieces_array
#   - you could use a for loop, range, list comprehension, append,... whatever you want!
#   - if you can, avoid using any NumPy methods
# make sure you save all the 100 different pieces (numpy arrays) inside the list pieces
# ...
# pieces = ...


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9b30ec79470704e7, points: 0, grade: true, solution: false -->
```python
import math

assert isinstance(panda, np.ndarray) and panda.shape == (460, 460) and panda.min() == 0 \
    and panda.max() == 255 and  _hash(panda[0, 10]) == '751b860653', \
    "OMG did you change the panda variable? How dare you make that to a panda! Reload it with the code: panda = load_panda()"
assert isinstance(pieces, list), "Wrong! The variable pieces should be a list of numpy arrays!"
assert len(pieces) == 100, "Wrong! You should have 100 pieces, no less no more!"

# check if all pieces are in the list
piece_hashes = set([_hash(piece) for piece in pieces])
assert piece_hashes == {
 '02be9d6ce2', '07848c537e', '0a2e28f440', '0af647d93c', '11d3b39a80', '13cfdfa2e9', '1a8344eb52', '1bd4a336b6',
 '1c601bdd89', '22bd495b3c', '26c7dc1d4c', '289180411b', '28a6e32b35', '2a6c71b8df', '2a729e70c9', '2b02c76a03',
 '2bfa86fc88', '2e95608216', '2f248cb1c8', '35fb7da6db', '3697193023', '376315dbbd', '3eb8f59cd2', '40128ced6e',
 '40d7edd99f', '41badc88cd', '430813f00e', '4397304874', '44c70a7d78', '470fb34b6d', '4979a9b607', '4e5bd0b84b',
 '4fe0fd7b10', '503419682b', '5048f5fd22', '51857ecdd5', '51a28f3aca', '57da59b8f8', '5835ab0b91', '58600ed82f',
 '5ac9427587', '5c84000093', '5d45dcf52a', '5d62bde0ab', '6756b0fc03', '68eb166d6a', '6ac4fb00e9', '6ddbc03b88',
 '706b4ae8ab', '747944d48d', '7675185716', '79af3029d6', '7caa71e4f3', '7da3a9d8c2', '847027c964', '84cea37281',
 '86457b6483', '89a69e1c4c', '8a052d1712', '8b8d7648ac', '97862ad710', '983979c8d2', '9989f29398', '9b82e1e2ee',
 '9c867619b4', 'a026d50550', 'a7f7329ad3', 'a94443d646', 'acf93bc106', 'b0e0eb6ba7', 'b572037421', 'b6be9e18e8',
 'b922487e57', 'bb418e046e', 'bedf071ed4', 'bf00c8999b', 'bfe2706cfc', 'c48ec85d2f', 'c4ae780575', 'c5b00870fb',
 'c7901354c2', 'cfc67efe75', 'd7251c83d4', 'd88fc392a9', 'da29f647d5', 'db3609fbb9', 'dcb1466706', 'de172beccc',
 'de79a642a2', 'deeee80582', 'df8e136de6', 'e15c1c9bc9', 'e5642145ee', 'e806e500ee', 'e975f82de7', 'ed3651f66f',
 'f0236ab443', 'f12bd5363a', 'f2e8138d40', 'fe3fbd881d'}, "No! :( Something's not right... The panda is crying."

print("You did it, you're awesome! Let me shuffle the pieces now. :)")
```

```python
# My turn - run this cell!!

# shuffle the pieces
import random
pieces_shuffled = pieces.copy()
random.shuffle(pieces_shuffled)
```

The moment of truth...

<!-- grade_id: cell-650545c11bb3b476, points: 0, grade: true, solution: false -->
```python
# SHOW THE PUZZLE!!
try:
    plot_img(np.concatenate([np.concatenate(pieces_shuffled[n:n+10], axis=1) for n in range(0, 100, 10)]))
    print("\nYey!! There it is. Our 100-piece puzzle!")
except:
    print("Did you pass all the asserts? Either you didn't or I didn't catch your errors.")
    print("Now our panda is sad...")
    plot_sad_panda()
```

### Last but not least, submit your work!

To submit your work, fill your slack ID in the `slack_id` variable (as a string).

Example: `slack_id = "x-men"`

Help: if you forgot your slack ID, [read this](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c).

<!-- grade_id: cell-1138f06eb745452a, points: 0, grade: false, solution: true -->
```python
# Submit your work!

#slack_id = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c9a56f8e4aaafdb4, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 12
submit(slack_id, slu)
```

---