```python
import csv
import json
import numpy as np
import utils
```

### Exercise 1 (5 points)

This exercise is a quiz with 14 multiple answer questions.
In the cell below, we've declared a dictionary called `answers`.

You should fill in that dictionary with your answers, using as keys the question number, like `question_x`, and as values the number from `1` to `4` that corresponds to the right answer.
For each question, you should provide only one answer, i.e, the dict values should have type **int**.

For example, if you want to answer Question 1 with choice number 2, then you do:
```
answers["question_1"] = 2
```

```python
answers = {}
```

#### Question 1
In the terminal of a Unix based machine, which of these commands will take you to your home directory?
1. cd ..
2. ls ..
3. cd ~/
4. ls ~/

<!-- grade_id: cell-98b6eeed9f5eb544, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()

# answers["question_1"] = ...
```

#### Question 2
In the terminal of a Unix based machine, which of these commands would you use to see the entire content of a file named `data.csv`?
1. head -5 data.csv
2. tail -5 data.csv
3. cat data.csv
4. ls -la data.csv

<!-- grade_id: cell-f75169d418f280ba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()
# answers["question_2"] = ...
```

#### Question 3
After running one of these commands in the terminal of a Unix based machine, you won't have a file named `data.csv` in your current directory. Which command is it?
1. cp dataset.csv data.csv
2. rm data.csv
3. touch data.csv
4. mv dataset.csv data.csv

<!-- grade_id: cell-d0ee1b32bb06368b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_3"] = ...
```

#### Question 4
In the terminal of a Unix based machine, if you want to see the size of the files in the current directory, in a human readable format, which command is the best option?
1. ls
2. ls .
3. ls -la
4. ls -lh

<!-- grade_id: cell-b48ab5ac85e1bf69, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_4"] = ...
```

#### Question 5
There are certain types of files that you usually don't want to have in your git repositories, like cache files, or local configuration files that contain passwords. In order to have git not tracking these files, you can use the `.gitignore` file. Say you're currently in the root directory of your project, and you have a `.gitignore` file there. In the terminal of a Unix based machine, which of these commands should you use to list the files and see the `.gitignore` file on that list?
1. ls -la
2. ls .
3. ls -lh
4. ls -l

<!-- grade_id: cell-d53190f9ce4dedbe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_5"] = ...
```

#### Question 6
Consider the following git commands: `add`, `commit`, `status`, `pull`, `push`. Which is the most natural order to use them?
1. `pull` > `status` > `add` > `commit` > `push`
2. `status` > `commit` > `add` > `push` > `pull`
3. `push` > `add` > `status` > `commit` > `pull`
4. `pull` > `push` > `status` > `add` > `commit`

<!-- grade_id: cell-44e7a38840f54f6c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_6"] = ...
```

#### Question 7
How do you change from the `master` branch to an existing `feature_dev` branch?
1. git checkout master
2. git checkout master feature_dev
3. git checkout -b feature_dev
4. git checkout feature_dev

<!-- grade_id: cell-8f6345404f52ff67, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_7"] = ...
```

#### Question 8
Which of these commands should you use to find the hash of your previous commit?
1. git status
2. git log
3. git diff
4. git commit

<!-- grade_id: cell-67e3948e70b7e795, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_8"] = ...
```

#### Question 9
Which one of these is false:
1. You can use a tuple as a dictionary key.
2. You can add new elements to a tuple.
3. Tuples can have duplicated elements.
4. You can define a tuple with length 0.

<!-- grade_id: cell-fd120313a3617c0f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_9"] = ...
```

#### Question 10
If you want to incorporate the master branch into another branch called `dev`, which 2 commands should you run sequentially?
1. git checkout dev; git merge master
2. git merge master; git checkout dev
3. git checkout master; git merge dev
4. git merge dev; git checkout master

<!-- grade_id: cell-03d49ee095b85596, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_10"] = ...
```

#### Question 11
Which command should you use to bring back changes that you've stashed?
1. git stash
2. git stash pop
3. git revert
4. git revert apply

<!-- grade_id: cell-3753c7fa82eba09c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_11"] = ...
```

#### Question 12
Which one of these files will you always find in the root directory of a Python package?
1. README.md
2. requirements.txt
3. \_\_init\_\_.py
4. setup.py

<!-- grade_id: cell-6c9b515dc4382694, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_12"] = ...
```

#### Question 13
How do you import function `my_function` from module `my_module`?
1. import my_function
2. import my_module
3. from my_function import my_module
4. from my_module import my_function

<!-- grade_id: cell-63f33fece990776a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_13"] = ...
```

#### Question 14
Which one of these should not raise an Error?
1. '4' + 'None'
2. 4 + None
3. 4 + 'None'
4. '4' + None

<!-- grade_id: cell-81aca48ed55e858b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()

# answers["question_14"] = ...
```

<!-- grade_id: cell-db5effd691135306, points: 5, grade: true, solution: false -->
```python
# Run this cell to see if your answers are correct
utils.exercise_1(answers)
```

### Exercise 2

Consider a csv file with four columns: `city`, `latitude`, `longitude`, `country`.

File `cities_locations.csv` is an example of such file and was included in the zip file you downloaded with this notebook, so you can preview it using some command line arguments.

#### Part I (2 points)

Implement a function that reads a file with the same format as the `cities_locations.csv` file, and stores the data in a dictionary with the following structure:

```
{
    "city_name": {
        "idx": row number in the input file, not considering the header,
               i.e, the first row with data values has idx=0 (type: int),
        "lat": latitude value (type: float),
        "lng": longitude value (type: float),
        "country": country (type: str)
    },
}

```

The function should be called `read_cities`, receives an argument called `filepath` (which is the file that the function should read the data from), and returns the dictionary that was created.

<!-- grade_id: cell-622c6adc26159244, points: 0, grade: false, solution: true -->
```python
def read_cities(filepath):
    """
    Reads the file in filepath, parses it and returns the data in a dictionary.

    Parameters:
    filepath (str): Path to the input file to be parsed.

    Returns:
    cities_parsed_data (dict): Cities locations dataset stored as a dictionary, where the keys are city names
          and the values are data about the city.
    """
    
    # YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6c24a4ad6f30ce55, points: 2, grade: true, solution: false -->
```python
cities_locations_data = read_cities('cities_locations.csv')

assert len(cities_locations_data) == 10

assert 'Barcelona' in cities_locations_data
barcelona = cities_locations_data['Barcelona']
assert barcelona['idx'] == 5
assert barcelona['country'] == 'Spain'
np.testing.assert_almost_equal(barcelona['lat'], 41.3833, decimal=3)
np.testing.assert_almost_equal(barcelona['lng'], 2.1834, decimal=3)


from more_tests import test_exercise_2_I
test_exercise_2_I(read_cities)
```

#### Part II (3 points)

We want to calculate the distance in km between each pair of cities from our dataset, and store this in a numpy array.

In order to calculate these distances, you should use the **geodesic** function from the python package [geopy](https://pypi.org/project/geopy/).

You'll need to read the package's docs, in order to install it in your local environment and to learn how to use the **geodesic** function. The version of geopy that is installed in our grader is `1.22.0`, so, to make sure that you'll have the same behaviour in your local environment and in our grader, make sure to install the same version.

The matrix with the distances will have shape `(n_cities, n_cities)`, where `n_cities` is the number of cities in the dataset. Each city's position in the matrix's rows and columns should be the `idx` value from the dataset. 

Implement a function called `create_dist_matrix` that receives as input a dictionary like the one we created in the first part of this exercise, and returns the numpy array representing the distances matrix.

<!-- grade_id: cell-7854488ee93d67b6, points: 0, grade: false, solution: true -->
```python
from geopy.distance import geodesic

def create_dist_matrix(cities_data):
    """
    Computes the geodesic distance between each pair of cities in the input dataset, and returns a matrix with the
    computed distances.

    Parameters:
    cities_data (dict): Dictionary with the city locations dataset.

    Returns:
    matrix (np.array): Matrix with shape (n_cities, n_cities) where the value in position (i, j) represents
        the geodesic distance in km between city_i and city_j.
    """
    
    # YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-56eb4a74f3544d03, points: 3, grade: true, solution: false -->
```python
cities_locations_data = {
    'Rome': {'idx': 0, 'lat': 41.896, 'lng': 12.4833, 'country': 'Italy'},
    'Milan': {'idx': 1, 'lat': 45.47, 'lng': 9.205, 'country': 'Italy'},
    'Lisbon': {'idx': 2, 'lat': 38.7227, 'lng': -9.1449, 'country': 'Portugal'},
    'Porto': {'idx': 3, 'lat': 41.15, 'lng': -8.62, 'country': 'Portugal'},
    'Madrid': {'idx': 4, 'lat': 40.4, 'lng': -3.6834, 'country': 'Spain'},
    'Barcelona': {'idx': 5, 'lat': 41.3833, 'lng': 2.1834, 'country': 'Spain'}
}

dists = create_dist_matrix(cities_locations_data)

assert type(dists) == np.ndarray
assert dists.shape == (6, 6)

for i in range(dists.shape[0]):
    assert dists[i][i] == 0
    
np.testing.assert_almost_equal(dists[0][1], 476.93, decimal=1)
np.testing.assert_almost_equal(dists[0][5], 859.51, decimal=1)


from more_tests import test_exercise_2_II
test_exercise_2_II(create_dist_matrix)
```

#### Part III (2 points)

Create a function called `furthest_cities` that finds in our distances matrix, the pair of cities furthest away from each other.

It should return a tuple of city indexes, where the first one is smaller than the second.

<!-- grade_id: cell-29b0e001d3a16b34, points: 0, grade: false, solution: true -->
```python
def furthest_cities(dist_matrix):
    """
    Finds the pair of cities that is furthest away from each other.

    Parameters:
    dist_matrix (np.array): Matrix that stores geodesic distances between cities.

    Returns:
    furthest_cities_idx (tuple): Indexes of the two cities. The first element in the tuple is the lower index. 
    """
    
    # YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7a082faf67626b85, points: 2, grade: true, solution: false -->
```python
# Creating a random distances matrix for this test
np.random.seed(19)
dists = np.random.rand(5, 5) * 1000
for i in range(dists.shape[0]):
    dists[i][i] = 0
    for j in range(dists.shape[1]):
        if j > i:
            dists[i][j] = dists[j][i]

idx_1, idx_2 = furthest_cities(dists)
assert idx_1 == 1
assert idx_2 == 4


from more_tests import test_exercise_2_III
test_exercise_2_III(furthest_cities)
```

### Exercise 3

In this exercise, we'll use object oriented programming concepts to model [Pokemons](https://en.wikipedia.org/wiki/Pok%C3%A9mon).

In case you're not familiar with Pokemon, they are wild (imaginary) creatures. These are some examples.

<img src="./images/pokemons.png" style="width:800px">

When they're living in the wild, Pokemon can be captured by Pokemon Trainers. Once they belong to a Trainer, they will obey the Trainer's commands. They're usually sent out to non-lethal battles against other Pokémon, in order to gain experience and level up.

<img src="./images/battle.png" style="width:300px">

When they reach certain levels, they can undergo a form of metamorphosis and transform into a similar but stronger species of Pokémon: this is called evolution.

<img src="./images/charmander_evolves.jpg" style="width:350px">

#### Part I (4 points)

Your first assignment is to implement a class that represents a `Pokemon`.

You'll need to store the following information about a Pokemon:
* `name`: the Pokemon's name.
* `max_health`: the number of health points that this Pokemon has with full health.
* `speed`: a measure of how fast this Pokemon is. Faster Pokemons usually attack first in battles.
* `hp`: current number of health points. Pokemons may lose health points during battles.
* `level`: Pokemon's current level. This is a measure of the Pokemon's experience. Pokemons in higher levels have more chances of winning battles.

Some additional information:
* All the stats points described above (`max_health`, `speed`, `hp`, `level`) should be measured with non-negative integers.
* When a Pokemon is born, its `level` is always 1 and its `hp` is always the same as `max_health`, but its `name`, `max_health`and `speed` vary from Pokemon to Pokemon.


Our Pokemon class should implement 4 methods, described below.

##### 1. Method `is_knocked_out`

During battles, Pokemon take damage, which translates into losing health points.
Method `is_knocked_out` receives no arguments, and checks whether the Pokemon is knocked out by checking if the Pokemon's `hp` is equal to 0.
This method should return a bool.


##### 2. Method `level_up`

When a Pokemon wins a battle, it will level up.
Method `level_up` receives no arguments, and doesn't return anything.

This method should:
* Increase the Pokemon's `level` by one.
* Increase the Pokemon's `max_health` by 20 points.
* Increase the Pokemon's `speed` by 10%. But speed must be an integer, so round it down.
* It shouldn't change the Pokemon's `hp`!


##### 3. Method `take_damage`

As explained above, during battles, Pokemon take damage, which translates into losing health points.
Method `take_damage` receives as argument the integer `damage_points`, and doesn't return anything.

This method should decrease the Pokemon's `hp` by `damage_points`. Make sure that the Pokemon's `hp`doesn't fall below 0.


##### 4. Method `attack_damage`

This method calculates the number of damage points that our Pokemon's attack will inflict on an enemy Pokemon, during a battle. More experienced Pokemon inflict more damage and have higher chances of having successful attacks than less experienced Pokemon. 
Method `attack_damage` receives as argument an `enemy_pokemon` (which is another instance of the Pokemon class) and returns an integer representing the number of damage points.

This method should compute the damage points in the following way:
* Create a variable `level_diff` that stores the difference in levels of the two Pokemon. If our Pokemon is more experienced, `level_diff` should be positive. If our Pokemon is less experienced, `level_diff` should be 0.
* Create a variable `max_level` that stores the maximum level between the levels of the two Pokemon.
* Create a variable called `p_success`, that represents the probability of success of your Pokemon's attack, and set it to: $$0.5 + \frac{level\_diff}{2 * max\_level}$$
* Create a variable called `attack_success`, that represents whether the attack was successful or not. Calculate its value by drawing one sample from a binomial distribution: use numpy's [binomial function](https://numpy.org/doc/stable/reference/random/generated/numpy.random.binomial.html), with parameters `n=1` and `p=p_success`. This function will output a 0 (which means the attack was not successful) or a 1 (which means the attack was successful). 
* Create a variable called `damage_points`. Calculate its value by multiplying your Pokemon's `level` by `attack_success`.
* Return `damage_points`.


<!-- grade_id: cell-25e4812b96a967d8, points: 0, grade: false, solution: true -->
```python
class Pokemon:
    # YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-92c4e256e70cfb10, points: 4, grade: true, solution: false -->
```python
score = 0

try:
    pika = Pokemon(name='Pikachu', max_health=20, speed=5)
    assert pika.name == 'Pikachu'
    assert pika.max_health == 20
    assert pika.speed == 5
    assert pika.hp == pika.max_health
    assert pika.level == 1
    assert not pika.is_knocked_out()
except AssertionError:
    pass
else:
    score += 1

try:
    pika.level_up()
    assert pika.max_health == 40
    assert pika.hp == 20
    assert pika.speed == 5
    assert pika.level == 2
except AssertionError:
    pass
else:
    score += 1


try:
    enemy = Pokemon(name='Squirtle', max_health=10, speed=10)
    np.random.seed(42)
    assert pika.attack_damage(enemy) == 2
    assert pika.attack_damage(enemy) == 0
    assert pika.attack_damage(enemy) == 2

    np.random.seed(13)
    assert pika.attack_damage(enemy) == 0
    assert pika.attack_damage(enemy) == 2
    assert pika.attack_damage(enemy) == 0
except AssertionError:
    pass
else:
    score += 1

try:
    pika.take_damage(5)
    assert pika.hp == 15

    pika.take_damage(20)
    assert pika.hp == 0
    assert pika.is_knocked_out()
except AssertionError:
    pass
else:
    score += 1

if score == 0:
    raise AssertionError("Not enough correct answers to score points :(")

score
```

#### Part II (2 points)

Now we'll implement a battle between two Pokemon.

Write a function called `battle` that receives as arguments `p1` and `p2`, both instances of the Pokemon class, and returns the Pokemon that wins the battle and the Pokemon that loses the battle (in this order!).

A battle is a sequence of attacks. In the first turn, the Pokemon who attacks is the fastest (if they have the same speed, `p1` attacks first); the slowest Pokemon attacks in the second turn, and then the fastest Pokemon attacks again. No Pokemon is allowed to attack twice, they always switch places as the attacker.

In an attack, the attacker inflicts as many damage points as indicated by the `attack_damage` method. Remember that this method has a random component, so you have to explicitly call it everytime the Pokemon attacks.
The defender may suffer damage points. In order to record that, you should use the `take_tamage` method.

The battle ends as soon as one of the Pokemon is knocked out.

Grader Tip: If you see a `KeyboardInterrupt` error on the grader feedback, that is because your cell is taking too long to run, which is probably due to an infinite loop.

<!-- grade_id: cell-6ee8d592008072c6, points: 0, grade: false, solution: true -->
```python
def battle(p1, p2):
    """
    Represents a battle between two Pokemon, where each Pokemon attacks at a time.

    Parameters:
    p1 (Pokemon): A Pokemon fighting in the battle.
    p2 (Pokemon): The other Pokemon fighting in the battle.

    Returns:
    winner (Pokemon): The winner Pokemon
    loser (Pokemon): The loser Pokemon
    """    
    # YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b055a1cab77f20ad, points: 2, grade: true, solution: false -->
```python
pika = Pokemon(name='Pikachu', max_health=20, speed=5)
squirtle = Pokemon(name='Squirtle', max_health=10, speed=10)

np.random.seed(19)
winner, loser = battle(pika, squirtle)

assert winner.name == "Pikachu"
assert loser.name == "Squirtle"
assert winner.hp == 15
assert loser.hp == 0

from more_tests import test_exercise_3_II
test_exercise_3_II(battle, Pokemon)
```

### Exercise 4

Consider that you have a matrix, that represents a field, where you can go for a walk.

In order to refer to a specific position in the matrix, we'll use the notation `(i,j)`, where `i` represents the row number, `j` represents the column number, using 0-indexing (the first row has index 0).

You start walking at some position in the matrix, and the value on that position will guide you towards your next position. The possible values in the matrix are:
* `l`: means that you'll move to the column on your **left**, in the same row
* `r`: means that you'll move to the column on your **right**, in the same row
* `u`: means that you'll move to the row **above**, in the same column
* `d`: means that you'll move to the row **below**, in the same column


You'll be able to keep walking until you're guided to a position outside the boundaries of the field.

Let's see an example. This is a field.

<code>[['r', 'r', 'd', 'l', 'u'],
 ['r', 'd', 'l', 'd', 'd'],
 ['u', 'r', 'd', 'r', 'l'],
 ['l', 'd', 'd', 'u', 'r']]
</code>

If you start your walk in position (0,0), your path will be:
* `(0,0)`
* `r` => `(0,1)`
* `r` => `(0,2)`
* `d` => `(1,2)`
* `l` => `(1,1)`
* `d` => `(2,1)`
* `r` => `(2,2)`
* `d` => `(3,2)`
* `d` => `end`

In red you can see the positions where you passed.

<code>[[<font color='red'>'r'</font>, <font color='red'>'r'</font>, <font color='red'>'d'</font>, 'l', 'u'],
 ['r', <font color='red'>'d'</font>, <font color='red'>'l'</font>, 'd', 'd'],
 ['u', <font color='red'>'r'</font>, <font color='red'>'d'</font>, 'r', 'l'],
 ['l', 'd', <font color='red'>'d'</font>, 'u', 'r']]
</code>

The length of your path is the number of positions where you have been in the field.
The path in the example has length 8.

***

#### Part I (1 point)

Write a function, named `walk` that receives as arguments:
* a numpy matrix that represents a **field**, and
* a tuple `(i, j)` that represents a **starting position**,

and returns the length of the path that you get if you start walking at the starting position and continue until you leave the field.

You can assume that the field is not empty, i.e, it has at least one position, and that the starting position is within the field.


#### Part II (1 point)

For an extra point, add the following validations to your function.

If you find an unknown instruction while walking your path, raise a **ValueError** with message: `Invalid instruction found on position (i, j)`, where `i` and `j` are the row and column numbers of the position where you found the unknown instruction. If the field has unknown instructions, but you don't stumble upon them while walking, no exception should be raised.

And be careful! Your field might tell you to start walking in circles! If during your walk, you find yourself in a position where you have been before, you should raise a **ValueError** with message: `"I'm running in circles"`

Grader Tip: If you see a `KeyboardInterrupt` error on the grader feedback, that is because your cell is taking too long to run, which is probably due to an infinite loop.

<!-- grade_id: cell-5a7eaacbb7097d35, points: 0, grade: false, solution: true -->
```python
def walk(field, start):
    """
    Returns the length of the path that begins on a starting position and continues according to the
    field's instructions.

    Parameters:
    field (np.array): Matrix with instructions that describes a field where one can go for a walk.
    start (tuple):    Position where to start walking.
                      The first value is the row index and the second is the column index.

    Returns:
    walk_length (int): Length of the walk.
    
    Raises:
    ValueError: if an invalid instruction is found.
    ValueError: if the walk has a loop.
    """
    
    # YOUR CODE HERE
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2eb102f666a51257, points: 1, grade: true, solution: false -->
```python
# test field with single position
field = np.array([['d']])
assert walk(field, (0, 0)) == 1


# test field with successful walk
field = np.array([
    ['d', 'l', 'd'],
    ['d', 'u', 'l'],
    ['d', 'r', 'd'],
    ['d', 'd', 'd'],
])
assert walk(field, (1, 2)) == 7
assert walk(field, (2, 2)) == 2


# test another field with successful walk
field = np.array([
    ['r', 'r', 'd', 'l', 'u'],
    ['r', 'd', 'l', 'd', 'd'],
    ['u', 'r', 'd', 'r', 'l'],
    ['l', 'd', 'd', 'u', 'r'],
])
assert walk(field, (0, 0)) == 8
```

<!-- grade_id: cell-8da0084a934a316d, points: 1, grade: true, solution: false -->
```python
# test field with loop
field = np.array([
        ['r', 'r', 'd', 'l', 'u'],
        ['r', 'd', 'l', 'd', 'd'],
        ['u', 'r', 'd', 'r', 'l'],
        ['l', 'd', 'd', 'u', 'r'],
    ])
try:
    walk(field, (1, 3))
    raise AssertionError("Should raise a ValueError when running in circles!")
except ValueError as e:
    assert str(e) == "I'm running in circles"
else:
    raise AssertionError("Should have raised the right error type!")

# test another field with loop
field = np.array([
    ['d', 'l', 'l'],
    ['d', 'w', 'u'],
    ['d', 'u', 'u'],
    ['r', 'r', 'u'],
])
try:
    walk(field, (3, 1))
    raise AssertionError("Should raise a ValueError when running in circles!")
except ValueError as e:
    assert str(e) == "I'm running in circles"
else:
    raise AssertionError("Should have raised the right error type!")

# test field with error
field = np.array([
    ['d', 'l', 'l'],
    ['d', 'w', 'u'],
    ['d', 'u', 'u'],
    ['r', 'r', 'u'],
])
try:
    walk(field, (2, 1))
    raise AssertionError("Should raise a ValueError when an invalid instruction is found!")
except ValueError as e:
    assert str(e) == "Invalid instruction found on position (1, 1)"
else:
    raise AssertionError("Should have raised the right error type!")

# test another field with error
field_with_error = np.array([
    ['r', 'x', 'd', 'l', 'u'],
    ['r', 'd', 'l', 'd', 'd'],
    ['u', 'r', 'd', 'r', 'l'],
    ['l', 'd', 'd', 'u', 'r'],
])
try:
    walk(field_with_error, (0, 0))
    raise AssertionError("Should raise a ValueError when an invalid instruction is found!")
except ValueError as e:
    assert str(e) == "Invalid instruction found on position (0, 1)"
else:
    raise AssertionError("Should have raised the right error type!")
```

---

### Last but not least, submit your work!

To submit your work, fill your slack ID in the `slack_id` variable (as a string).

Example: `slack_id = "x-men"`

Help: if you forgot your slack ID, [read this](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c).

<!-- grade_id: cell-af941850b12aa77e, points: 0, grade: false, solution: true -->
```python
# Submit your work!

#slack_id = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-92bbc7cc332a69f5, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 18_3
submit(slack_id, slu)
```

---