## SLU08 - Functions Intermediate: Exercise notebook

```python
# Students don't need to worry about this cell
# Just make sure that the first thing you do is
# to run it or the rest of the grading won't work!

import math
import hashlib
import inspect

def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()
```

## Exercise 1

Open up your favorite editor (word, google docs, notes, whatever) and write answers to the following questions in your own words:

1. Explain what local scope is
1. Explain what global scope is

You should aim to become capable of explaining it (even if you don't understand it yet) using the right words and concepts. Having a good idea yourself about what all of these are is essential.

<!-- grade_id: cell-feba5c25d67069c0, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 2

Consider the following code:


```python
def hello_world():
    a_var = 10
    return a_var

another_var = 100
```

Identify which variable is in a local scope and which is in the global scope

<!-- grade_id: cell-9e7c4f368ba822c3, points: 0, grade: false, solution: true -->
```python
# assign the string values of either 'local' or 'global'

# Assign a_var_scope the value of 'local' or 'global'
a_var_scope = ''

# Assign another_var_scope the value of 'local' or 'global'
another_var_scope = ''

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b9fe7c14fb27eeb7, points: 4, grade: true, solution: false -->
```python
assert _hash(a_var_scope) == '702153dc06', 'incorrect scope, try again!'
assert _hash(another_var_scope) == 'ec9f60ee7b', 'incorrect scope, try again!'
```

## Exercise 3

In this exercise, you'll learn the difference between scopes regardless of the name of a variable. The name of the variable is not the important part, it's all about being inside or outside of a function when the variable is assigned it's value.

```python
outside_var = 100

def hello_world_complicated():
    print(outside_var) # first print
    a_var = 10
    outside_var = 200
    print(outside_var) # second print
```

<!-- grade_id: cell-ee96a24cc19231ea, points: 0, grade: false, solution: true -->
```python
# The print(outside_var) which is the first line in the function is printing
# the value of a global or locally scoped variable
# Assign outside_var_first_print_scope the value of 'local' or 'global'

outside_var_first_print_scope = ''

# The print(outside_var) which is the last line in the function is printing
# the value of a global or locally scoped variable
# Assign outside_var_second_print_scope the value of 'local' or 'global'

outside_var_second_print_scope = ''


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-15f4aacaf3f4acf2, points: 4, grade: true, solution: false -->
```python
assert _hash('a' + outside_var_first_print_scope) == 'fc6003ae12', 'incorrect scope, try again!'
assert _hash('b' + outside_var_second_print_scope) == '650bfd57ec', 'incorrect scope, try again!'
```

## Exercise 4

Now you'll see first hand how local and global scope can clash (and potentially create problems) when a local function accesses variables in the global scope.

### DO NOT RUN THIS CODE

The point of this exercise is to read and understand what is happening between the variables and the scopes. You can easily run this code and come up with the answers but if you do it before thinking about it first, you'll be robbing yourself of a valueable learning opportunity!

```py
a_list = [1, 2, 3]

def hello():
    another_list = a_list
    another_list.append(4)
    
my_value_1 = a_list + [5]
hello()
my_value_2 = a_list
```

<!-- grade_id: cell-1d991da3b4fc577d, points: 0, grade: false, solution: true -->
```python
# Fill in the values after ONLY READING the code.
# DO NOT RUN THE CELL ABOVE

my_value_1 = [] # fill in the values in the list for my_value_1
my_value_2 = [] # fill in the values in the list for my_value_2

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8f1b5ee5de1cd405, points: 4, grade: true, solution: false -->
```python
assert _hash(my_value_1) == '90cd62da52', 'not quite, follow the code and the concepts!'
assert _hash(my_value_2) == '76d2a49034', 'not quite, follow the code and the concepts!'
```

## Exercise 5

In this exercise, you will create a variable in the global scope and then use it inside the local scope
of another function. 

Create a function with the following specifications:
1. Create a global variable called `a_number` and assign it the value of `2`
1. Create a function called `use_global_scope` that takes one argument
1. The argument is called `multiply_by` and should be an integer
1. If `multiply_by` is not an integer, the function returns `None`
1. This function should return the value of the global variable `a_number` multiplied by `multiply_by`

<!-- grade_id: cell-4d7ca7728f0a1ab1, points: 0, grade: false, solution: true -->
```python
# first define a global variable called a_number

# next define your function use_global_scope that multiplies a_number by multiply_by


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7eeb010ffbd8b2d1, points: 4, grade: true, solution: false -->
```python
sig = inspect.signature(use_global_scope)
source = inspect.getsource(use_global_scope)
params = sig.parameters

assert 'use_global_scope' in globals(), 'did you name the function correctly?'

assert len(sig.parameters) == 1, 'your function should take one argument'
assert 'multiply_by' in params, 'you must have an argument called multiply_by'
assert 'a_number' in globals(), 'be sure to define a_number in the global scope'
assert use_global_scope(10) == 20, 'check if your implementation fulfills all specifications'
assert use_global_scope('hello') == None, 'check if your implementation fulfills all specifications'
assert 'a_number' in source, 'be sure to use a_number in side of the function'
assert '=' not in source, 'you do not need to assign any values to any variables inside of this function'
assert '10' not in source, 'are you cheating by hard-coding?'
assert '20' not in source, 'no cheating by hard-coding stuff!'
```

## Exercise 6

In this exercise, you will create a function that multiplies all the elements of a list
by a given number

1. Create a function named `mult_by`
1. It takes one positional argument called `my_list`
1. It take one keyword argument called `multiply_by` and is a number
    1. The default value is `2`
1. The function must multiply all elements of the list by `multiply_by` and return the result(which will also be a list)

<!-- grade_id: cell-a20188b198ba965e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c9c233f3b6f3e544, points: 4, grade: true, solution: false -->
```python
assert 'mult_by' in globals(), 'did you name the function correctly?'

sig = inspect.signature(mult_by)
source = inspect.getsource(mult_by)
params = sig.parameters

assert len(sig.parameters) == 2, 'your function should define two arguments'
assert 'my_list' in params, 'you must have an argument called my_list'
assert 'multiply_by' in params, 'you must have an argument called multiply_by'
assert sig.parameters['my_list'].default == inspect._empty, 'my_list should be a positional argument'
assert sig.parameters['multiply_by'].default == 2, 'multiply_by should be a keyword argument with a default value of 2'
assert (
    '8' not in source
    and
    '10' not in source
), 'are you cheeeeeating?'


assert mult_by([4, 5]) == [8, 10], 'wrong return value, check if your implementation fulfills all specifications!'
assert mult_by([4, 5], multiply_by=1) == [4, 5], 'wrong return value, check if your implementation fulfills all specifications!'
```

<!-- grade_id: cell-8fadee9bf39bad5e, points: 0, grade: false, solution: true -->
```python
# Submit your work!

#slack_id = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3136d76d011bcbc9, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 8

submit(slack_id, slu)
```