# SLU18 - Exam preparation II

## Batch 3 #Wave 2

This is the 2nd Python exam from batch 3. Exam Duration: 3h

```python
import numpy as np
import pytest
import io
import math
import utils
```

## Exercise 1

**Orthogonal arrays**

In this exercise, we want to write a function that validates if two arrays are orthogonal.

For two arrays to be orthogonal, we need their inner product to be zero.

In this exercise, we'll represent arrays as tuples.

But let's do this in steps!

-------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------
**1.1 Validate tuples**

First, define a function called **validate_tuples** that receives two arguments, x and y, and checks if both x and y are tuples. 

- If both x and y are tuples, the function should return **True**.

- Otherwise, it should raise a ValueError with message: "The input arguments must be tuples."

- If the points don’t have the same number of coordinates, raise a ValueError with message: "Cannot compute distance between points with different number of coordinates."


**Grader tips**

In this exercise, you should just define function **validate_tuples**, you should not invoke the function!
Make sure that:
- the name of the function is exactly what is requested
- the message in the ValueError is exactly what is requested
- you don't have any print statements in your code

<!-- grade_id: cell-4d00af57c3b45c44, points: 0, grade: false, solution: true -->
```python
# create function called validate_tuples

# def validate_tuples(a, b)
# ...
# ...

### BEGIN SOLUTION
def validate_tuples(a, b):
    if (not isinstance(a, tuple)) or (not isinstance(b, tuple)):
        raise ValueError("The input arguments must be tuples.")

    return True
### END SOLUTION
```

**1.2 Validate dimensions** 

In this second step, we want to check if two arrays have the same number of dimensions. For that, you'll have to define a function called **validate_dimensions**, that:

- Receives two arguments, x and y, which are tuples (and represent our two arrays)
- First, check if both x and y have at least two dimensions, i.e, check if the length of the tuples is at least two. In case they don't, raise a ValueError with message: "Input arrays must have at least two dimensions."
- Then, check if the two arrays have the same number of dimensions, i.e, check if the two tuples have the same length. In case they don't, "chop out" the last dimensions from the longer array, until they have the same number of dimensions
- Finally return a list with the resulting arrays

<!-- grade_id: cell-4d00af57c3b45c45, points: 0, grade: false, solution: true -->
```python
# create function called validate_dimensions

# def validate_dimensions(a, b)
# ...
# ...


### BEGIN SOLUTION

def validate_dimensions(a, b):
    if (len(a) < 2) or (len(b) < 2):
        raise ValueError("Input arrays must have at least two dimensions.")

    min_dims = min(len(a), len(b))

    return [a[:min_dims], b[:min_dims]]

### END SOLUTION
```

**1.3 Inner product** 

In this third step, you'll have to define a function called **inner_product**, that receives two arguments, x and y, which are n-dimensional tuples, and computes their inner product as:


$$<x,y> = \sum_{n=1}^N {x_iy_i}={x_1y_1 + x_2y_2 + ... +  x_n y_n}$$

<!-- grade_id: cell-4d00af57c3b45c46, points: 0, grade: false, solution: true -->
```python
# create function called inner_product

# def inner_product(a, b)
# ...
# ...

### BEGIN SOLUTION

def inner_product(a, b):
    return sum([x_i * y_i for x_i, y_i in zip(a, b)])

### END SOLUTION
```

**1.4 Are orthogonal** 

In the final step we'll put everything together.

Define a function called **are_orthogonal**, that receives two arguments, x and y, and does the following:

- Calls function **validate_tuples** with arguments x and y
- Calls function **validate_dimensions** with arguments x and y
- Calls function **inner_product** with arguments x' and y', which were returned by function **validate_dimensions**
- Check if the inner product of the two arrays is 0. Return **True** if it is, and **False** if it's not.
In this step we've defined functions **validate_tuples**, **validate_dimensions** and **inner_product** for you. Thus you can call them even if you didn't pass the previous steps!

<!-- grade_id: cell-4d00af57c3b45c43, points: 0, grade: false, solution: true -->
```python
# create function called are_orthogonal

# def are_orthogonal(x, y)
# ...
# ...


### BEGIN SOLUTION
def are_orthogonal(x, y):
    validate_tuples(x, y)
    x, y = validate_dimensions(x, y)

    return inner_product(x, y) == 0
### END SOLUTION
```

<!-- grade_id: cell-0da2a9cc467a82cc, points: 2, grade: true, solution: false -->
```python

test_cases = [
    {'x': (-1, -2, -3), 'y': (4, 5, 6), 'expected_result': False},
    {'x': (-1, 0), 'y': (0, 1), 'expected_result': True},
    {'x': (-1, 0, 3), 'y': (0, 1), 'expected_result': True},
    {'x': (-1, 0, 3), 'y': (0, 1, 4, 5), 'expected_result': False}
]

for t in test_cases:

    assert are_orthogonal(t['x'], t['y']) == t['expected_result'] , "The update from your function was: \n "+ are_orthogonal(t['x'], t['y']) +"We were expecting: \n"+t['expected_result']

x = (1,)
y = (2,)
with pytest.raises(ValueError) as e:
    are_orthogonal(x, y)
assert str(e.value) == 'Input arrays must have at least two dimensions.'

x = (1,)
y = (2, -4, -5)

with pytest.raises(ValueError) as e:
    are_orthogonal(x, y)
assert str(e.value) == 'Input arrays must have at least two dimensions.'


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

---

## Exercise 2

**Sorting word**

**1.1. Select strings**

In this exercise, we want to write a function that sorts a list of words.

Again, we'll do this in steps.

First, define a function called **select_strings** that:

- receives a list as argument, let's call it list1
- creates a new list with the elements from list1 that are strings, let's call the new list list2
- returns list2


**Grader tips**

In this exercise, you should just define function select_strings, you should not invoke the function!

Make sure that:

- the name of the function is exactly what is requested
- you don't have any print statements in your code

<!-- grade_id: cell-f089360ee63f8d14, points: 0, grade: false, solution: true -->
```python
# create function called select_strings

# def select_strings(l)
# ...
# ...

### BEGIN SOLUTION
def select_strings(l):
    return [i for i in l if isinstance(i, str)]
### END SOLUTION
```

<!-- grade_id: cell-84ee2ec237b2a4c2, points: 1, grade: true, solution: false -->
```python
x = ["hello", "1", 4, "goodbye", None, (1, 2)]
expected_x = ["hello", "1", "goodbye"]
assert select_strings(x) == expected_x

x = list(range(3))
expected_x = []
assert select_strings(x) == expected_x


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**1.2 Clean strings.**

Now define a function called **clean_string** that:

- receives a string as argument, let's call it x
- removes any white spaces at the beginning and end of x 
- lowercases x
- returns the clean version of x


<!-- grade_id: cell-f089360ee63f8d15, points: 0, grade: false, solution: true -->
```python
# create function called clean_string

# def clean_string(l)
# ...
# ...


### BEGIN SOLUTION
def clean_string(x):
    return x.strip().lower()
### END SOLUTION
```

<!-- grade_id: cell-84ee2ec237b2a4c3, points: 1, grade: true, solution: false -->
```python
x = "  HELLO"
expected_x = "hello"
assert clean_string(x) == expected_x

x = "   ByE1  "
expected_x = "bye1"
assert clean_string(x) == expected_x

x = "ldsa rules"
expected_x = x
assert clean_string(x) == expected_x

x = ""
expected_x = x
assert clean_string(x) == expected_x


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**1.3. Sort string list**

Now define a function called **sort_string_list** that:

- receives a list of strings as argument, let's call it list1
- gets a sorted version of list1, in alphabetical order, **without changing the original list**, let's call it - list2
- returns list2

<!-- grade_id: cell-f089360ee63f8d16, points: 0, grade: false, solution: true -->
```python
# create function called sort_string_list

# def sort_string_list(l)
# ...
# ...

### BEGIN SOLUTION
def sort_string_list(l):
    return sorted(l)
### END SOLUTION
```

<!-- grade_id: cell-84ee2ec237b2a4c4, points: 1, grade: true, solution: false -->
```python
x = ["a", "d", "c", "b"]
copy_of_x = x.copy()
expected_x = ["a", "b", "c", "d"]
assert sort_string_list(x) == expected_x
assert x == copy_of_x

x = []
copy_of_x = x.copy()
expected_x = []
assert sort_string_list(x) == expected_x
assert x == copy_of_x

x = ["random", "words", "to", "sort!"]
copy_of_x = x.copy()
expected_x = ["random", "sort!", "to", "words"]
assert sort_string_list(x) == expected_x
assert x == copy_of_x


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**1.4. Sort this**

Now define a function called **sort_this** that:

- receives a list as argument, let's call it list1
- calls function **select_strings** in order to get a new version of the list with only strings, let's call it list2
- applies function **clean_string** to each element of the list2
- calls function **sort_string_list** to get a sorted version of list2, let's call it list3
- returns list3

In this exercise, we have defined functions **select_strings**, **clean_string** and **sort_string_list** for you, so you don't need to reimplement them. And you can still pass this step even if you didn't complete any of the previous steps.


<!-- grade_id: cell-f089360ee63f8d17, points: 0, grade: false, solution: true -->
```python
# create function called sort_this

# def sort_this(l):
# ...
# ...

### BEGIN SOLUTION
def sort_this(l):
    return sort_string_list([clean_string(i) for i in select_strings(l)])
### END SOLUTION
```

<!-- grade_id: cell-84ee2ec237b2a4c1, points: 1, grade: true, solution: false -->
```python
x = [" Aleluia", 1, "None!   ", None]
expected_x = ["aleluia", "none!"]
assert sort_this(x) == expected_x

x = []
expected_x = []
assert sort_this(x) == expected_x

x = [1, "words", "    to", (1, 2)]
expected_x = ["to", "words"]
assert sort_this(x) == expected_x

x = [1, 2, 3]
expected_x = []
assert sort_this(x) == expected_x


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

---

## Exercise 3

**Process File**
In this exercise you should define a function called **process_file**, that processes a file.

The function should be defined so that:

- it receives two arguments: 1) a file descriptor, for the file to be processed,  2) a string that defines a separator
- you should define a variable named total, initially with the value 0
- you should iterate over the lines of the file:

1. each line is a string with multiple values. In order to split the string into its values, you should consider the separator
2. if the first value of a line is neither "+" nor "-", ignore the line
3. then, all the other values in the line are valid integers, and you'll add them to total as positive numbers if the line started with a "+", or as negative numbers, if the line started with a "-"

Note: a file descriptor is what's returned by the open( ) function when you call it on a file. This means that you don't need to call open( ) anymore, the file is already open! In fact, if you try to open the file descriptor, you should see something like this error:

`TypeError: expected str, bytes or os.PathLike object, not io.StringIO`

<!-- grade_id: cell-dba1c3448d02f3d4, points: 0, grade: false, solution: true -->
```python
# create function called multiply_matrices

#def multiply_matrices(matrix_1, matrix_2):
# ...
# ...

### BEGIN SOLUTION
def process_file(fp, separator):
    total = 0
    for line in fp:
        line_split = line.split(separator)
        line_sum = sum([int(i) for i in line_split[1:]])
        if line_split[0] == '+':
            total += line_sum
        if line_split[0] == '-':
            total -= line_sum

    return total
### END SOLUTION
```

<!-- grade_id: cell-4887ed8b8490950d, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
def mock_fp(test_input, separator):
    lines_list = [
        separator.join(str(item) for item in raw_line)
        for raw_line in test_input
    ]
    lines = '\n'.join(lines_list)

    return io.StringIO(lines)


def get_error_msg(test_input, separator, expected_output, computed):
    return (
        f"data used to generate file: {test_input}\n"
        f'with separator: "{separator}"\n'
        'ERROR!'
        f"expected: {expected_output}\n"
        f"got: {computed}"
    )


def test_process_file():
    test_cases = {
        (("+", 3, 4, 10), ("+", 1)): 18,
        (("+", 3, 4, 5), ("-", 2, 6)): 4,
        (("+", 1, 1), ("-", 5)): -3,
        (("-", 1, 1, 3),): -5,
        (("-", 11, 4), ('+', 15)): 0,
        (("-", 11, 4), ('+', 15), ('=', 10)): 0,
    }
    separators = [' ', ',', '::']

    for test_input, expected_output in test_cases.items():
        for sep in separators:
            fp = mock_fp(test_input, sep)
            answer = process_file(fp, sep)

            error_msg = get_error_msg(test_input, sep, expected_output, answer)
### END HIDDEN TEST
            assert answer == expected_output, error_msg


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

---

## Exercise 4

### Shapes

**4.1. Perimeter**

In this exercise we are going to use Python classes to represent geometrical shapes.

The first step is a building block for the next exercise.

You should define a method called **perimeter**, that:

- receives as argument a list of points, which are represented as 2D tuples (2D means that they have length = 2), and that represent the vertices of a polygon
- you don't need to do any validations on the input, just assume that all the tuples in the list have 2 elements and that they are sorted in a way that our algorithm to compute the perimeter works
- if the list of points has less than 3 points, just return 0
- otherwise, compute the perimeter of the polygon:
 
1. iterate through the list of points
2. compute the euclidean distance between each point and the next one
3. also compute the distance between the last point and the first one
4. sum all the distances and return this value

Hint: you can use numpy on this exercise.




**Grader tips**

In this exercise, you should just define function **perimeter**, you should not invoke the function!

Make sure that:

- the name of the function is exactly what is requested
- you don't have any print statements in your code 
- you are handling all the cases described in the problem statement

<!-- grade_id: cell-2a151430c72dc66f, points: 0, grade: false, solution: true -->
```python
# create function called perimeter

# def perimeter(points)
# ...
# ...

### BEGIN SOLUTION

def perimeter(points):
    if len(points) < 3:
        return 0

    dists = []

    points_round = points + [points[0]]
    for i in range(len(points)):
        p1 = np.array(points_round[i])
        p2 = np.array(points_round[i+1])
        dists.append(np.linalg.norm(p1 - p2))

    return sum(dists)

### END SOLUTION
```

<!-- grade_id: cell-9cee4c6578deb3a4, points: 2, grade: true, solution: false -->
```python
points = [(0, 0), (1, 0), (1, 1), (0, 1)]
assert perimeter(points) == 4

points = [(0, 0), (3, 0), (1.5, 7)]
assert math.isclose(perimeter(points), 17.31782, rel_tol=1e-2)

points = [(-1, -1), (3, -4), (4, 5), (4.5, 7), (-1.5, 10)]
assert math.isclose(perimeter(points), 33.83649, rel_tol=1e-2)

points = []
assert perimeter(points) == 0

points = [(1, 1)]
assert perimeter(points) == 0

points = [(1, 1), (2, 2)]
assert perimeter(points) == 0


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**4.3. Polygon**

In this exercise we are going to use Python classes to represent geometrical shapes.

Create a class **Polygon**, that:

- has an attribute called points that should be initialized through an argument passed into the constructor. points is the list of 2 dimensional points that define the polygon.
- defines a method called **area**, that simply raises a NotImplementedError, with message: "Method area should be defined in the child class."
- defines a method called **perimeter**
Method **perimeter** should behave exactly like the one you defined on the previous step. We have defined function **compute_perimeter**, that receives a list of points as argument, and that you can call in your **perimeter** method. This way, you can still complete this exercise if you didn't figure out how to implement the function from last step.

**Grader tips**

In this exercise, you should just implement class Polygon!

Make sure that:

- the name of the class and the methods are exactly what is requested
- you don't have any print statements in your code
- you are throwing the right type of error, with the right error message


<!-- grade_id: cell-edc0084ca8ba2211, points: 0, grade: false, solution: true -->
```python
# create class called Polygon

# def class Polygon:
# ...
# ...

### BEGIN SOLUTION


class Polygon:
    def __init__(self, points):
        self.points = points

    def area(self):
        msg = 'Method area should be defined in the child class.'
        raise NotImplementedError(msg)

    def perimeter(self):
        return perimeter(self.points)
    
### END SOLUTION
```

<!-- grade_id: cell-65397abfd2557135, points: 1, grade: true, solution: false -->
```python

points = [(0, 0), (2, 0), (2, 2), (0, 2)]
polygon = Polygon(points)

assert isinstance(polygon, Polygon)
assert isinstance(polygon.points, list)
assert polygon.points == points

points = [(0, 0), (2, 0), (2, 2), (0, 2)]
polygon = Polygon(points)

with pytest.raises(NotImplementedError) as e:
    polygon.area()

assert str(e.value) == 'Method area should be defined in the child class.'

points = [(0, 0), (2, 0), (2, 2), (0, 2)]
polygon = Polygon(points)

assert polygon.perimeter() == 8


points = []
polygon = Polygon(points)

assert isinstance(polygon, Polygon)
assert isinstance(polygon.points, list)
assert polygon.points == points
assert polygon.perimeter() == 0


points = [(0, 0), (10, 9)]
polygon = Polygon(points)

assert isinstance(polygon, Polygon)
assert isinstance(polygon.points, list)
assert polygon.points == points
assert polygon.perimeter() == 0
    

print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**4.3. Square** 

Now create a class **Square**, that:

- extends class **Polygon**
- re-defines method **area**

Method **area** should compute the area of the square defined by the class, by doing the following:

- we are going to believe that the points actually define a square
- so, we'll just compute the euclidean distance between the first two points in points
- and return the square of its value

In this exercise we have implemented class Polygon for you, so it will be already defined if you refer to it.

We've also implemented a function called **euclidean_distance**, that receives two tuples as arguments and returns the euclidean distance between the pair of points (tuples). You can use it in this exercise.

Finally, if you don't know what extending a class means, check the Python docs here: https://docs.python.org/3/tutorial/classes.html#inheritance.


**Grader tips**

In this exercise, you should just implement class **Square**!

Make sure that:

- the name of the class and the methods are exactly what is requested
- you don't have any print statements in your code

<!-- grade_id: cell-eacfb152edb1f698, points: 0, grade: false, solution: true -->
```python
def euclidean_distance(a, b):
    array_a = np.array(a)
    array_b = np.array(b)

    return np.linalg.norm(array_a - array_b)

# create class called Square

# class Square(Polygon):
# ...
# ...

### BEGIN SOLUTION



# this function was given in the exercise

class Square(Polygon):
    def area(self):
        side = euclidean_distance(self.points[0], self.points[1])

        return side**2
    
### END SOLUTION
```

<!-- grade_id: cell-430a45dfcb1413e1, points: 1, grade: true, solution: false -->
```python
points = [(0, 0), (2, 0), (2, 2), (0, 2)]
square = Square(points)

assert isinstance(square, Square)
assert issubclass(Square, Polygon)
assert isinstance(square.points, list)
assert square.points == points

points = [(0, 0), (2, 0), (2, 2), (0, 2)]
square = Square(points)

assert square.area() == 4

points = [(0, 0), (2, 0), (2, 2), (0, 2)]
square = Square(points)

assert square.perimeter() == 8
    

print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

### 4.4. Circle

Now create a class Circle, that:

- has two attributes, centre and radius, that should be initialized through arguments passed into the constructor in this specified order. centre will be a 2 dimensional point represented as a tuple and radius will be a number (int or float)
- has a method perimeter that returns the perimeter of the circle, as 
$$perimeter = 2 π×radius$$

- has a method area that returns the area of the circle, as 
$$ area=π×radius^2$$


 

**Grader tips**

In this exercise, you should just implement class **Circle**!

Make sure that:

the name of the class and the methods are exactly what is requested
you don't have any print statements in your code

<!-- grade_id: cell-fc30f3aa92e79eb9, points: 0, grade: false, solution: true -->
```python
# create class called Circle

# class Circle:
# ...
# ...

### BEGIN SOLUTION

class Circle:
    def __init__(self, centre, radius):
        self.centre = centre
        self.radius = radius

    def perimeter(self):
        return 2 * math.pi * self.radius

    def area(self):
        return math.pi * self.radius**2

### END SOLUTION
```

<!-- grade_id: cell-e477ce094c4043f4, points: 1, grade: true, solution: false -->
```python
centre = (1, 1)
radius = 3
circle = Circle(centre, radius)

assert isinstance(circle, Circle)
assert isinstance(circle.centre, tuple)
assert circle.centre == centre
assert circle.radius == radius

centre = (1, 1)
radius = 3
circle = Circle(centre, radius)

assert math.isclose(circle.area(), 28.27433, rel_tol=1e-2)

centre = (1, 1)
radius = 3
circle = Circle(centre, radius)

assert math.isclose(circle.perimeter(), 18.84955, rel_tol=1e-2)

centre = (1, 1)
radius = 0
circle = Circle(centre, radius)

assert circle.perimeter() == 0
assert circle.area() == 0
    

print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

### 4.5. Shapes

This is a multiple choice questions. 

What you have to do here is to flag the correct sentence  (there is only one).


1. Class Polygon is an abstract class

2. Class Circle is as a child of class Polygon

3. The relationship between classes Polygon and Square is an example of inheritance

4. Class Square is the parent of class Polygon
 

<!-- grade_id: cell-f9786f4be2292225, points: 0, grade: false, solution: true -->
```python
# answer_4_5 = ...

### BEGIN SOLUTION
answer_4_5 = 3

### END SOLUTION
```

<!-- grade_id: cell-8bdc604d7ca18ac0, points: 1, grade: true, solution: false -->
```python
# Run this cell to see if your answer is correct

assert answer_4_5 == 3

print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

---

## Exercise 5

## Numpy Questions

**5.1 - Create array**

Implement a function called **create_array**, that:

1. Receives two arguments: a list of numbers and a 2D tuple (has length = 2)
2. The goal of this function is to create a numpy array using the data from the list, with the shape defined in the tuple. For that, we need to validate if it's possible or not to create an array with that data and that shape. If it's not possible to create the array with these conditions, the function should just return None
3. Otherwise, return a numpy array with the data from the list and with the shape defined in the tuple


**Grader tips**

In this exercise, you should just define function **create_array**, you should not invoke the function!

Make sure that:

- the name of the function is exactly what is requested
- you don't have any print statements in your code


<!-- grade_id: cell-c89c51ca779da3dc, points: 0, grade: false, solution: true -->
```python
# create function called create_array

# def create_array(numbers, shape):
# ...
# ...

### BEGIN SOLUTION
def create_array(numbers, shape):
    if len(numbers) != shape[0] * shape[1]:
        return None

    return np.array(numbers).reshape(shape)


# alternative solution
def create_array_aternative(numbers, shape):
    n_elements = shape[0] * shape[1]
    if len(numbers) < n_elements:
        return None

    if len(numbers) > n_elements:
        numbers = numbers[:n_elements]

    return np.array(numbers).reshape(shape)

### END SOLUTION
```

<!-- grade_id: cell-2a4834cad293d726, points: 1, grade: true, solution: false -->
```python
numbers = [1, 2, 3, 4]
shape = (2, 2)

expected_result = np.array([[1, 2], [3, 4]])
result = create_array(numbers, shape)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)


numbers = [1, 2, 3, 4]
shape = (3, 3)

assert create_array(numbers, shape) is None


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**5.2 - Substract matrices**

Implement a function called **subtract_matrices**, that:

1. Receives two arguments that are numpy arrays: *a* and *b*
2. Checks if the shapes of *a* and *b* are the same
3. If not, the function should just return None
4. Otherwise, return a numpy array where each element is the subtraction of the correspondent element in *a* and *b*

<!-- grade_id: cell-3e293051e761c4f2, points: 0, grade: false, solution: true -->
```python
# create function called subtract_matrices

# def subtract_matrices(a, b):
# ...
# ...


### BEGIN SOLUTION

def subtract_matrices(a, b):
    if a.shape != b.shape:
        return None

    return a - b


### END SOLUTION
```

<!-- grade_id: cell-8c6bff316acc6eb1, points: 1, grade: true, solution: false -->
```python
a = np.array([[1, 2], [3, 4]])
b = np.array([[1, 1], [1, 1]])

expected_result = np.array([[0, 1], [2, 3]])
result = subtract_matrices(a, b)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)


a = np.array([[1, 2], [3, 4]])
b = np.array([[1, 1]])

assert subtract_matrices(a, b) is None


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**5.3 - Sum rows**

Implement a function called **sum_rows**, that:

1. Receives an argument that is a numpy array
2. Sums the values in each row of the array
3. Returns the resulting array

<!-- grade_id: cell-861b51250ce79c44, points: 0, grade: false, solution: true -->
```python
# create function called sum_rows

# def sum_rows(a):
# ...
# ...


### BEGIN SOLUTION
def sum_rows(a):
    return a.sum(axis=1)

### END SOLUTION



```

<!-- grade_id: cell-c981c118374add77, points: 1, grade: true, solution: false -->
```python
a = np.array([[1, 2], [3, 4]])
expected_result = np.array([3, 7])
result = sum_rows(a)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)

a = np.array([[1, -2], [0, 4], [6, -10]])
expected_result = np.array([-1, 4, -4])
result = sum_rows(a)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**5.4 - Maximum per column**

Implement a function called **maximum_per_column**, that:

1. Receives an argument that is a numpy array x
2. Returns an array where each value is the maximum value in each column of x

<!-- grade_id: cell-e080838a6b305989, points: 0, grade: false, solution: true -->
```python
# create function called maximum_per_column

# def maximum_per_column(x):
# ...
# ...

### BEGIN SOLUTION
def maximum_per_column(x):
    return x.max(axis=0)
### END SOLUTION


```

<!-- grade_id: cell-c981c118374add78, points: 1, grade: true, solution: false -->
```python
a = np.array([[1, 2], [3, -5]])
expected_result = np.array([3, 2])
result = maximum_per_column(a)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)

a = np.array([[1, -2], [0, 4], [6, -10]])
expected_result = np.array([6, 4])
result = maximum_per_column(a)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

**5.5 - Replace negatives**

Implement a function called **replace_negatives**, that:

1. Receives two argument: a numpy array, *x*, and a number, replacement
2. Gets a new numpy array which is the a version of *x*, with all the negative values replaced by replacement
3. Return the new numpy array
4. This function should not change the input numpy array!

<!-- grade_id: cell-e080838a6b305988, points: 0, grade: false, solution: true -->
```python
# create function called replace_negatives

# def replace_negatives(x, replacement):
# ...
# ...

### BEGIN SOLUTION
def replace_negatives(x, replacement):
    return np.where(x > 0, x, replacement)
### END SOLUTION


```

<!-- grade_id: cell-c981c118374add79, points: 1, grade: true, solution: false -->
```python
a = np.array([[1, 2], [3, -5]])
b = 100
original_a = a.copy()

expected_result = np.array([[1, 2], [3, 100]])
result = replace_negatives(a, b)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)

error_msg = 'The original array has changed'
assert np.array_equal(a, original_a), error_msg

a = np.array([[-100, 2], [3.1, -5]])
b = 7
original_a = a.copy()

expected_result = np.array([[7, 2], [3.1, 7]])
result = replace_negatives(a, b)

assert isinstance(result, np.ndarray)
assert np.array_equal(result, expected_result)

error_msg = 'The original array has changed'
assert np.array_equal(a, original_a), error_msg


print("Answer is correct. Good Job.")

```

Output:
```
Answer is correct. Good Job.
```

---

## Exercise 6 - Git questions


This exercise is a quiz with 5 multiple answer questions.
In the cell below, we've declared a dictionary called `answers`.

You should fill in that dictionary with your answers, using as keys the question number, like `question_x`, and as values the number from `1` to `4` that corresponds to the right answer.
For each question, you should provide only one answer, i.e, the dict values should have type **int**.

For example, if you want to answer Question 1 with choice number 2, then you do:
```
answers["question_1"] = 2
```

```python
answers = {}
```

**Question 1** 

In your local machine, you have a clone of git repository. You haven't touched it in a while and now want to work a bit on this project.

What is the first thing that you should do?

1. git push

2. git add

3. git commit

4. git pull

<!-- grade_id: cell-944db9daa15f31b3, points: 0, grade: false, solution: true -->
```python
# answers["question_1"] = ...

### BEGIN SOLUTION

answers["question_1"] = 4

### END SOLUTION
```

**Question 2** 
Now that you made sure that your clone of the git repository is up to date, you will start working on it.

You are a very careful person, so you want to make sure that your code is not going to mess up the current version of the code.

What's the first step that you should take when working on this project, while feeling safe that you're not going to mess up the current version of the code?


1. Open a pull request on GitHub

2. Ask for comments on your pull request before you commit and push your changes

3. Create a new branch and work on your changes there instead of in master branch

4. Ask for comments on your pull request after you commit and push your changes

<!-- grade_id: cell-944db9daa15f31b2, points: 0, grade: false, solution: true -->
```python
# answers["question_2"] = ...

### BEGIN SOLUTION

answers["question_2"] = 3

### END SOLUTION
```

**Question 3** 

What does the -m flag does in the git commit command?


1. It means that your both committing and pushing your changes in the same command

2. It allows you to add a commit message

3. It means that you'll commit against the master branch

4. It means that your both adding the changed files and committing your changes in the same command

<!-- grade_id: cell-884dc207680723be, points: 0, grade: false, solution: true -->
```python
# answers["question_3"] = ...

### BEGIN SOLUTION

answers["question_3"] = 2

### END SOLUTION
```

**Question 4** 

What git command should you use if you wanted to add all the python files in the current directory to the staging area?

1. git add .

2. git add . py

3. git add *.py

4. git add .py

<!-- grade_id: cell-20a72a68ef2738e4, points: 0, grade: false, solution: true -->
```python
# answers["question_4"] = ...

### BEGIN SOLUTION

answers["question_4"] = 3

### END SOLUTION
```

**Question 5** 

What git command should you use to update the remote repository after you've commited your changes to the master branch?

1. git stash

2. git push

3. git pull

4. git commit

<!-- grade_id: cell-149bc74590f9b7f9, points: 0, grade: false, solution: true -->
```python
# answers["question_5"] = ...

### BEGIN SOLUTION

answers["question_5"] = 2

### END SOLUTION
```

<!-- grade_id: cell-360cc5dfa82f167f, points: 1, grade: true, solution: false -->
```python
# Run this cell to see if your answers are correct
utils.exercise_6(answers)

```

Output:
```
Answer is correct. Good Job.
```

---

### Last but not least, submit your work!

To submit your work, fill your slack ID in the `slack_id` variable (as a string).

Example: `slack_id = "x-men"`

Help: if you forgot your slack ID, [read this](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c).

<!-- grade_id: cell-cc9a39c41ca8e434, points: 0, grade: false, solution: true -->
```python
# Submit your work!

#slack_id = 

### BEGIN SOLUTION
slack_id = "ADMINSLU18_2"
### END SOLUTION
```

<!-- grade_id: cell-eaa3733afdf6743c, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 18_2
submit(slack_id, slu)
```

Output:
```
Success
```

---