# SLU04 - Data Structures

### Start by importing the following packages

```python
#used for evaluation
import hashlib
import json
import random
```

In this notebook the following is tested:
- Tuples   
- Lists   
- Dictionaries
- Sets

**IMPORTANT:** Some exercises require you to use a variable that was defined in a previous cell, so don't forget to run it before!

For example, in Exercise 1.2 you need to run the cell with `color = ("red", "blue", "green", "yellow", "black", "white")` before running the following cell with your solution.

## Exercise 1: Tuples  <a name="1"></a>

This exercise covers topics learned regarding tuples.

### 1.1) Create a tuple 

Create a tuple of __floats__ named `this_tuple` with size 5.

<!-- grade_id: cell-3a4c0b59df7bcf57, points: 0, grade: false, solution: true -->
```python
#this_tuple = ...
### BEGIN SOLUTION
this_tuple = (1., 2., 3., 4., 5.)
print(type(this_tuple))
print(len(this_tuple))
print(type(this_tuple[3]))
### END SOLUTION
```

Output:
```
<class 'tuple'>
5
<class 'float'>
```

<!-- grade_id: cell-b63230a58719fa25, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_tuple, tuple), "Are you sure this_tuple is a tuple?"
assert len(this_tuple) == 5, "The length is not quite right."
for i in this_tuple:
    assert isinstance(i, float), "Did you write floats?"

```

### 1.2) Index a tuple

Consider the following tuple:    

```python
color = ("red", "blue", "green", "yellow", "black", "white")
```

Using __negative__ indexing, what is the index of the element `"green"`? Assign this value to the variable `green_index`.

<!-- grade_id: cell-268b39f922ddc862, points: 0, grade: false, solution: true -->
```python
# Hint: the solution should be a number
# green_index = ...
### BEGIN SOLUTION
green_index = -4
print(color[green_index])
print(
    'green_index hashed:',
    hashlib.sha256(
        json.dumps(green_index).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
green
green_index hashed: e5e0093f285a4fb94c3fcc2ad7fd04edd10d429ccda87a9aa5e4718efadf182e
```

<!-- grade_id: cell-d3eb2f19b4f68632, points: 0.5, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(green_index).encode()).hexdigest() == 'e5e0093f285a4fb94c3fcc2ad7fd04edd10d429ccda87a9aa5e4718efadf182e', "The index is not correct. Are you using negative indexing?"
```

### 1.3) Slice a tuple

Extract `("green", "blue", "red")` (in this same order) from the tuple `color`.
Assign the results to a variable called `rgb`.

```python
color = ("red", "blue", "green", "yellow", "black", "white")
```

<!-- grade_id: cell-205bd03d74d9e9f5, points: 0, grade: false, solution: true -->
```python
# Hint: use backwards slicing
# rgb = ...

### BEGIN SOLUTION
rgb = color[-4:-7:-1]
print(rgb)
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(rgb).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
('green', 'blue', 'red')
rgb hashed: 43da8949efc00c51f7d96130a25a6b902f6cfd6157817015ca6ee524e3085374
```

<!-- grade_id: cell-f17f899e57df23ad, points: 1, grade: true, solution: false -->
```python
assert isinstance(rgb, tuple), "Is your result a tuple?"
assert len(rgb) == 3, "You aren't selecting the correct number of elements."
assert rgb[0] == "green", "Check which elements you are selecting and their order."
assert hashlib.sha256(json.dumps(rgb).encode()).hexdigest() == '43da8949efc00c51f7d96130a25a6b902f6cfd6157817015ca6ee524e3085374'
```

###  1.4) Index a tuple of tuples

Considering the following tuple of tuples:

```python
random_numbers = ((1, 2, 3),(4, 5, 6),(7, 8, 9),(10, 11, 12))
```

What is the right way to extract the number 8 from `random_numbers`?

a) `random_numbers[3][2]`   
b) `random_numbers[-1][-1]`   
c) `random_numbers[-2][1]`  
d) `random_numbers[-1][1]`   

<!-- grade_id: cell-7e6632972efca3db, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
### BEGIN SOLUTION
random_numbers = ( (1, 2, 3), (4, 5, 6), (7, 8, 9), (10, 11, 12) )
answer = "c"
print(answer)
print(random_numbers[-2][1])
print(
    'answer hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
c
8
answer hashed: 879923da020d1533f4d8e921ea7bac61e8ba41d3c89d17a4d14e3a89c6780d5d
```

<!-- grade_id: cell-a17e088c506ff934, points: 0.5, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '879923da020d1533f4d8e921ea7bac61e8ba41d3c89d17a4d14e3a89c6780d5d', "Wrong answer."
```

### 1.5) Tuple of size one

How can we create the tuple `example` of size 1 __without__ using the function __tuple__?

a) example = (5)   
b) example = 5   
c) example = [5]   
d) example = 5,   

<!-- grade_id: cell-d239485262f1edcd, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
### BEGIN SOLUTION
answer = "d"
print(answer)
this_tuple = 5,
print(type(this_tuple))
print(
    'answer hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
d
<class 'tuple'>
answer hashed: 3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633
```

<!-- grade_id: cell-3e575a27a20113ad, points: 0.5, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633', "Wrong answer."
```

###  1.6) Replace values in a tuple

Can a tuple be modified after its creation?

<!-- grade_id: cell-1d7ed0f610a63b1a, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "yes"
#answer = "no"
### BEGIN SOLUTION
answer = "no"
print(answer)
print(
    'answer hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
no
answer hashed: 04a06452677210a3cdaec376fd5ebbca1714cb7af9e62bf5cce1644310a9086a
```

<!-- grade_id: cell-c7188110ce71150e, points: 0.5, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '04a06452677210a3cdaec376fd5ebbca1714cb7af9e62bf5cce1644310a9086a', "Wrong answer."
```

### 1.7) Merge two tuples

Considering the following tuples:

```python
left = (1,11,22215,7,14,1,11,9,1,6,2,5)
right = (1,24,50,45,2,45,1,1,2,1,2,1,88,9,9,9,44,5,2)
```

Create a new tuple  `this_tuple` by merging the tuples above. The elements in the tuple `left` should precede the ones on the tuple `right`.

<!-- grade_id: cell-05376cff3e8e7950, points: 0, grade: false, solution: true -->
```python
# Hint: Using operations between tuples.
#this_tuple = ...

### BEGIN SOLUTION
this_tuple = left + right
print(this_tuple)
print(type(this_tuple))
print(len(this_tuple))
print(this_tuple[-3])
print(
    'this_tuple hashed:',
    hashlib.sha256(
        json.dumps(this_tuple).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
(1, 11, 22215, 7, 14, 1, 11, 9, 1, 6, 2, 5, 1, 24, 50, 45, 2, 45, 1, 1, 2, 1, 2, 1, 88, 9, 9, 9, 44, 5, 2)
<class 'tuple'>
31
44
this_tuple hashed: 0c1afd35431992aba438b9382b08332f4d4fed7d4380e538c047e42a223a9dd5
```

<!-- grade_id: cell-74db790f2e9e6836, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_tuple, tuple), "The result should be a tuple."
assert len(this_tuple) == 31, "The merging is not right."
assert this_tuple[-3] == 44, "Re-check the order of the tuples."
assert hashlib.sha256(json.dumps(this_tuple).encode()).hexdigest() == '0c1afd35431992aba438b9382b08332f4d4fed7d4380e538c047e42a223a9dd5'
```

## Exercise 2: Lists  <a name="2"></a> 

This exercise covers topics learned regarding lists.

###  2.1) List Creation

Write a list named `this_list` with length five and with the string `"bananas"` on negative index -4.

<!-- grade_id: cell-f8854927f9656f23, points: 0, grade: false, solution: true -->
```python
#bananas = ...
#this_list = ...

### BEGIN SOLUTION
bananas = "bananas"
this_list = ["kiwi", bananas, "melon", "lemon", "kiwi"]
print(this_list)
print(type(this_list))
print(len(this_list))
print(this_list[-4])
### END SOLUTION
```

Output:
```
['kiwi', 'bananas', 'melon', 'lemon', 'kiwi']
<class 'list'>
5
bananas
```

<!-- grade_id: cell-5de18c0078e98c2e, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_list, list), "The result should be a list."
assert len(this_list) == 5, "The length is not quite right."
assert this_list[-4] == "bananas", "There is no \"bananas\" on position -4."
```

###  2.2) Delete and append values in a list

Given the list `ice_cream`, delete `"chocolate"` and append `"dulce de leche"`.

```python
ice_cream = ["lemon", "stracciatella", "pistacchio", "chocolate", "vanilla"]
```

<!-- grade_id: cell-c22be3e4356c6e22, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
del ice_cream[-2]
ice_cream.append("dulce de leche")
print("chocolate" in ice_cream)
print(ice_cream[-1])

print(
    'ice_cream hashed:',
    hashlib.sha256(
        json.dumps(ice_cream).encode()
    ).hexdigest()
)

### END SOLUTION
```

Output:
```
False
dulce de leche
ice_cream hashed: 083936d7994a5de336fe655aa14bb4b14ad95e30605fa4252937ddc8776cfa09
```

<!-- grade_id: cell-298523708d65510d, points: 1, grade: true, solution: false -->
```python
assert "chocolate" not in ice_cream, "There is still \"chocolate\" in the ice_cream."
assert ice_cream[-1] == "dulce de leche", "Did you append the right ingredient?"
assert hashlib.sha256(json.dumps(ice_cream).encode()).hexdigest() == '083936d7994a5de336fe655aa14bb4b14ad95e30605fa4252937ddc8776cfa09'
```

###  2.3) Delete the last value in a list

Considering the following list, `ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]`, what is the right answer if we want to delete `"vanilla"` in the last position?

a) `del ice_cream[1]`   
b) `del ice_cream[-1]`    
c) `ice_cream.remove("vanilla")`   
d) `ice_cream[-1] = None`   

<!-- grade_id: cell-fafe1f47881998b6, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"

### BEGIN SOLUTION
answer = "b"

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
del ice_cream[1]
print(ice_cream)

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
del ice_cream[-1] 
print(ice_cream)

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
ice_cream.remove("vanilla")
print(ice_cream)

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
ice_cream[-1] = None
print(ice_cream)


print(
    'answer hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
['lemon', 'stracciatella', 'pistacchio', 'chocolate', 'vanilla']
['lemon', 'vanilla', 'stracciatella', 'pistacchio', 'chocolate']
['lemon', 'stracciatella', 'pistacchio', 'chocolate', 'vanilla']
['lemon', 'vanilla', 'stracciatella', 'pistacchio', 'chocolate', None]
answer hashed: c100f95c1913f9c72fc1f4ef0847e1e723ffe0bde0b36e5f36c13f81fe8c26ed
```

<!-- grade_id: cell-6589184220a059bc, points: 0.5, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == 'c100f95c1913f9c72fc1f4ef0847e1e723ffe0bde0b36e5f36c13f81fe8c26ed', "Wrong answer."
```

###  2.4) List Operations

Using list operations, create a list called `hello_world` of size 10000. It should have just two unique values, `"hello"` and `"world"`. Starting with `"hello"` then `"world"`, then `"hello"` and so on.

<!-- grade_id: cell-0989f9a6a1c1f63e, points: 0, grade: false, solution: true -->
```python
# Hint: multiply a list by a int 
#hello_world = ...

### BEGIN SOLUTION
hello_world = ["hello", "world"] * 5000

print(len(hello_world))
print(hello_world[0])
print(hello_world[3])
print(hello_world[4532])
print(
    'hello_world hashed:',
    hashlib.sha256(
        json.dumps(hello_world).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
10000
hello
world
hello
hello_world hashed: ba01ef9ff3158651119ad5f13ef61b559dd73eff752f4a2cc05e3912c2c85d32
```

<!-- grade_id: cell-4ac5d3b434c8a490, points: 1, grade: true, solution: false -->
```python
assert len(hello_world) == 10000, "The size is not right."
assert hello_world[0] == "hello", "I guess you started on the wrong foot."

for idx,word in enumerate(hello_world):
    if (idx % 2) == 0:
        assert word == "hello", "Are the words alternating?"
    else:
        assert word == "world", "Are the words alternating?"
        
assert hashlib.sha256(json.dumps(hello_world).encode()).hexdigest() == 'ba01ef9ff3158651119ad5f13ef61b559dd73eff752f4a2cc05e3912c2c85d32', "Wrong answer."
```

###  2.5) Replace values and sort a list

Replace `"pistacchio"` in the list `ice_cream` by `"cream"`.

Add the sublist `others` to the end of the list `ice_cream`.   

After these operations, sort the elements in the list and convert it to a __tuple__.

```python
ice_cream = ["lemon", "stracciatella", "pistacchio", "chocolate", "vanilla"]

others = ["dulce de leche", "caramel", "cookies", "peanut butter"] * 100
```

<!-- grade_id: cell-7327a77943c8fe67, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ice_cream[2] = "cream"
ice_cream = ice_cream + others
ice_cream.sort()
ice_cream = tuple(ice_cream)

print(len(ice_cream))

print(
    'ice_cream hashed:',
    hashlib.sha256(
        json.dumps(ice_cream).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
405
ice_cream hashed: bbe23485294172ad0fab1f37eee62d4f518a296aad06ea944d6d7a31e6ef58a6
```

<!-- grade_id: cell-955f097b2db82b73, points: 1, grade: true, solution: false -->
```python
assert len(ice_cream) == 405, "The size of the tuple is not right."
assert ice_cream.index("cream") == 201, "Did you replace \"pistacchio\"?"
assert isinstance(ice_cream, tuple), "The result is not a tuple."
assert hashlib.sha256(json.dumps(ice_cream).encode()).hexdigest() == 'bbe23485294172ad0fab1f37eee62d4f518a296aad06ea944d6d7a31e6ef58a6'
```

## Exercise 3: Dictionaries  <a name="3"></a>

This exercise covers topics learned regarding dictionaries.

###  3.1) Create a dictionary

Create a dictionary called `this_dict` with 5 key-value pairs where the keys are strings and the values are lists.

<!-- grade_id: cell-bb5b978acdaaae16, points: 0, grade: false, solution: true -->
```python
#this_dict = ...
### BEGIN SOLUTION
this_dict = {"a":[],"b":[],"c":[],"d":[],"e":[]}
print(len(this_dict))
print(isinstance(list(this_dict.keys())[0], str))
print(isinstance(list(this_dict.values())[0], list))
print(isinstance(list(this_dict.keys())[2], str))
### END SOLUTION
```

Output:
```
5
True
True
True
```

<!-- grade_id: cell-0f3a2761c0922b5f, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_dict, dict), "The result is not a dictionary."
assert len(this_dict) == 5, "The dictionary doesn't have 5 key-value pairs."
for key, value in this_dict.items():
    assert isinstance(key, str), "The dictionary keys are not strings: %s" % key
    assert isinstance(value, list), "The dictionary values are not lists: %s" % value
```

### 3.2) Extract a value from a dictionary

Considering the following dictionary named `groceries`:

```python
groceries = {             
             "bread": {"type": "grains", "price_per_unit": 2, "quantity_purchased": 1},   
             "onions": {"type": "vegetables", "price_per_unit": 0.5, "quantity_purchased": 2},   
             "spinaches": {"type": "vegetables" , "price_per_unit": 1.5, "quantity_purchased": 1}   
            }
```

What is the notation that we should use in order to extract "grains"?

a) `groceries["type"]`   
b) `groceries["bread"]`   
c) `groceries[0][0]`   
d) `groceries["bread"]["type"]`

<!-- grade_id: cell-fd6a36d5f663708b, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
### BEGIN SOLUTION
answer = "d"
print(answer)
print(groceries["bread"]["type"])

print(
    'answer hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
d
grains
answer hashed: 3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633
```

<!-- grade_id: cell-978adad567f2a3cb, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633', "Wrong answer."
```

###  3.3) Replace, Append and Delete operations on dictionaries

Considering the following dictionary named groceries:

```python
groceries = {   
             "bread": {"type": "grains", "price_per_unit": 2, "quantity_purchased": 1},   
             "onions": {"type": "vegetables", "price_per_unit": 0.5, "quantity_purchased": 2},   
             "spinages": {"type": "vegetables" , "price_per_unit": 1.5, "quantity_purchased": 1}   
            }
```

Complete the following questions regarding Replacing, Appending and Deleting key-value pairs to/from a dictionary.

####  3.3.1) Replace a value on a key-value pair

Considering the dictionary `groceries`, update the `price_per_unit` for bread from 2 to 3.   

<!-- grade_id: cell-095cb4bce1528b15, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
groceries["bread"]["price_per_unit"] = 3

print(groceries["bread"]["price_per_unit"])

print(
    'groceries hashed:',
    hashlib.sha256(
        json.dumps(groceries).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
3
groceries hashed: be21838b4e9ebe7f267a1a172ec093414999f321b6c027fa0a8a02f9f4f5174d
```

<!-- grade_id: cell-c5c653d824ffcda9, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(groceries).encode()).hexdigest() == 'be21838b4e9ebe7f267a1a172ec093414999f321b6c027fa0a8a02f9f4f5174d', "Wrong answer."
```

####  3.3.2) Add a key-value pair

Let's add `"rice"` to our `groceries` dictionary.

The `"type"` should be `"grains"`, its `"price_per_unit"` is 1 and the `"quantity_purchased"` is 2.

<!-- grade_id: cell-bacb5fca35772b35, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
groceries["rice"] = {"type": "grains", "price_per_unit": 1, "quantity_purchased": 2}
print(groceries["rice"])

print(
    'groceries_rice hashed:',
    hashlib.sha256(
        json.dumps(groceries["rice"]).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
{'type': 'grains', 'price_per_unit': 1, 'quantity_purchased': 2}
groceries_rice hashed: 8cfb3220f7b8b23cc8c9e578d6a7a5680dfa301c90c90d287764517e42850a82
```

<!-- grade_id: cell-59704708ceb3f98c, points: 1, grade: true, solution: false -->
```python
assert "rice" in groceries.keys(), "Are you sure you've added rice to the dictionary?"
assert hashlib.sha256(json.dumps(groceries["rice"]).encode()).hexdigest() == '8cfb3220f7b8b23cc8c9e578d6a7a5680dfa301c90c90d287764517e42850a82', "Wrong answer."
```

####  3.3.3) Delete a key-value pair

Delete onions from our `groceries` dictionary.

<!-- grade_id: cell-c8c6802714175503, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
del groceries["onions"]
"onions" not in groceries.keys()
### END SOLUTION
```

Output:
```
True
```

<!-- grade_id: cell-9b202febf2268e6d, points: 0.5, grade: true, solution: false -->
```python
assert "onions" not in groceries.keys(), "There's still \"onions\" on the groceries."
```

###  3.4) Extract keys and values using methods

Considering the dictionary `seniority`, extract the keys and values of the dictionary to the variables `names` and `ages`, respectively.

Convert the variable `ages` to a list called `list_ages` and sort this list.

In the end, calculate how many people are senior and assign the value to `n_seniors`.

```python
seniority = {"joao":"senior", "bernardo":"adult", "gabriel":"child", "antonio":"senior", "maria":"senior", "joel":"adult", "ines":"adult", "alberto":"adult", "amilcar":"senior", "emilia":"adult", "ana":"adult", "margarida":"adult" }
```

<!-- grade_id: cell-1f5f20a99e9e053b, points: 0, grade: false, solution: true -->
```python
#names = ...
#ages = ...
#list_ages = ..
#n_seniors = ...

### BEGIN SOLUTION
names = seniority.keys()
print(names)
print(len(names))

ages = seniority.values()
print(ages)
print(len(ages))

list_ages = list(ages)
list_ages.sort()
n_seniors = list_ages.count("senior")

print(
    'list(names) hashed:',
    hashlib.sha256(
        json.dumps(list(names)).encode()
    ).hexdigest()
)
print(
    'list_ages hashed:',
    hashlib.sha256(
        json.dumps(list_ages).encode()
    ).hexdigest()
)
print(
    'n_seniors hashed:',
    hashlib.sha256(
        json.dumps(n_seniors).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
dict_keys(['joao', 'bernardo', 'gabriel', 'antonio', 'maria', 'joel', 'ines', 'alberto', 'amilcar', 'emilia', 'ana', 'margarida'])
12
dict_values(['senior', 'adult', 'child', 'senior', 'senior', 'adult', 'adult', 'adult', 'senior', 'adult', 'adult', 'adult'])
12
list(names) hashed: 9a2a381eef2e8212d8c8b48ab6531a62c064e9fbc39091015c8131138174eaf1
list_ages hashed: 1fe2f77f82fc12b06e339263d31d6d6d261994a598e6c870152429dfc00a1ca5
n_seniors hashed: 4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a
```

<!-- grade_id: cell-2977c2eb6bba3963, points: 1, grade: true, solution: false -->
```python
assert len(names) == 12, "The number of names is not correct."
assert hashlib.sha256(json.dumps(list(names)).encode()).hexdigest() == '9a2a381eef2e8212d8c8b48ab6531a62c064e9fbc39091015c8131138174eaf1', "The variable names is incorrect."
assert len(ages) == 12, "The number of ages is incorrect."
assert hashlib.sha256(json.dumps(list_ages).encode()).hexdigest() == '1fe2f77f82fc12b06e339263d31d6d6d261994a598e6c870152429dfc00a1ca5', "The variable list_ages is not correct. Did you sort it?"
assert isinstance(list_ages, list), "list_ages should be a list."
assert hashlib.sha256(json.dumps(n_seniors).encode()).hexdigest() == '4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a', "The number of seniors is incorrect."
```

## Exercise 4: Sets  <a name="4"></a>

This exercise covers topics learned regarding sets.

### 4.1) Create a set 

Create an empty set called `mary_collection`.

<!-- grade_id: cell-939dc6aefe259e77, points: 0, grade: false, solution: true -->
```python
#mary_collection = ...
### BEGIN SOLUTION
mary_collection = set()
print(type(mary_collection))
print(len(mary_collection))
### END SOLUTION
```

Output:
```
<class 'set'>
0
```

<!-- grade_id: cell-66f3420a0f38d7f9, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(mary_collection, set), "Are you sure mary_collection is a set?"
assert len(mary_collection) == 0, "The length is not quite right."
```

### 4.2) Add values to a set 

####  4.2.1) Add values from a list to a set

Add the values in the list `mary_albums` to the set `mary_collection` created in the previous step.

```python
mary_albums = ["Repeater", "The Queen Is Dead", "Atom Heart Mother", "Loveless", "Milo Goes To College", "Revolver"]
```

<!-- grade_id: cell-56dd2081f31559cc, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mary_collection.update(mary_albums)

print(
    'len(mary_collection) hashed:',
    hashlib.sha256(
        json.dumps(len(mary_collection)).encode()
    ).hexdigest()
)

print(
    'list(mary_collection) hashed:',
    hashlib.sha256(
        json.dumps(sorted(list(mary_collection))).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
len(mary_collection) hashed: e7f6c011776e8db7cd330b54174fd76f7d0216b612387a5ffcfb81e6f0919683
list(mary_collection) hashed: 3e697cd768486645a889b03f2041357b2cdaf517d2c2ea222410506a0286cedc
```

<!-- grade_id: cell-9d163b623ea9cb27, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(len(mary_collection)).encode()).hexdigest() == 'e7f6c011776e8db7cd330b54174fd76f7d0216b612387a5ffcfb81e6f0919683', "The number of elements in the set is not correct."
assert hashlib.sha256(json.dumps(sorted(list(mary_collection))).encode()).hexdigest() == '3e697cd768486645a889b03f2041357b2cdaf517d2c2ea222410506a0286cedc', "The result is not correct..."
```

####  4.2.2) Add a single value to a set

Add the album `"Antics"` to the set `mary_collection`.

<!-- grade_id: cell-2f1785dbc24b6d2c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mary_collection.add("Antics")
print(
    'len(mary_collection) hashed:',
    hashlib.sha256(
        json.dumps(len(mary_collection)).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
len(mary_collection) hashed: 7902699be42c8a8e46fbbb4501726517e86b22c56a189f7625a6da49081b2451
```

<!-- grade_id: cell-af70c7273fff9dca, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(len(mary_collection)).encode()).hexdigest() == '7902699be42c8a8e46fbbb4501726517e86b22c56a189f7625a6da49081b2451', "The number of elements in the set is not correct. Did you add the new album?"
assert "Antics" in mary_collection, "The album is not there... Did you make a typo?"
```

### 4.3) Operations with sets

####  4.3.1) Find common values

Now, please find out which albums Mary and Fred have in common, using the sets `mary_collection` and `fred_collection`.

Assign the result to the variable `same_albums`.

```python
fred_collection = {"Still Life", "Atom Heart Mother", "Blackwater Park", "Lateralus", "Revolver", "Foxtrot"}
```

<!-- grade_id: cell-0865490a991ecbcb, points: 0, grade: false, solution: true -->
```python
#same_albums = 
### BEGIN SOLUTION
same_albums = mary_collection.intersection(fred_collection)

print(same_albums)

print(
    'same_albums hashed:',
    hashlib.sha256(
        json.dumps(sorted(list(same_albums))).encode()
    ).hexdigest()
)

### END SOLUTION
```

Output:
```
{'Revolver', 'Atom Heart Mother'}
same_albums hashed: 244aa15017be9070eafbf75ded26c1cdb11c103a886c0a77f398554a7849c4ed
```

<!-- grade_id: cell-caaff753f2640a4f, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(sorted(list(same_albums))).encode()).hexdigest() == '244aa15017be9070eafbf75ded26c1cdb11c103a886c0a77f398554a7849c4ed', "The final set is not correct..."
```

####  4.3.2) Get all values from 2 sets

Now, let's create the set `full_collection` that contains all albums from `mary_collection` and `fred_collection`.
What operation should we use?

a) `mary_collection & fred_collection`  
b) `fred_collection - mary_collection`    
c) `mary_collection.union(fred_collection)`   
d) `mary_collection.difference(fred_collection)`

<!-- grade_id: cell-cd331dcd713eaf11, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
### BEGIN SOLUTION
answer = "c"
full_collection = mary_collection.union(fred_collection)

print(full_collection)

print(
    'answer hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
{'Revolver', 'Antics', 'Repeater', 'Still Life', 'Blackwater Park', 'The Queen Is Dead', 'Atom Heart Mother', 'Foxtrot', 'Loveless', 'Milo Goes To College', 'Lateralus'}
answer hashed: 879923da020d1533f4d8e921ea7bac61e8ba41d3c89d17a4d14e3a89c6780d5d
```

<!-- grade_id: cell-4f884d0d1a257efd, points: 0.5, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '879923da020d1533f4d8e921ea7bac61e8ba41d3c89d17a4d14e3a89c6780d5d', "Oops...that's not the right answer"
```

### 4.4) Use set to check duplicates in a list

Given the list `lucky_numbers`, create the variable `total_size` corresponding to the total number of values in the list and create the variable `unique_size` that corresponds to the number of **unique** values there are in the list (i.e. excluding duplicates). 

```python
lucky_numbers = [1,6,3,8,4237,243,2,73,3,753,531,8,2,982,12,5,6,73,531,0,642,568,0,132]
```

<!-- grade_id: cell-fa5c1cbe99743894, points: 0, grade: false, solution: true -->
```python
#total_size = ...
#unique_size = ...
### BEGIN SOLUTION
total_size = len(lucky_numbers)
unique_size = len(set(lucky_numbers))

print(total_size)
print(unique_size)

print(
    'total_size hashed:',
    hashlib.sha256(
        json.dumps(total_size).encode()
    ).hexdigest()
)

print(
    'unique_size hashed:',
    hashlib.sha256(
        json.dumps(unique_size).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
24
17
total_size hashed: c2356069e9d1e79ca924378153cfbbfb4d4416b1f99d41a2940bfdb66c5319db
unique_size hashed: 4523540f1504cd17100c4835e85b7eefd49911580f8efff0599a8f283be6b9e3
```

<!-- grade_id: cell-ca3480f44778254d, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(total_size).encode()).hexdigest() == 'c2356069e9d1e79ca924378153cfbbfb4d4416b1f99d41a2940bfdb66c5319db', 'The total size of the list is not correct'
assert hashlib.sha256(json.dumps(unique_size).encode()).hexdigest() == '4523540f1504cd17100c4835e85b7eefd49911580f8efff0599a8f283be6b9e3', 'The number of unique values in the list in not correct...'
```

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-fba4a36826306298, points: 0, grade: false, solution: true -->
```python
#slack_id = ...
### BEGIN SOLUTION
slack_id = None
### END SOLUTION
```

<!-- grade_id: cell-9c18d9bc564b7e70, points: 0, grade: true, solution: false -->
```python
from submit import submit
submit(slack_id, 4)
```