# SLU03 - Git Basics - Exercise notebook

```python
# Import for evaluating the results.
import hashlib
from submit import submit
```

I hope you enjoyed the **Learning Notebook**! The Exercise Notebook is multiple choice so it's really important that you have a good grasp of everything before starting it. **Good luck!**

## Exercise 1: Version Control
What is the main principle behind a VCS:

- Option **A**: Changes are tracked in a given location, allowing you to track, revert and compare all changes to the files and folders within it

- Option **B**: All files and tracked changes are shared publicly in a remote repository

- Option **C**: Changes are tracked in all files of a repository

- Option **D**: You cannot start a repository without a README


<!-- grade_id: cell-0ad663f9cf9def82, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_1 the option letter that you think is correct.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_1 =
```

<!-- grade_id: cell-28da65d304eb76f9, points: 2, grade: true, solution: false -->
```python
exercise_1_hash = "559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_1,str), "The variable should be a string."
assert exercise_1.isupper(), "Please use uppercase for your answer."
assert exercise_1.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_1_hash == hashlib.sha256(bytes(exercise_1, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 2: Git

Which one of the following statements is false: 

- Option **A**: Git is a decentralized version control tool

- Option **B**: Users can clone entire repositories to their local systems

- Option **C**: Commits are possible even if offline

- Option **D**: Works are shared to the remote repository automatically by commit


<!-- grade_id: cell-71db91f1c18eebb3, points: 0, grade: false, solution: true -->
```python
# Assign the option letter of the option that you think is correct to the variable exercise_2.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_2 =
```

<!-- grade_id: cell-ad5700b91f015868, points: 2, grade: true, solution: false -->
```python
exercise_2_hash = "3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_2,str), "The variable should be a string."
assert exercise_2.isupper(), "Please use uppercase for your answer."
assert exercise_2.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_2_hash == hashlib.sha256(bytes(exercise_2, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 3: Staging
Which of the following statements is true:
- Option **A**: You can only have one file staged at a time
- Option **B**: You can only stage files after a commit
- Option **C**: The staging area is an intermediate storage for changes that will be part of the next commit
- Option **D**: All files are added to the staging area automatically

<!-- grade_id: cell-e9c408d24e570684, points: 0, grade: false, solution: true -->
```python
# Assign the option letter of the option that you think is correct to the variable exercise_3.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_3 =
```

<!-- grade_id: cell-6a4c40061538306a, points: 3, grade: true, solution: false -->
```python
exercise_3_hash = "6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_3,str), "The variable should be a string."
assert exercise_3.isupper(), "Please use uppercase for your answer."
assert exercise_3.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_3_hash == hashlib.sha256(bytes(exercise_3, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 4: Git add

What command should you use when you want to stage all text files changes in the workspace?

- Option **A**: `git add --all`
- Option **B**: `git add .`
- Option **C**: `git add *.txt`

<!-- grade_id: cell-ebcecd4226f1dce6, points: 0, grade: false, solution: true -->
```python
# Assign the option letter of the option that you think is correct to the variable exercise_4.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_4 =
```

<!-- grade_id: cell-df39aeb2d728fcd8, points: 2, grade: true, solution: false -->
```python
exercise_4_hash = "6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d"
options = ["A", "B", "C"]
assert isinstance(exercise_4,str), "The variable should be a string."
assert exercise_4.isupper(), "Please use uppercase for your answer."
assert exercise_4.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_4_hash == hashlib.sha256(bytes(exercise_4, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 5: Git add and Git commit

What is the difference between `git add .`; `git commit` and `git commit -a -m`?

- Option **A**: The first tracks and stages all files in the repo, the second commits all staged files in the repo, and the last stages and commits all changed files, followed by a message
- Option **B**: The first stages all files that start with '.', the second commits all staged files in the repo, and the last stages and commits all changed files, followed by a message
- Option **C**: The first stages all files in the repo, the second commits all modified files in the repo, and the last stages and commits all changed files, followed by a message

<!-- grade_id: cell-c696448a0429c1f8, points: 0, grade: false, solution: true -->
```python
# Assign the option letter of the option that you think is correct to the variable exercise_5.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_5 =
```

<!-- grade_id: cell-aff7a12614986551, points: 3, grade: true, solution: false -->
```python
exercise_5_hash = "559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd"
options = ["A", "B", "C"]
assert isinstance(exercise_5,str), "The variable should be a string."
assert exercise_5.isupper(), "Please use uppercase for your answer."
assert exercise_5.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_5_hash == hashlib.sha256(bytes(exercise_5, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 6: Git status

The `git status` command shows you the difference between the working directory and the index, where changes on tracked files are kept. 

In a repository where 5 files are being tracked, a new file is added. Additionally, the first line of one of the files is changed. What would the `git status` command show?

- Option **A**: Staged: 1, Modified: 1, Untracked: 0
- Option **B**: Staged: 0, Modified: 1, Untracked: 1
- Option **C**: Staged: 5, Modified: 1, Untracked: 1
- Option **D**: Staged: 5, Modified: 0, Untracked: 1

<!-- grade_id: cell-ad3eb1b9f7256b3d, points: 0, grade: false, solution: true -->
```python
# Assign the option letter of the option that you think is correct to the variable exercise_6.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_6 =
```

<!-- grade_id: cell-c266b16ef62872b4, points: 3, grade: true, solution: false -->
```python
exercise_6_hash = "df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_6,str), "The variable should be a string."
assert exercise_6.isupper(), "Please use uppercase for your answer."
assert exercise_6.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_6_hash == hashlib.sha256(bytes(exercise_6, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 7: .gitignore

Which one of the following statements is false: 

- Option **A**: You can stage a file that is on the .gitignore if you add it explicitly to the `git add` command
- Option **B**: You can't ignore the .gitignore file itself
- Option **C**: You can use patterns, such as wildcards, in the .gitignore



<!-- grade_id: cell-ce41a1f975fae767, points: 0, grade: false, solution: true -->
```python
# Assign the option letter of the option that you think is correct to the variable exercise_7.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_7 =
```

<!-- grade_id: cell-3dc6db0d118d1a92, points: 2, grade: true, solution: false -->
```python
exercise_7_hash = "df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c"
options = ["A", "B", "C"]
assert isinstance(exercise_7,str), "The variable should be a string."
assert exercise_7.isupper(), "Please use uppercase for your answer."
assert exercise_7.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_7_hash == hashlib.sha256(bytes(exercise_7, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 8: git push

What would happen if you tried to `git push` changes to a remote repository that is ahead on commits compared to your local repository?

- Option **A**: The remote repository will handle everything like magic
- Option **B**: The remote repository will throw an error because versions don't match
- Option **C**: The remote repository will throw an error because the remote repository is ahead on commits 

<!-- grade_id: cell-d8e48506ae23cd8e, points: 0, grade: false, solution: true -->
```python
# Assign the option letter of the option that you think is correct to the variable exercise_8.

# YOUR CODE HERE
raise NotImplementedError()
# exercise_8 =
```

<!-- grade_id: cell-808483824d7e116a, points: 3, grade: true, solution: false -->
```python
exercise_8_hash = "6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d"
options = ["A", "B", "C"]
assert isinstance(exercise_8,str), "The variable should be a string."
assert exercise_8.isupper(), "Please use uppercase for your answer."
assert exercise_8.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_8_hash == hashlib.sha256(bytes(exercise_8, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 9: Git Workflow
This exercise is ungraded. It's a set of instructions for you to follow and see for yourself how some git commands work and have some fun with it!

- Save this file
- On the root of your Prep Course repository, run the `git status` command
- Check the output, you should see that the **Exercise Notebook** file is modified
- Run the `git add` command to add **just** the Exercise Notebook file
- Run the `git commit` command with the option that allows you to write the commit message in the command. Write the following message: "I've conquered the power of git!"


Congratulations! You've just finished the first steps of your journey into the *wonderful world of version control*! These skills will accompany you for the rest of the Prep Course, so make good use of them! 

**See you in the Git Intermediate SLU for some more VCS fun!**

## Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-2c781dcf91e7533f, points: 0, grade: false, solution: true -->
```python
#slack_id = ...  
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4db9424844fb4cce, points: 0, grade: true, solution: false -->
```python
assert isinstance(slack_id,str)
submit(slack_id, 3)
```