# SLU11 | String & File Handling: Exercise Notebook

***

Now we're going to test how well you understood the learning notebook.

Also, this notebook is going to often require some **googling skills**. 

It's very important to learn **`how to google anything you don't remember or don't know how to do.`**

## Start by importing these packages

```python
# for evaluation purposes
import hashlib
import os
from IPython.display import display, Image

def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()

src = os.path.join('.','assets','eba68db97a.jpg') 
```

## PART 1

### Exercise 1
- Open a local file called **`'assignment.txt'`** and store it in a variable called **`'f'`**.
- Read the **whole file** and store it in a variable called **`'text'`**.

⚠️ _The variable should consist of **only one string representing all the lines from the file.**_

<!-- grade_id: cell-5c616f01a6b14b63, points: 0, grade: false, solution: true -->
```python
# f = ...
# text = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e2d149b3016dc2fc, points: 1, grade: true, solution: false -->
```python
assert isinstance(text, str), "Are you sure you read the whole file as one string (not as list of strings)?"
assert _hash(len(text)) == 'a980c09096', "The length of the string doesn't match"
assert _hash(text[0]) == '5ca379c9c1', "Did you read the correct file?"

print("---- all asserts passed ---- ")
```

### Exercise 2
- Move the **`read cursor`** to the **beginning of the file** stored in variable **`f`**.
- Read **each line** of the file and store it in a variable called **`lines`**.

⚠️ _The variable should be a **list of strings.**_

<!-- grade_id: cell-78e73bd6fb0635fc, points: 0, grade: false, solution: true -->
```python
# ...
# lines = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1ac86663d79cef7, points: 1, grade: true, solution: false -->
```python
assert _hash(len(lines)) != '5b4838043f', "Did you move the cursor to the beginning of the file?"
assert isinstance(lines, list), "Are you sure you stored the file as a list of strings?"
assert _hash(len(lines)) == 'd4fe3de58a', "The number of string elements in the list doesn't match"
assert _hash(lines[0]) =='0d3c2c6fb2', "Are you sure you read the correct file?"

print("---- all asserts passed ---- ")
```

***

## PART 2

### Exercise 3
Let's preprocess our file.
- For each **string** in the **`lines`** **list**:
    > remove all the **`newline characters`** and save the result in **the same variable**.

<!-- grade_id: cell-d3af9181b11c5b5b, points: 0, grade: false, solution: true -->
```python
# lines = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-62eba51f14c12c49, points: 1, grade: true, solution: false -->
```python
assert isinstance(lines, list), "Lines has to be a list of strings"
assert _hash(len(lines)) == 'd4fe3de58a', "The number of strings in the list doesn't match"
assert _hash(len(lines[0])) =='fdeb53ac14', "Are you sure you removed the newline characters?"

print("---- all asserts passed ---- ")
```

### Exercise 4

Remove all the **empty lines** in the **`lines`** **list** and save the result in **the same variable**:

<!-- grade_id: cell-cca50a46d15fd13d, points: 0, grade: false, solution: true -->
```python
# lines = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-048486c832b1aeb9, points: 1, grade: true, solution: false -->
```python
assert _hash(len(lines)) == '427566cc6c', "Are you sure that you removed all the empty strings?"

print("---- all asserts passed ---- ")
```

### Exercise 5
Concatenate all the **strings** in the **`lines` list** into one **long string**.

<!-- grade_id: cell-121fbc38e772e04a, points: 0, grade: false, solution: true -->
```python
# lines = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d6a132cb0e86cc5d, points: 1, grade: true, solution: false -->
```python
assert isinstance(lines, str), "It's not a string"
assert _hash(len(lines)) == 'd085807d6c', "The length of the string doesn't match"

print("---- all asserts passed ---- ")
```

### Exercise 6
Convert the **`lines`** (**string**) to a **list of words** and save it in a variable called **`'words'`**.

🌟 **Hint:** _Use space symbol as a separator_ 🌟

<!-- grade_id: cell-0023979f541f52e0, points: 0, grade: false, solution: true -->
```python
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84bffa431d75f2f6, points: 1, grade: true, solution: false -->
```python
assert isinstance(words, list), "It's not a list"
assert isinstance(words[0], str), "The list has to consist of strings"
assert _hash(len(words)) == '8be07733f2', "The number of words doesn't match"

print("---- all asserts passed ---- ")
```

### Exercise 7
Convert all the **strings** in **`words`** to lowercase.

<!-- grade_id: cell-cb00efd631ee3d37, points: 0, grade: false, solution: true -->
```python
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ab311a01acff014c, points: 1, grade: true, solution: false -->
```python
assert all([word.islower() for word in words]), "Not all the words are lowercase"

print("---- all asserts passed ---- ")
```

### Exercise 8
In **Natural Language Processing** - **`NLP`** we usually want to remove some words that have no meaning. They are usually called **`stop words`**.

We're not dealing with **NLP** yet, but let's also **remove such words**.
- From the **`words`** list remove the words that are in the **`stop_words`** list:
> stop_words = [ **`'on'`, `'an'`, `'in'`, `'a'`, `'of'`, `'the'`, `'and'`** ]

<!-- grade_id: cell-064d872c33ebafdd, points: 0, grade: false, solution: true -->
```python
# stop_words = ['on','an','in','a','of','the','and']
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e51d4948d757891, points: 1, grade: true, solution: false -->
```python
assert _hash(len(words)) == 'd03f94b383', "The number of words doesn't match"

print("---- all asserts passed ---- ")
```

### Exercise 9
Keep only unique words in the **`words`** list created on **exercise 8**.

Store them in the same variable. **It has to be a list of unique words!**

🌟 **Hint:** _You'll need some **Googling** on this question_ 🌟

<!-- grade_id: cell-aa8cd00bd90721fb, points: 0, grade: false, solution: true -->
```python
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e354a40c7e837c45, points: 1, grade: true, solution: false -->
```python
assert isinstance(words, list), "We need to store the unique words in a list"
assert _hash(len(words)) == 'f062f478b4', "The number of words doesn't match"

print("---- all asserts passed ---- ")
```

### Exercise 10
Create a function to find **if a word has more characters** than a given **number**:
- Create the function **`is_long()`**, which must:
- Receive **2 parameters:**
> * **`string`**: representing the word that we're checking.
> * **`integer`**: representing the minimum length that word should have.
- Return:
> * **`True`**: if length of the word **>=** `integer`;
> * **`False`**: if length of the word **<** `integer`;

<!-- grade_id: cell-2cfca9c135f57b08, points: 0, grade: false, solution: true -->
```python
# def is_long(word, number):
#    ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f83d24ebe38a8c0e, points: 1, grade: true, solution: false -->
```python
assert is_long('Exercise', 5), "The function didn't pass the test for the following parameters: word='exercise', num=5"
assert not is_long('Exercise', 10), "The function didn't pass the test for the following parameters: word='exercise', num=10"

print("---- all asserts passed ---- ")
```

### Exercise 11
- Apply **`is_long()`** function to **each word** in **`'words'`** variable you have created in **Exercise 9**.
- Use the value **`5`** for the **`integer`** parameter.
- Store the result in a variable called **`long_words`**.

<!-- grade_id: cell-d6886505b94e4ce5, points: 0, grade: false, solution: true -->
```python
# long_words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c31b9ac6bc23548, points: 1, grade: true, solution: false -->
```python
assert _hash(len(long_words)) == 'dadbc194c4', "The number of words doesn't match"

print("---- all asserts passed ---- ")
```

### Exercise 12
Create a function that **calculates how many words in a list start with a given letter**.
- Create the function **`start_with()`**, which must:
- Accept **2 parameters:**
> * **`words_list`**: list of words.
> * **`letter`**: letter to check.
- Return an **integer**, which is the **number of cases** when this letter is the **first letter in a word**.

<!-- grade_id: cell-2088771824f8e2a1, points: 0, grade: false, solution: true -->
```python
# def start_with(words_list, letter):
#    ...

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-8bb4a1579925fe13, points: 1, grade: true, solution: false -->
```python
assert isinstance(start_with(long_words, 'b'), int), "The function must return an integer."
assert _hash(start_with(long_words, 'b')) == 'cf2d85ea1d', "The function didn't work for long_words and letter='b'"
assert _hash(start_with(words, 'a')) == 'a72043e53f', "The function didn't work for words and letter='a'"
assert _hash(start_with(words, 'd')) == 'e4ffcdac53', "The function didn't work for words and letter='d'"

print("---- all asserts passed ---- ")
```

### Exercise 13

This task is going to be a bit **more difficult**, but also **pretty realistic.** 

We're going to create our own **`encoder`**!

> **But what is that?!**
>
> _**`Encoder`** is a program that transforms one format of data into another format. You're going to use them all the time on the academy. Now we're going to write a simple encoder by ourselves._

Our encoder is going to **check if a sentence has some predifined words.** Imagine that we have a lot of sentences and we want to understand if they are `positive` or `negative`.

Our idea is the following:

> If the sentence has words like **`'good'`**, **`'awesome'`**, **`'fantastic'`**, **`'hilarious'`** etc., it's **positive.** And it's **negative** otherwise.
> 
> Let's call these **good** words our **`vocabulary`**

Create a function, which:
- Is called **`encoder()`**.
- Accepts **2 inputs:**
> 1. sentence (string).
> 2. vocabulary (list of words).
- For each **word** in the **`vocabulary`**:
> checks whether this **word** is in the sentence.
- Returns a list with the same length of our **`vocabulary`**, but instead of **words** it has only **zeros and ones**:
> **`1`** means that this word was in the sentence.
>
> **`0`** otherwise.

**Example:**

`vocabulary = ['good', 'better', 'awesome', 'fantastic']
sentence = 'This day was fantastic. Tomorrow will be even better!'
encoder(sentence, vocabulary)`

> **Output:**
>
> `[0,1,0,1]`

**Explanation:**

**`'good'`** was **not** in the sentence &#8594; **0** <br>
**`'better'`** was in the sentence &#8594; **1** <br>
**`'awesome'`** was **not** in the sentence &#8594; **0** <br>
**`'fantastic'`** was in the sentence &#8594; **1** <br>

⚠️ **WARNING:** Don't forget to **`lowercase`** the words in the **sentence** and in the **vocabulary** ⚠️

Assume that all the words are **separated by `space symbol`**. If you're stuck with this task, **try to divide it in smaller parts.** 

***
**Tips:** 

We can divide this task into a few **smaller steps**:

1) First, we might need to **lowercase** each **word** in the **`vocabulary`**.

2) **Split** the **sentence** into a **list of words**:
> Let's call it **`words_list`**

3) Lowercase each **word** in this list.

4) For each **word** in our **`vocabulary`**, ask the program whether this word is in words_list:
> If **yes**, the answer is **1**.
> If it's **not**, the answer is **0**.

5) Return the **list of answers**.

***

<!-- grade_id: cell-24946c705e259797, points: 0, grade: false, solution: true -->
```python
# def encoder(sentence, vocabulary):
#    ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-28362929ccc11e01, points: 2, grade: true, solution: false -->
```python
assert isinstance(encoder('The day was good', ['good']), list), 'The function has to return a list'
assert _hash(len(encoder('The day was good', ['good']))) != '5b4838043f', 'Your function returns an empty list'
assert isinstance(encoder('The day was good', ['good'])[0], int), 'The elements of the output list have to be integers'
assert _hash(len(encoder('The day was good', ['good', 'bad', 'awesome']))) == 'e4ffcdac53', 'Number of elements in the output list does not match the length of the vocabulary'
assert _hash(encoder('The day was good', ['good', 'bad', 'awesome'])) == 'dec39382cd', "The function failed on the test: encoder('The day was good', ['good', 'bad', 'awesome'])"

print("---- all asserts passed ---- ")
```

***

## PART 3

### Exercise 14

`def function(a):
    assert isinstance(a, int)
    print(a + 1)`
    
⚠️ **`isinstance()`** function checks whether a variable has a specific format. Remember that from **`SLU-10`**?
    
What is going to happen if you call:

`function(3.5)`

Options:
- **a.** 4.5 will be printed.
- **b.** `AssertionError` will be raised.
- **c.** `AssertionError` will be printed.
- **d.** Nothing. The code is not valid.

Write the letter with the correct answer to a variable called **`answer_14`** as a **string**, for example:

**`answer_14 = 'a'`**

⚠️ **Note:** _try to think before writing anything. If you don't know the right answer, check the learning material._ ⚠️

But please try not to guess the answer by testing the function itself (**don't call it**).

_You're learning it for yourself, not for correct answers or some grades!_ 🙂 

<!-- grade_id: cell-213d3992590c59fb, points: 0, grade: false, solution: true -->
```python
# answer_14 = ''

# YOUR CODE HERE
raise NotImplementedError()
```

### Exercise 15

What's the output of the following line?

`'I am' + 17 + 'years old'`

Options:
- **a.** Name error.
- **b.** 'I am 17 years old'.
- **c.** Type error.
- **d.** 'I am years old 17'.

Save the answer in a variable called **`answer_15`**. Example:

**`answer_15 = 'a'`**

<!-- grade_id: cell-dcc00f6894cf84b2, points: 0, grade: false, solution: true -->
```python
# answer_15 = ''

# YOUR CODE HERE
raise NotImplementedError()
```

### Exercise 16

What's the result of the following lines?

`text = ['I', 'am', 'happy', 'today']
print(text[4])`

Options:
- **a.** 'today'.
- **b.** 'h'.
- **c.** Type error.
- **d.** Index error.

Save the answer in a variable called **`answer_16`**. Example:

**`answer_16 = 'a'`**

<!-- grade_id: cell-c26141d1871dd231, points: 0, grade: false, solution: true -->
```python
# answer_16 = ''

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6ea081f9ee1904d5, points: 2, grade: true, solution: false -->
```python
assert isinstance(answer_14, str), 'The answer should be a string'
assert isinstance(answer_15, str), 'The answer should be a string'
assert isinstance(answer_16, str), 'The answer should be a string'

answers = [answer_14, answer_15, answer_16]
assert _hash(answers) == '401711e05e', 'Oops. One or more answers are not correct! Try it again! :)'

print("---- all asserts passed ---- ")
```

### Exercise 17

Create a function called **`join_words()`** which:
* Accepts a **list of strings** called `list_strings` and a **string** `separator`.
* **Joins** all strings on `list_strings` with string `separator` between all words to a **new variable** called `sentence`.
* Returns `sentence`.
* If the operation **goes wrong** returns the following string **`"function inputs don't match the requirements"`**

<!-- grade_id: cell-ef39f20f9c0de7f0, points: 0, grade: false, solution: true -->
```python
#def join_words(list_strings, separator):

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-726e1e35c284e034, points: 1, grade: true, solution: false -->
```python
assert isinstance(join_words(["this", "is", "a", "string"], " "), str), "String output is expected"
assert join_words(["this", "is", "a", "string"], " ") == "this is a string", "The answer is wrong for input=['This', 'is', 'a', 'string']"
assert join_words(["this", "is", False, "a", "string"], " ") == "function inputs don't match the requirements", "Error message was not returned"

print("---- all asserts passed ---- ")
```

### Exercise 18

Create a function called **`multiply_not_five()`**, which:
- Accepts an **integer** `number` as a parameter.
- Returns the `number` multipled **by 2**.
- If the `number` is **equal to 5**, **raises an `AssertionError`**.

<!-- grade_id: cell-e778b52bc04648f8, points: 0, grade: false, solution: true -->
```python
#def multiply_not_five(number):

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-845a471dd5f6e303, points: 1, grade: true, solution: false -->
```python
assert _hash(multiply_not_five(10)) == 'b38f28b46b', 'Wrong answer for number = 10'
assert _hash(multiply_not_five(-5)) == '578946e522', 'Wrong answer for number = -5'

def check_answer(function):
    try:
        function(5)
        raise Exception("The function shouldn't work for number = 5")
    except AssertionError:
        pass
    except:
        raise Exception("You need to return an AssertionError, not another exception type")
        
check_answer(multiply_not_five)

print("---- all asserts passed ---- ")
```

### Exercise 19

The last exercise is pretty similar to **exercise 18.**

Create a function called **`multiply_string()`**, which:
- Accepts a **string** as a parameter.
- Returns the **string** repeated **15 times**.
> For example, `multiply_string('a') = 'aaaaaaaaaaaaaaa'`

❗ _You might want to **google** how to repeat strings in python_ ❗

- If the input is not a string, raise an **Exception** with some **text output.** 
> For example, `Exception("It's not a string, dude!")`

❗ _You might want to **google** how to check whether a variable is a string in python_ ❗

<!-- grade_id: cell-6f4c85c9880ef5d8, points: 0, grade: false, solution: true -->
```python
#def multiply_string(string):

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a055f00629e597a8, points: 2, grade: true, solution: false -->
```python
assert _hash(multiply_string('a')) == '11f86a430e', "The function didn't work for input='a'"
assert _hash(multiply_string('z')) == '6c6ed9ef14', "The function didn't work for input='z'"

def check_answer(function):
    try:
        function(5)
        raise Exception("The function shouldn't work for integers")
    except:
        pass
    
check_answer(multiply_string)

def finalize():    
    display(Image(filename=src, width=500))

print("---- all asserts passed ---- ")

finalize()
```

### And now... Mission Accomplished!!! 🎉🎉🎉
Now you can rest for this week, because your work was amaaaaaaazing (Again)!

**But don't forget...**

***

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-7281d78bfd08794b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# slack_id =
```

<!-- grade_id: cell-383debe5cc5ea67b, points: 0, grade: true, solution: false -->
```python
from submit import submit

assert slack_id is not None
submit(slack_id=slack_id, learning_unit=11)
```