# SLU10 | Object Oriented Programming - Inheritance: Exercise Notebook

***

## Start by importing these packages

```python
# for evaluation purposes
import hashlib
def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()

import math
import inspect

from utils import get_random_number
from utils import Spell
```

## PART 1

### Exercise 1

Create a class named `Worker()` and define the `__init__()` method of the created class to have three arguments:

1. name
2. age
3. department

Now, do the following:

- Assign these variables to attributes by using the `self` keyword;

- Define an extra attribute in the `__init__()` method, **without user input**, called `income` and with default value of **`1200`**;

- Define another one attribute in the `__init__()` method, called `company` and with default value  **`"Dunder Mifflin Paper Company"`**;

- Create a `presentation()` method that will **`return`** a string `speech` with the message:
<br><br>
    **_"Hi, my name is `name`! I'm work at `company` on the `department` department, nice to meet you."_**
<br><br>

- Create another method called `show_income()` that will **`return`** a string `salary` with the message:
<br><br>
    **_"My salary is € `income` per month"_**
<br><br>
In the end, create the `employer` object for a **"Jim Halpert"**, **25** from **"Sales"** department.


<!-- grade_id: cell-1201721e632289d8, points: 0, grade: false, solution: true -->
```python
class Worker():
    """
    Abstract base class representing a Worker
    """

    # def __init__(...)    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # def presentation(...)    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # def show_income(...)    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Creating our first Worker!

employer = Worker(name='Jim Halpert',
                  age=25, 
                  department='Sales')
```

<!-- grade_id: cell-3a68fb1816240298, points: 1, grade: true, solution: false -->
```python
assert _hash(employer.name) == 'aeb6411c48', 'Ops!! Did you set the name?'
assert _hash(employer.age) == 'f4073f1323', 'Something is wrong with the age...'
assert _hash(employer.company) == '656a034445', 'Ops!! Did you set the company as we said to?'
assert _hash(employer.department) == 'f0f668e3a2', 'Where are the department?'
assert _hash(employer.income) == 'f9ef548c29', 'Ops!! We are talking about money, right?'

assert _hash(type(employer.age)) == 'd461979a1d', 'Are you sure that `age` is a numerical value?'
assert _hash(type(employer.income)) == 'd461979a1d', 'Are you sure that `income` is a numerical value?'

assert _hash(employer.presentation()) == 'e75d412926', "Something going wrong with the presentation() method"
assert _hash(employer.show_income()) == '5f5cc6ca00', "Something going wrong with the show_income() method"

print("---- Yay! All asserts passed ---- ")
```

```python
# Please dear employer, present yourself:
print(employer.presentation())

# And sorry for the indiscretion, but how much do they pay to you?
print(employer.show_income())

# Oh, ok, Thanks!
```

<img src="./assets/exe_01.jpg" width="700"/>

### Exercise 2

One of the mainly benefits of **Inheritance** is that we can create more complicated classes that inherit variables or methods from their **parent** classes. This saves us time and helps us build more complex objects, since these **child** classes can also include additional variables or methods. 


- Create another class, **which inherits `Worker`**, and call it `Manager`;
- Define a new method `promote()` that will multiply the `income` **by** a (pseudo)random integer from 2 to 10; ❗
- Then, implement an object called `boss` for a **"Michael Scott"**, **35** from **"Management"**;

So, make sure that `boss` belongs to `Manager` class.

❗ _use the **`get_random_number()`** function, already imported in the first line, for generating pseudo-random integers._ ❗

<!-- grade_id: cell-3dd201ca47eb5496, points: 0, grade: false, solution: true -->
```python
# class Manager(...)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Now let create our Manager, and call him a Boss!

boss = Manager(name='Michael Scott',
                  age=35, 
                  department='Management')
```

<!-- grade_id: cell-c4f20bb4fd2d42cf, points: 1, grade: true, solution: false -->
```python
assert _hash(boss.name) == '7d84d32222', 'Ops!! Did you changed something with the name?'
assert _hash(boss.age) == '7d8e4de993', 'Ops!! Did you changed something with the age?'
assert _hash(boss.company) == '656a034445', 'Ops! The company should not change'
assert _hash(boss.department) == 'c24f31960b', 'Ops!! Did you changed something with the department?'

# Now all our boss need is a promotion!
# Let's promote him...
boss.promote()
assert _hash((boss.income >= (1200*2)) & (boss.income <= (1200*10))) == 'dd931fabaa', 'Ops! The promotion did not work'

assert _hash(boss.presentation()) == '11be1d4ed9', 'Ops!! Did you changed something with the presentation() method?'
assert _hash((isinstance(boss, Manager))) == 'dd931fabaa', 'boss must belong to the Manager Child Class'
assert _hash((isinstance(boss, Worker))) == 'dd931fabaa', 'boss must belong to the Worker Parent Class'
assert _hash((isinstance(employer, Worker))) == 'dd931fabaa', 'Did you changed something with our employer?'
assert _hash((isinstance(employer, Manager))) == '3ab05b19da', 'Did you changed something with our employer?'

# Well, take a look at these details!
print("Instance test:", isinstance(boss, Manager),'| boss belongs to the Manager Child Class')
print("Instance test:", isinstance(boss, Worker),'| boss belongs to the Worker Parent Class')
print('')
print("Type test:", type(boss),'| boss is the type of Manager Class')
print("Type test:", type(employer),'| employer is the type of Worker Class')
print('')
print("---- Well Done, all asserts passed ---- ")
```

```python
# Hello, amazing boss, can you also present yourself?
boss.presentation()

# I heard you were promoted, so tell me the news!
boss.show_income()
```

<img src="./assets/exe_02.jpg" width="700"/>

***

## Part 2

### Exercise 3

Now, let's take a look at these enchanted charms: 

<img src="./assets/spell.PNG" width="700"/>

What's the **`parent`** class in this code above?

Options:

* a. Spell.
* b. Hocus_Pocus.
* c. Tontus_Talontus.
* d. put_spell.

Write the letter with the **correct answer** to a variable called **`answer_3`** as a string, for example:

    answer_3 = 'c'

⚠️ **Note:** _try to think before writing anything. If you don't know the right answer, check the learning material._ ⚠️

But please try not to guess the answer by testing the function itself (**don't call it**).

_You're learning it for yourself, not for correct answers or some grades!_ 🙂

<!-- grade_id: cell-cd3f8ea886177e3f, points: 0, grade: false, solution: true -->
```python
# answer_3 = ''

# YOUR CODE HERE
raise NotImplementedError()
```

### Exercise 4

Which option is correct about the **`child`** class in this code above?

Options:

* a. Spell.
* d. Hocus_Pocus.
* c. Tontus_Talontus.
* d. Spell and Hocus_Pocus.
* e. Hocus_Pocus and Tontus_Talontus.
* f. Tontus_Talontus and Spell.

Save the answer in a variable called **`answer_4`**. Example:

    answer_4 = 'c'

<!-- grade_id: cell-34c6856836b6fe52, points: 0, grade: false, solution: true -->
```python
# answer_4 = ''

# YOUR CODE HERE
raise NotImplementedError()
```

### Exercise 5

What's the output of the following line?

**`put_spell(Tontus_Talontus())`**

Options:

* a. `'Pocus Charm | Hocus Pocus | Causes the victim to become confused and dizzy'` will be printed.
* b. `'Pocus Charm | Hocus Pocus | No description'` will be printed.
* c. Nothing. The code is not valid.
* d. `'Tontus Charm | Tontus_Talontus | Causes the victim to become confused and dizzy'` will be printed.
* e. `'Tontus Charm | Tontus_Talontus | No description'` will be printed.

Save the answer in a variable called **`answer_5`**. Example:

    answer_5 = 'c'

<!-- grade_id: cell-f7d2eb502855efb2, points: 0, grade: false, solution: true -->
```python
# answer_5 = ''

# YOUR CODE HERE
raise NotImplementedError()
```

### Exercise 6

What's the output of the following lines?

**`spell = Hocus_Pocus()
put_spell(spell)`**

Options:

* a. `'Pocus Charm | Hocus Pocus | Causes the victim to become confused and dizzy'` will be printed.
* b. `'Pocus Charm | Hocus Pocus | No description'` will be printed.
* c. Nothing. The code is not valid.
* d. `'Tontus Charm | Tontus_Talontus | Causes the victim to become confused and dizzy'` will be printed.
* e. `'Tontus Charm | Tontus_Talontus | No description'` will be printed.

Save the answer in a variable called **`answer_6`**. Example:

    answer_6 = 'c'

<!-- grade_id: cell-7e7cf3b638acec98, points: 0, grade: false, solution: true -->
```python
# answer_6 = ''

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-47eb6eb6bd431088, points: 2, grade: true, solution: false -->
```python
assert isinstance(answer_3, str), 'The answer should be a string'
assert isinstance(answer_4, str), 'The answer should be a string'
assert isinstance(answer_5, str), 'The answer should be a string'
assert isinstance(answer_6, str), 'The answer should be a string'

answers = [answer_3, answer_4, answer_5, answer_6]
assert _hash(answers) == 'a437db6d5b', 'Oops. One or more answers are not correct! Try it again! :)'

print("---- all asserts passed ---- ")
```

<img src="./assets/hocus_pocus.jpg" width="500"/>

Do you know another better way to get more practice with **OOP Inheritance** besides biology?

_Yeap... Exactly with Math (and geometry)!_ 📐

### Exercise 7

Start by creating a new class `Rectangle` with the following:
+ Attributes
    - **length**: float
    - **width**: float
    
    
+ Methods
    - `get_area()`: returns the product of the rectangle **length** with **width**
    - `get_perimeter()`: returns the sum of double the **lenght** and double the **width**

<!-- grade_id: cell-e1f05519ef39b0f6, points: 0, grade: false, solution: true -->
```python
# class Rectangle():

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c0da80882a9d38bc, points: 2, grade: true, solution: false -->
```python
# considering a 28,04 m by 17,11 m rectangle:
r = Rectangle(length=28.04, width=17.11)

assert math.isclose(r.get_area(), 479.765, abs_tol=0.001), 'Area computation is not correct!'
assert math.isclose(r.get_perimeter(), 90.300, abs_tol=0.001), 'Perimeter computation is not correct!'
assert _hash(type(r)) == 'f533583264', 'You should implement a Rectangle Class!'

print("Retangle Area:", r.get_area(),'| The Area is correct!')
print("Retangle Perimeter:", r.get_perimeter(),'| The Perimeter is correct!')
print('')
print("---- UhuuuuL! all asserts passed ---- ")
```

### Exercise 8

Now, let's implement another object:

- Create a `Square` class, as a **child** of the `Rectangle` class. 

- Change **only** the `__init__()` method in a way that now it **only** accepts one argument **`side`**. 

- You can call the `__init__()` method from the **parent** class and assign both **`length`** and **`width`** to be equal to **`side`**

<!-- grade_id: cell-9262f02e916a66f1, points: 0, grade: false, solution: true -->
```python
# class Square(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0d959343e7bd5637, points: 1, grade: true, solution: false -->
```python
# considering a 28,04 m by 28,04 m square:
s = Square(side=28.04)

assert math.isclose(s.get_area(), 786.241, abs_tol=0.001), 'Area computation is not correct!'
assert math.isclose(s.get_perimeter(), 112.16, abs_tol=0.001), 'Perimeter computation is not correct!'

error_message = 'Did you use the super() function?'
assert 'self.width' not in inspect.getsource(Square.__init__), error_message
assert 'self.length' not in inspect.getsource(Square.__init__), error_message
assert "super()." in inspect.getsource(Square.__init__), error_message

error_message =  'Do not re-implement the methods, use the ones from the Rectangle :)'
assert 'self.length' in inspect.getsource(Square.get_area), error_message
assert 'self.width' in inspect.getsource(Square.get_area),  error_message
assert 'self.length' in inspect.getsource(Square.get_perimeter), error_message
assert 'self.width' in inspect.getsource(Square.get_perimeter), error_message


print("Square Area:", s.get_area(),'| The Area is correct!')
print("Square Perimeter:", s.get_perimeter(),'| The Perimeter is correct!')
print('')
print("---- all asserts passed ---- ")
```

### Exercise 9

Create a new class `Triangle`, independent from the other classes:

+ Attributes
    - **base**: float
    - **height**: float
    
    
+ Methods
    - `get_tri_area()` which calculates the area of the triangle (do you remember the formula? 🤓)

<!-- grade_id: cell-361d990792a542f3, points: 0, grade: false, solution: true -->
```python
# class Triangle():

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a0f3c54e307aeee9, points: 1, grade: true, solution: false -->
```python
# considering a 28 m by 17 m triangle:
t = Triangle(base=28, height=17)

assert math.isclose(t.get_tri_area(), 238, abs_tol=0.001), 'The area of the triangle is miscalculated!'

print("Triangle Area:", t.get_tri_area(),'| The Area is correct!')
print('')
print("---- all asserts passed ---- ")
```

### Exercise 10

Now for the `RightPyramid` class, create methods that compute the area of this geometric solid using both the `Square` and `Triangle` classes as **parents**.

**Override** the `get_area()` method in order to compute the equation of the **RightPyramid area**. 

🌟 **Hint:** _use the `super()` function to  make everything easier._ 🌟

So, you're in luck, geometry is not the subject for this SLU! Here are the formulas that you need:

* The area of a Right Pyramid is equal to the sum of the base_area with the its lateral_area.
* The Lateral area can be computed as the next image shows (1/2 * perimeter * slant_height)

<img src="assets/pyramid_lateral_area.jpg" width="500"/>

Other options do exist to compute the lateral area, like using the **triangle area**. Feel free to use what is easier for you as **this is about python, not geometry!**

+ Attributes:
    - **square_side** (this is equivalent to one side of the square base, just like the previous exercises. So, base is also equal to 1/4 of the perimeter `p` of the square base of the pyramid)
    - **slant_height**
    
    
+ Methods:
    - `get_area()`

<!-- grade_id: cell-18e78aa00977daf8, points: 0, grade: false, solution: true -->
```python
# class RightPyramid(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7a3b92fb1669680d, points: 3, grade: true, solution: false -->
```python
# Creating one Right Pyramid, right now!
rp = RightPyramid(square_side=17, slant_height=28)


assert isinstance(rp, RightPyramid), 'ops, are you sure it is a right pyramid object?'
assert isinstance(rp, Square), 'Probably your class is not properly defined, should be a child class of the Square class'
assert isinstance(rp, Triangle), 'Probably your class is not properly defined, should be a child class of the Triangle class'

error_message = 'Did you use the super approach?'
assert 'super()' in inspect.getsource(rp.__init__), error_message
assert 'super()' in inspect.getsource(rp.get_area), error_message

assert math.isclose(rp.get_area(), 1241, abs_tol=0.001), 'The area of the right pyramid is miscalculated!'

print("Right Pyramid Area:", rp.get_area(),'| The Area is correct!')
print('')
print("---- all asserts passed ---- ")
```

Feeling of the day:
<img src="assets/calculus.png" width="500"/>

***

## Part 3

Uhuuul! 🎉 This is the last part of this notebook exercise! And you're rocking it! 🤘🤘🏻🤘🏼🤘🏽🤘🏾🤘🏿



In the following exercice the goal it is to create a **BankAccount program**, composed by **one grand-parent class**, **one parent class** and **one child class**, in order to manage the clients accounts. This is the general idea:

<img src="assets/exe_08.jpeg" width="600"/>



### Exercise 11

First things first, the `BankAccount` class:

+ Attributes
    - **customer**
    - **number**
    - **balance** with default value of 0 (zero) but passed as an argument anyway (you'll need it for **Ex.9** and **Ex.10**)
    - **currency** variable is always € (no need to pass it as an argument)
    
    
+ Methods
    - `deposit()`: 
        * Deposit **amount** into the bank account. (increase the balance)
        * If the **amount** is invalid (< 0) it should `return` an alert message to the user
        
      ⚠️ **Warning**: it should not deduce the amount and `return`, such as **`Invalid deposit amount`** message. <br><br>
    - `withdraw()`:
        * Withdraw **amount** from the bank account. (decrease the balance)
        * This method must ensure that the account has sufficient funds for the asked **amount**. 
        
      ⚠️ **Warning**: Again, it should be > 0 and also `return` an alert if the value is less than zero or if it's bigger than the current balance 
      
      (a.k.a **`Invalid withdraw amount`** message) <br><br>
    - `check_balance()`: `print()` a statement of the account balance, example **`"The balance of account number 48522888 is €1000"`**

<!-- grade_id: cell-aec63d5088ce5f30, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    """
    Abstract base class representing a bank account
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5fb3aa9dc9bb6e4c, points: 2, grade: true, solution: false -->
```python
# Now, let's create our account!
account = BankAccount(customer='Dwight Schrute', number=21457666)

assert isinstance(account, BankAccount), 'Your class is not properly defined'
assert account.number == 21457666, 'You have to define an "number" attribute for your class!'
assert _hash(account.customer) == 'dd1e4dfba6', 'You have to define a "customer" attribute for your class!'
assert account.balance == 0, 'The default value of "balance" should be 0!'
assert 'balance=0' in inspect.getsource(account.__init__), 'balance=0 must be passed as an argument anyway'
assert account.currency == '€', 'Currency symbol is not working... Did you change the existing line?'

print("The account for %s was successfully created. \nHis account number is %s and for now his account is empty."
     % (account.customer, account.number))
print('')
print("---- Horay! All asserts passed ---- ")
```

<!-- grade_id: cell-f8e3492932f9b714, points: 1, grade: true, solution: false -->
```python
# Let's deposit that big, fat check! 
account.deposit(amount=25000)

assert account.balance == 25000, 'Your balance, after making a deposit should change!'

print("%s just made a deposit of €25'000. \nHis new account balance is: %s%s"
      %(account.customer, account.currency, account.balance))
print('')
print("---- Yeah! all asserts passed ---- ")
```

<!-- grade_id: cell-a12d39bd6fbae5c7, points: 1, grade: true, solution: false -->
```python
# And then just go shopping
# online... better be online, right?!

account.withdraw(amount=2250)

assert account.balance == 22750, 'Your balance, after making a withdrawal should change or it is not correct!'

print("%s bought the best fertilizers for his beet production. \nHis new account balance is: %s%s"
      %(account.customer, account.currency, account.balance))
print('')
print("---- Yoooo! all asserts passed ---- ")
```

<!-- grade_id: cell-aaeb81da6d6cc191, points: 1, grade: true, solution: false -->
```python
# checking the returned message for invalid deposit/withdraw:
assert type(account.deposit(-10)) == str, 'Are you returning a string message for invalid deposit?'
assert type(account.withdraw(-10)) == str, 'Are you returning a string message for invalid withdraw?'

# Finally let's ensure that your program does not accept invalid operations:
account.deposit(amount=-10)
assert account.balance == 22750, 'Your balance, after making an invalid deposit should not change!'
account.withdraw(amount=-10)
assert account.balance == 22750, 'Your balance, after making an invalid withdraw should not change!'

print("---- Amaaazing! all asserts passed ---- ")
```

### Exercise 12

Now we need to implement the `CheckingAccount` class, as a **child class** of `BankAccount`, that has a similar structure of the previous exercise:

+ Attributes:
    - **annual_fee**
    - **withdraw_limit**
    - All the BankAccount attributes
    
    
+ Methods:
    - `withdraw()`: 
        **override** the **parent class** method and create an `if` clause with an extra condition, in order to limit the possible amount that a person can withdraw daily.<br><Br>
        
    - `charge_fee()`: 
        Deduct the annual fee from the account balance, making sure that the account **never goes negative**! 
    
       🌟 **Hint:** use python's `max()` function 🌟

<!-- grade_id: cell-6d21acc19869fec0, points: 0, grade: false, solution: true -->
```python
# class CheckingAccount(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4f6963493a0e8fe9, points: 1, grade: true, solution: false -->
```python
# Updating our account to be a Checking Account
cheking_account = CheckingAccount(customer='Dwight Schrute', number=21457666, annual_fee=20., withdraw_limit=200)

assert isinstance(cheking_account, CheckingAccount) == True, 'Your class is not properly defined'
assert isinstance(cheking_account, BankAccount) == True, 'Your class has no relation with the BankAccount class'

assert cheking_account.number == 21457666, 'You have to define an "number" attribute for your class!'
assert _hash(cheking_account.customer) == 'dd1e4dfba6', 'You have to define a "customer" attribute for your class!'
assert cheking_account.balance == 0, 'The default value of "balance" should be 0!'
assert cheking_account.currency == '€', 'Currency symbol is not working... Did you changed the existing line?'
assert cheking_account.annual_fee == 20, 'You have to define a "annual fee" attribute for your class!'


print("---- Weee! all asserts passed ---- ")
print('')
print("Now, the new cheking account for %s was successfully created. \nHis account number is %s and for now his account is empty."
     % (cheking_account.customer, cheking_account.number))
```

<!-- grade_id: cell-8f5b292e2c0a7fc0, points: 1, grade: true, solution: false -->
```python
#Let's do something crazy! Witdhraw more than balance...
cheking_account.withdraw(amount=220)

assert cheking_account.balance == 0, "The created current account must not authorize withdraws when the amount is higher than the account balance"

# Making a new deposit and checking balance:
cheking_account.deposit(amount=750)
cheking_account.check_balance()
assert cheking_account.balance == 750, 'Ops! Did you change something on deposit() method?'

# Another withdraw that should be blocked by withdraw limit:
cheking_account.withdraw(amount=220)
assert cheking_account.balance == 750

# Charge current account annual fee:
cheking_account.charge_fee()
assert cheking_account.balance == 730

print("---- Great!! all asserts passed ---- ")
print('')
print(cheking_account.check_balance())
print('Applying annual fee and...')
cheking_account.charge_fee()
print(cheking_account.check_balance())
```

### Exercise 13

Finally, let's create the `InvestimentAccount` class, as a **child class** of `CheckingAccount` (🧩 _Yup! It's right, we have a multilevel inheritance here_ 🧩) 

+ Attributes:
    - **ticker** 🔍 _a code that represents a company in the stock market_ 🔎
    - **investment_fee**
    - All the CheckingAccount attributes ⚠️ (the **`withdraw_limit`** needs always to be the same as the balance)
    
    
+ Methods:
    - `withdraw()`: 
        **override** the method **again**, and now, deduce from the account balance **`amount + investment_fee`**.
    
       🌟 **Hint:** remember that the account **never goes negative**!   🌟
       <br><Br>
    
    - `bet_on_market()`: 
        will multiply the balance **by** a (pseudo)random integer from 1 to 10; ❗ 
    
       ❗ **Hint:** use **`get_random_number()`** function again, remember?! ❗
    

<!-- grade_id: cell-2b45d8e46269b879, points: 0, grade: false, solution: true -->
```python
# class InvestimentAccount(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fbeaf6f831b1a0c4, points: 2, grade: true, solution: false -->
```python
# Let's open the our new investment account:
investment_account = InvestimentAccount(
    customer='Dwight Schrute',number=21457666,annual_fee=20,ticker='DMP',investment_fee=10
)

assert isinstance(investment_account, InvestimentAccount) == True, 'Your class is not properly defined'
assert isinstance(investment_account, CheckingAccount) == True, 'Your class has no relation with the CheckingAccount class'
assert isinstance(investment_account, BankAccount) == True, 'Your class has no relation with the BankAccount class'

assert investment_account.number == 21457666, 'You have to define an "number" attribute for your class!'
assert _hash(investment_account.customer) == 'dd1e4dfba6', 'You have to define a "customer" attribute for your class!'
assert investment_account.balance == 0, 'The default value of "balance" should be 0!'
assert investment_account.currency == '€', 'Currency symbol is not working... Did you changed the existing line?'
assert investment_account.annual_fee == 20, 'You have to define a "annual fee" attribute for your class!'
assert investment_account.investment_fee == 10, 'Ops! Where is the "investment fee" attribute of your class?!'


assert _hash(investment_account.ticker) == '941024a0fd', 'Ops! what ticker are we investing?'
assert investment_account.withdraw_limit == 0, 'Investiment account does not have a properly withdraw limit...'

investment_account.deposit(amount=1570)
assert investment_account.balance == 1570, 'Ops! Did you change something on deposit() method?'

investment_account.bet_on_market()
assert _hash((investment_account.balance >= (1570*1)) & (investment_account.balance <= (1570*10))) == 'dd931fabaa', 'Ops! The bet_on_market() did not work'

investment_account.withdraw(250)
assert _hash((investment_account.balance >= ((1570*1)-10) & (investment_account.balance <= ((1570*10)-10)))) == 'dd931fabaa', 'Did you deduce "investment fee" on withdraw() method?'

investment_account.withdraw(investment_account.balance+5)
assert _hash((investment_account.balance >= ((1570*1)-10) & (investment_account.balance <= ((1570*10)-10)))) == 'dd931fabaa','You cannot withdraw more funds than you have...'

print("---- Everything is awesome!! all asserts passed ---- ")
print('')
print(investment_account.check_balance())
print('{} is playing on market again...'.format(investment_account.customer))
print('And now...')
investment_account.bet_on_market()
print(investment_account.check_balance())
```

<img src="./assets/exe_10.jpg" width="500"/>

***

### And now... Mission Accomplished!!! 🎉🎉🎉


Now you can rest proudly,
because your work was excellent!

**But don't forget...**

***

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-a48c05d8f2f06b81, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# slack_id =
```

<!-- grade_id: cell-6e99685f6f8c1895, points: 0, grade: true, solution: false -->
```python
from submit import submit

assert slack_id is not None
submit(slack_id=slack_id, learning_unit=10)
```