# SLU 17 - Exam Prep I

## Batch 2 - Wave 1 Python Exam 
This is the Python exam from batch 2 - wave 1. Exam Duration: 2h

```python
from operator import itemgetter
#used for evaluation
import utils
```

## Question 1

If you took a look at the dataset, you may have noticed that some of the numeric values are "polluted" with characters.

In this task you'll have to write a function, `clean_percentage`, that should do the following:

- Receive as argument a list of dictionaries;
 - Each of the dictionaries keeps the % of female students in higher education per area, for a specific year;
- Remove all the characters and spaces from the numeric values;
- Return the cleaned data;

Hint: mind the naughty (R)'s :)

<!-- grade_id: cell-c7177b04bd669cff, points: 0, grade: false, solution: true -->
```python
def clean_percentage(data):
    ### BEGIN SOLUTION
    res = [{k: v.replace('(R) ', '') for k, v in r.items()} for r in data]
    return res
    ###END SOLUTION
```

<!-- grade_id: cell-18bc7007043d3ec2, points: 3, grade: true, solution: false -->
```python
utils.b2w1_exerc_1_grading(clean_percentage)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 2

In this task, we want to clean the names of the education areas, so they are all lower case and without white spaces.

For instance, Social Sciences, Management and Law should become social_sciences_management_and_law.

So you'll have to write two functions.

The first function, `clean_header_string`, should do the following:

- Receive a string as input;
- Remove all the commas from the string;
- Replace the white spaces in the string with underscores;
- Convert the string to lower case;
- Return the transformed string;

The second function, `clean_header`, should do the following:

- Receive as input, a list of dictionaries, similar to the one in the last question;
- Use the first function to transform the education area names;
- Return the transformed data;


<!-- grade_id: cell-a11732fd19b01c68, points: 0, grade: false, solution: true -->
```python
def clean_header_string(string):
    ### BEGIN SOLUTION
    string = string.replace(',', '').replace(' ', '_') 
    return string.lower()
    ### END SOLUTION

def clean_header(d):
    ### BEGIN SOLUTION
    return [{clean_header_string(k): v for k, v in r.items()} for r in d]
    ### END SOLUTION
```

<!-- grade_id: cell-d7215a1617be8293, points: 3, grade: true, solution: false -->
```python
utils.b2w1_exerc_2_grading(clean_header_string, clean_header)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 3

In this task, we want to convert our data to the right data types.

For this you'll have to write two functions.


The first function, `commas`, should do the following:

- Receive as input a list of dictionaries, similar to the previous questions;
- Replace the commas with dots in the numeric values;
- Return the transformed data;


The second function `data_types` should do the following:

- Receive as input a list of dictionaries (specifically the one returned by the first function);
- Convert the year values to integer;
- Convert the % of female students values to float;
- Return the transformed data;

<!-- grade_id: cell-76ca0448c5209f19, points: 0, grade: false, solution: true -->
```python
def commas(d):
    ###BEGIN SOLUTION
    return [{k: v.replace(',', '.') for k, v in r.items()} for r in d]
    ###END SOLUTION


def data_types(d):
    ###BEGIN SOLUTION
    return [{k: int(v) if k == 'year' else float(v) for k, v in r.items()} for r in d]
    ###END SOLUTION
```

<!-- grade_id: cell-c0f514108d4869b2, points: 3, grade: true, solution: false -->
```python
utils.b2w1_exerc_3_grading(commas, data_types)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 4

In this task, we want to focus on data about a specific education area.

So you'll have to write a function, `education_years`  that should do the following:

- Receive as input a list of dictionaries (similar to the ones of the previous questions), and an education area (string);
- Return a list of tuples where the first value of the tuple is the year and the second value is the % of females in the given education area;

In this task you're required to use the following:

- `map` function;
- `lambda` function;

<!-- grade_id: cell-0546a2b6445fd75b, points: 0, grade: false, solution: true -->
```python
def education_years(data, education_area):
    ###BEGIN SOLUTION
    res = list(map(lambda x: (x['year'], x[education_area]), data))
    return res
    ###END SOLUTION
```

<!-- grade_id: cell-0bf157ce02b1b085, points: 3, grade: true, solution: false -->
```python
utils.b2w1_exerc_4_grading(education_years)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 5

In this task, for a certain education area, we want to understand what were the years where more females were enrolled in higher education.

So you'll have to write a function named `female_enrolled` that should do the following:

- Receive as input a list of tuples with year and % of females (like the output of last task) and a threshold;
- Filter **out** all the tuples that correspond to a % of females strictly below the threshold;
- Sort the remaining tuples from higher to lower % of females;
- Return the resulting list of tuples;

In this task you're required to use the following:

- `sorted` function

Hint: `itemgetter` may be handy

<!-- grade_id: cell-3c339e619872e065, points: 0, grade: false, solution: true -->
```python
def female_enrolled(data, threshold):
    ###BEGIN SOLUTION
    filtered_field_data = [x for x in data if x[1] >= threshold]
    solution = sorted(filtered_field_data, key=itemgetter(1), reverse=True)
    return solution
    ###END SOLUTION
```

<!-- grade_id: cell-30f67a7608d49607, points: 3, grade: true, solution: false -->
```python
utils.b2w1_exerc_5_grading(female_enrolled, 58.1)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 6

In this task we'll create a class to play the lottery!

First, get to know the `raffle_key`. The `raffle_key` is a dictionary, associating ticket numbers with prizes. All the numbers that are not in the raffle_key won't get any prize.

```python
raffle_key = {1:"10 euros", 5:"100 euros", 9:"40 euros"}
```

### Step 1:

- Create a `Lottery` class that takes one constructor parameter called `raffle_key`;
- In the `Lottery` class create a method named _get_lucky_ that receives the lucky_number as parameter;
    and returns the strings "{X} euros" or "Better luck next time!", depending on the presence of the lucky number in the raffle key;
- In our example above, our lucky number 9 corresponds to the response string "40 euros"
    if in some case we get a lucky number that doesn't exist in the raffle key, we should print out "Better luck next time!";


### Step 2:

- Instantiate the `Lottery` class with name `simple_lottery` and with the given raffle key;
- Call the method `get_lucky` with the given lucky number and print out the result;


In this task you're required to use the following:

- `try/except` statements;


<!-- grade_id: cell-dc8077515454abe0, points: 0, grade: false, solution: true -->
```python
class Lottery():
###BEGIN SOLUTION
    def __init__(self, raffle_key):
        self.raffle_key = raffle_key
    
    def get_lucky(self, lucky_number):
        try:
            return self.raffle_key[lucky_number]
        except KeyError:
            return "Better luck next time!"


simple_lottery = Lottery(raffle_key)      
###END SOLUTION
```

<!-- grade_id: cell-96c7d9be71562f8a, points: 5, grade: true, solution: false -->
```python
utils.b2w1_exerc_6_grading(Lottery, simple_lottery, raffle_key)
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

### Last but not least, submit your work!

To submit your work, fill your slack ID in the `slack_id` variable (as a string).

Example: `slack_id = "x-men"`

Help: if you forgot your slack ID, [read this](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c).

<!-- grade_id: cell-140bf2d2262ef46e, points: 0, grade: false, solution: true -->
```python
# Submit your work!

#slack_id = 

### BEGIN SOLUTION
slack_id = "ADMINSLU17_2"
### END SOLUTION
```

<!-- grade_id: cell-570d9013454a9e37, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 17_2
submit(slack_id, slu)
```

Output:
```
Success
```