# SLU 17 - Exam Prep I

## Batch 2 - Wave 2 Python Exam
This is the Python exam from batch 2 - wave 2. Exam Duration: 2h

```python
import json
from operator import itemgetter
import random
#used for evaluation
import utils
```

## Question 1

The variable `data` is a list of dictionaries.

Create a function `replace_unknown` that:
- Receives as argument a list of dictionaries;
- In each dictionary, it replaces all the '?' you find with the string value '1';
- Returns the modified list of dictionaires;

Note: values should retain the original datatypes.

<!-- grade_id: cell-7d5d16f632d819c3, points: 0, grade: false, solution: true -->
```python
def replace_unknown(data):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2ec0be01019aa1ce, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_1_grading(replace_unknown)
print("Answer is correct. Good Job.")
```

## Question 2

The variable `data` is a list of dictionaries.

Create a function `convert_int` that:
- Receives as argument a list of dictionaries;
- In each dictionary, find all the numeric values and make sure their data type is `int`;
- Returns the modified list of dictionaires;

Hint: `isdigit()`

<!-- grade_id: cell-5b768d0072b87818, points: 0, grade: false, solution: true -->
```python
def convert_int(data):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9138a4e7a89f3adc, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_2_grading(convert_int)
print("Answer is correct. Good Job.")
```

## Question 3

The variable `data` is a list of dictionaries.

Create a function `lower_class` that:

- Receives as argument a list of dictionaries;
- For each element of the list, find the key 'class_type' and convert the correspondent value to lower case;
- Returns the modified list of dictionaires;

<!-- grade_id: cell-0f7774d47c55eb30, points: 0, grade: false, solution: true -->
```python
def lower_class(data):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3a68fa6f6e905977, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_3_grading(lower_class)
print("Answer is correct. Good Job.")
```

## Question 4

The variable `data` is a list of dictionaries.

Create a function `order_animals` that:

- Receives as argument a list of dictionaries;
- Order the animals by number of legs;
- Return them as a list of (animal_name, class_type, legs) tuples;

<!-- grade_id: cell-122a9ec5cefa138a, points: 0, grade: false, solution: true -->
```python
def order_animals(data):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4c4b419744a6ebf6, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_4_grading(order_animals)
print("Answer is correct. Good Job.")
```

## Question 5

The variable `data` is a list of dictionaries.

Create a function `mammal_letters` that:

- Receives as argument a list of dictionaries;
- Filter all the mammals;
- Return the average number of letters in the mammals' names. The return type must be a float;

<!-- grade_id: cell-81f72032281eb943, points: 0, grade: false, solution: true -->
```python
def mammals_letters(data):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-03842b58f0727f29, points: 3, grade: true, solution: false -->
```python
utils.b2w2_exerc_5_grading(mammals_letters)
print("Answer is correct. Good Job.")
```

## Question 6

Create an `Animal` class.

The class has three class parameters: `animal_name` (str), `animal_type` (str) and `eggs` (bool).
Implement a method called `lay_eggs` that:
- Returns random number between 0 and 10 if the animal can lay eggs (eggs parameter = True)
- Returns the sentence "I can't lay eggs..." if the animal doesn't lay eggs

<!-- grade_id: cell-786ac51bf37eef4c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f50a68ec0e515188, points: 5, grade: true, solution: false -->
```python
utils.b2w2_exerc_6_grading(Animal)
print("Answer is correct. Good Job.")
```

### Last but not least, submit your work!

To submit your work, fill your slack ID in the `slack_id` variable (as a string).

Example: `slack_id = "x-men"`

Help: if you forgot your slack ID, [read this](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c).

<!-- grade_id: cell-1a02f44222e2455e, points: 0, grade: false, solution: true -->
```python
# Submit your work!

#slack_id = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c461aefc9884225e, points: 0, grade: true, solution: false -->
```python
from submit import submit
assert isinstance(slack_id, str)

slu = 17_3
submit(slack_id, slu)
```