# SLU15 | Debugging. Exercise notebook

***

Now we're going to solve a couple of bugs ourselves.

Use the suggestions from the learning notebook, your own logic and Google (if required).

**WARNING**: 

We commented some cells out here. The reason is that they produce exceptions (and this behavior is expected), but our internal system doesn't allow exercise notebooks to contain cells that produce exceptions.

Uncomment these cells when you want to execute them, and comment them back out when you submit your solution.

## Start by importing these packages

```python
import hashlib

def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()
```

Our task is to write a program that reads a text file that contains a batch of numbers, calculate its moving average and write to another text file called `result.txt`.

The final `result.txt` file should contain all the moving averages placed at new lines.

Moving averages of a list `nums` is a list of numbers, where element `i` is equal to the average of all numbers from `nums` until index `i`(including).

For example, if we calculate moving averages for a list of numbers `[1, 2, 6, 3]` , the result will be `[1, 1.5, 3, 3]`, because in the first step we calculate the average of `[1]`, which equals `1`, in the second step we calculate the average of `[1, 2]`, which equals `1.5`, in the third step we calculate the average of `[1, 2, 6]`, which equals `3`, and in the fourth step we calculate the average of `[1, 2, 6, 3]`, which equals `3`.

For this particular exercise, it's important to implement the moving average calculations ourselves, so we shouldn't use functions like len(), sum(), mean() for the average calculation. It's okay to use these functions in any other situation.

If you uncomment and try to execute the program below, you'll get an error:

`TypeError: unsupported operand type(s) for +=: 'int' and 'str'`.


```python
# with open('document.txt', 'r') as f:
#     numbers = f.readlines()
# n_items = 0
# current_sum = 0
# means = []
# for i in numbers:
#     n_items += 1
#     current_sum += i
#     means.append(current_sum / n_items)
# with open('result.txt', 'w') as f:
#     for num in means:
#         f.write(num)
```

This error already suggests where to look for issues, but let's start from something different.

Let's divide our code into functions, so that it's easier to find bugs. It will also allow us to easily modify the program on the next steps.

### Task 1:
Create functions `read_nums`, `calculate_moving_avg`, `write_nums` from the code above.

Function `read_nums` needs to take a string `file_name` as a parameter and return the result of reading this file as a list of values.

Function `calculate_moving_avg` needs to take a list of integers `nums` as parameter and return a list of moving averages as a result.

Function `write_nums` needs to take a string `file_name` and a list of floats `nums` as parameters and write each number to the file. This function shouldn't return anything.

You don't need to change anything in the code yet. Just divide the code into functions.

<!-- grade_id: cell-c806f18640ab4195, points: 0, grade: false, solution: true -->
```python
# def read_nums(file_name):
    # ...

# def calculate_moving_avg(nums):
    # ...

# def write_nums(file_name, nums):
#     ...


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c69436878e0ec76e, points: 2, grade: true, solution: false -->
```python
nums = read_nums('document.txt')
assert isinstance(nums, list), "Are you sure your read_nums function returns a list?"
assert len(nums) == 19, "Wrong number of items in the list returned by read_nums function"
assert _hash(str(nums)) == "0de4458d18", "The result of reading the file doesn't match the expected output"
```



Now as we divided our code into functions, it's much easier to understand what's wrong with our code.

If you execute the code below, you should see the following error:

```
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)
<ipython-input-26-2c8baf7cfac4> in <module>
      1 nums = read_nums('document.txt')
----> 2 means = calculate_moving_avg(nums)
      3 write_nums('result.txt', means)

<ipython-input-25-bb04f4b79a0c> in calculate_moving_avg(nums)
     10     for i in numbers:
     11         n_items += 1
---> 12         current_sum += i
     13         means.append(current_sum / n_items)
     14     return means

TypeError: unsupported operand type(s) for +=: 'int' and 'str'
```

```python
# nums = read_nums('document.txt')
# means = calculate_moving_avg(nums)
# write_nums('result.txt', means)
```



### Task 2:

It's important to read the errors, so read the error above and answer the questions below.

Answer the questions by assigning the letter of the correct option to the `answer_i` variable, where i is the question number.

For example, `answer_1="A"`

Some questions may have more than one correct answer, so in that case type all the letters that match without spaces: 

`answer_i = "ABC"`

1. What's the type of the Exception?

```
A. IndexError
B. TypeError
C. NameError
D. SyntaxError
```

2. What's the error message.

```
A. calculate_moving_avg(nums)
B. Traceback (most recent call last)
C. unsupported operand type(s) for +=: 'int' and 'str'
D. TypeError
```

3. What is/are the line(s) of code that cause the errors?

```
A. current_sum += i 
B. means.append(current_sum / n_items)
C. write_nums('result.txt', means)
D. means = calculate_moving_avg(nums)
```

<!-- grade_id: cell-568fd9fe92bd054b, points: 0, grade: false, solution: true -->
```python
# answer_1 = ""
# answer_2 = ""
# answer_3 = ""

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2038524ce338dfa9, points: 2, grade: true, solution: false -->
```python
assert _hash(''.join(sorted(answer_1.upper()))) == '90a8cbb480', "Wrong exception type"
assert _hash(''.join(sorted(answer_2.upper()))) == 'f72385315c', "Wrong error message"
assert _hash(''.join(sorted(answer_3.upper()))) == 'd685f0ed6b', "Wrong error code lines"
```

As you should have noticed, the last error is `TypeError: unsupported operand type(s) for +=: 'int' and 'str'`, and it's the result of calling `current_sum += i` in `calculate_moving_avg`.

What does and what doesn't it mean?

It doesn't mean that the `calculate_moving_avg` function contains bugs. 

However, there was something wrong while executing `current_sum += i` inside `calculate_moving_avg`. Being more precise, we add some integer to some string in that line, and Python doesn't allow such operations.

Let's figure out why we have a string there, because our plan was to add an integer to an integer, right?



### Task 3:




Your task is to go through the code line by line, analyzing what happens with `current_sum`, what happens with `i`, and why we have a string there.

Modify the created functions in such a way, so that both `current_sum` and `i` would be integers and so that this error would disappear.

**NOTE:**

There are still some other bugs in our code, but we'll only solve them in the next tasks.

Don't edit anything that's not related to this particular exception.

<!-- grade_id: cell-949aca96afb49149, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-69d3c58b504eb6f1, points: 3, grade: true, solution: false -->
```python
nums = read_nums('document.txt')
means = calculate_moving_avg(nums)
assert isinstance(means, list), "calculate_moving_avg function is supposed to return a list"
assert len(means) == 19, "wrong number of items in the list returned by calculate_moving_avg function"
for num in means:
    assert isinstance(num, float), "Wrong value type. calculate_moving_avg function is supposed to return a list of floats"
```

If you solved the previous exercise correctly, the previos exception should have dissappeared, but there should still be another exception if you execute the code below:

```
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)
<ipython-input-27-2c8baf7cfac4> in <module>
      1 nums = read_nums('document.txt')
      2 means = calculate_moving_avg(nums)
----> 3 write_nums('result.txt', means)

<ipython-input-25-8ffcbd828a96> in write_nums(file_name, nums)
     23     with open(file_name, 'w') as f:
     24         for num in nums:
---> 25             f.write(num)
     26 
     27 

TypeError: write() argument must be str, not float
```


```python
# nums = read_nums('document.txt')
# means = calculate_moving_avg(nums)
# write_nums('result.txt', means)
```



# Task 4:


Answer the following questions by assigning the letter of the correct option to the `answer_i` variable, where i is the question number.

For example, `answer_1="A"`

Some questions may have more than one correct answer, so in that case type all the letters that match without spaces: 

`answer_i = "ABC"`



1. What's the type of the Exception?

```
A. IndexError
B. TypeError
C. NameError
D. SyntaxError
```

2. What's the error message.

```
A. write() argument must be str, not float
B. Traceback (most recent call last)
C. f.write(num)
D. write_nums('result.txt', means)
```

3. What is/are the line(s) of code that cause the errors?

```
A. current_sum += i 
B. means.append(current_sum / n_items)
C. write_nums('result.txt', means)
D. f.write(num)
```

<!-- grade_id: cell-9e13dd102b96d202, points: 0, grade: false, solution: true -->
```python
# answer_1 = ""
# answer_2 = ""
# answer_3 = ""

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-a2a266bf76e087e1, points: 2, grade: true, solution: false -->
```python
assert _hash(''.join(sorted(answer_1.upper()))) == '90a8cbb480', "Wrong exception type"
assert _hash(''.join(sorted(answer_2.upper()))) == 'b685b2f862', "Wrong error message"
assert _hash(''.join(sorted(answer_3.upper()))) == '94c05104ae', "Wrong error code lines"
```

## Task 5:

As you might have noticed, the last error is `TypeError: write() argument must be str, not float`, and it's the result of calling `f.write(num)` in `write_nums` function.

What does and what doesn't it mean?

It doesn't mean that the `write_nums` function contains bugs. 

However, there was something wrong while executing `f.write(num)` inside `calculate_moving_avg`. Being more precise, we're trying to write a float, but it's expected to be a string.

In fact, if we check the `write` method documentation [here](https://www.w3schools.com/python/ref_file_write.asp), we can find the following:

`parameter: byte --- The text or byte object that will be inserted.`

It means that write method doesn't allow to write anything besides strings and byte objects.

Seems like `num` is a float number.


Your task is to modify the code in such a way, so that `write` method works.

**NOTE:** 

There might still be some other bugs in our code, but we'll only solve them in the next tasks. Don't edit anything that's not related to this particular exception.

<!-- grade_id: cell-082945ee7566cb3c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

```python
nums = read_nums('document.txt')
means = calculate_moving_avg(nums)
write_nums('result.txt', means)
```

## Task 5

If you did everything correctly in the previous exercise, there should be no more exceptions.

But it doesn't mean that our program works 100% correctly, right?

We need to validate that the results are correct.

Let's read the task once again.

```
Our task is to write a program that read a text file that contains a batch of numbers, calculate its moving average and write to another text file called `result.txt`.

The final `result.txt` file should contain all the moving averages placed on new lines.
```

Did we meet the requirements? If you opened the `result.txt` file, you'd find out that it looks like this:

<img src="./assets/result.png" width="900">

And the task description says that the results file should contain numbers placed on new lines.

Your task is to modify the code in such a way, so that each number is on a new line.
The last line should be a number, not an empty string.

e.g., correct:

`10.5`

`3.5`

`1.2`


wrong:

`10.5`

`3.5`

`1.2`

` `

<!-- grade_id: cell-ed7689c30aba19db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-c9e2a417b3135004, points: 3, grade: true, solution: false -->
```python
nums = read_nums('document.txt')
means = calculate_moving_avg(nums)
write_nums('result.txt', means)
with open('result.txt', 'r') as f:
    content = f.readlines()
assert len(content) == 19, "Are you sure there is the same number of lines in the text file, as in the moving averages list?"
for i in range(len(content)):
    line = content[i]
    if i != len(content) - 1:
        assert line[-1:] == "\n", "Are you sure there is a newline symbol after each number (besides the last one)?"
    else:
        assert line[-1:] != "\n", "There should be no newline symbol at the end of the file"
    num = float(line.strip()) # if line.strip() can be converted to float, everything is alright
```

### And now... Mission Accomplished!!! 🎉🎉🎉
Now you can rest for this week, because your work was amaaaaaaazing (Again)!

**But don't forget...**

***

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-7281d78bfd08794b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# slack_id =
```

<!-- grade_id: cell-383debe5cc5ea67b, points: 0, grade: true, solution: false -->
```python
from submit import submit

assert slack_id is not None
submit(slack_id=slack_id, learning_unit=15)
```

## Help us improve this course!

Please help us provide even better experience for the next year students by filling up this form:
https://forms.gle/kp36emHkxPJXGqKh8

```python

```