# SLU02 Command Line & Text Editor - Exercise Notebook
***

```python
# For evaluation purposes, please import this package:
import hashlib
```

### Exercise 1: Where is home?

<img src="./assets/tortoise_home.png" width="400"/>

Imagine that you have a Linux user called `tortoise`, what is the correct path for this user's *home directory*?

* Option A: */root*
* Option B: */tortoise*
* Option C: */root/tortoise*
* Option D: */home/tortoise*

<!-- grade_id: cell-a894a58ceb39eaa2, points: 0, grade: false, solution: true -->
```python
# exercise_1 =
### BEGIN SOLUTION
exercise_1 = "D"
### END SOLUTION
```

<!-- grade_id: cell-5a2440598289b1e3, points: 1, grade: true, solution: false -->
```python
ex_1 = '50c9e8d5fc98727b4bbc93cf5d64a68db647f04f'
assert isinstance(ex_1, str), "The variable should be a string."
assert ex_1 == hashlib.sha1(bytes(exercise_1, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 2.1: Changing directory
Now, for our same `tortoise` user, if our current working directory is `/var/lib/systemd/`, how do we go this user's *home directory* ?

* Option A: pwd
* Option B: cd /home
* Option C: cd ~
* Option D: cd tortoise

<!-- grade_id: cell-29db176b41ae1c07, points: 0, grade: false, solution: true -->
```python
# exercise_2_1 =
### BEGIN SOLUTION
exercise_2_1 = "C"
### END SOLUTION
```

<!-- grade_id: cell-e4ea53657d321c19, points: 1, grade: true, solution: false -->
```python
ex_2_1 = '484d52691fcadbfabec5a318d1cf9692c7f293cbc8c1d5f22b2d839b'
assert isinstance(ex_2_1,str), "The variable should be a string."
assert ex_2_1 == hashlib.sha224(bytes(exercise_2_1, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 2.2: Changing directory
Now, for our same `tortoise` user, 
if our current working directory is `/var/lib/systemd/` and we run 
`cd ~/../..`,
what will our current working directory be? 
* Option A: /home/tortoise
* Option B: /home/tortoise/home/tortoise
* Option C: ~
* Option D: /

<!-- grade_id: cell-29db176b41ae1c00, points: 0, grade: false, solution: true -->
```python
# exercise_2_2 =
### BEGIN SOLUTION
exercise_2_2 = "D"
### END SOLUTION
```

<!-- grade_id: cell-e4ea55657d321c19, points: 1, grade: true, solution: false -->
```python
ex_2_2 = '7497e88b8c64062edfb2f30b9c5f845552b841d1fbae97ab3f224f74'
assert isinstance(ex_2_2, str), "The variable should be a string."
assert ex_2_2 == hashlib.sha224(bytes(exercise_2_2, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 3: Rename

How do we rename a non empty folder from `tortoise` to `tortoise-new`?

- Option A: *rename tortoise tortoise-new*
- Option B: *rn tortoise tortoise-new*
- Option C: *mv tortoise tortoise-new*
- Option D: *mv -r tortoise tortoise-new*

<!-- grade_id: cell-befccac89f6c99df, points: 0, grade: false, solution: true -->
```python
# exercise_3 =
### BEGIN SOLUTION
exercise_3 = "C"
### END SOLUTION
```

<!-- grade_id: cell-8e25f4730f887ad4, points: 1, grade: true, solution: false -->
```python
ex_3 = '484d52691fcadbfabec5a318d1cf9692c7f293cbc8c1d5f22b2d839b'
assert isinstance(ex_3, str), "The variable should be a string."
assert ex_3 == hashlib.sha224(bytes(exercise_3, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 4: Copy
`tortoise` is not happy. This user needs to copy the `baby-tortoise` folder inside the `tortoise-family` folder. 

What command should we use?

- Option A: *mv baby-tortoise tortoise-family*
- Option B: *cp baby-tortoise tortoise-family*
- Option C: *cp -a baby-tortoise tortoise-family*
- Option D: *cp -r baby-tortoise tortoise-family*

<!-- grade_id: cell-d629c9437fa027f7, points: 0, grade: false, solution: true -->
```python
# exercise_4 =
### BEGIN SOLUTION
exercise_4 = "D"
### END SOLUTION
```

<!-- grade_id: cell-9b757ee86b6b31a4, points: 2, grade: true, solution: false -->
```python
ex_4 = '7497e88b8c64062edfb2f30b9c5f845552b841d1fbae97ab3f224f74'
assert isinstance(ex_4, str), "The variable should be a string."
assert ex_4 == hashlib.sha224(bytes(exercise_4, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 5: Create Directories and files

<img src="./assets/directories.png" width="400"/>

Create the Directory structure in the image above, including the files.

were you able to do it?

* Option A: YES
* Option B: NO, and I'll ask some question on Slack to help me figure this out!

<!-- grade_id: cell-d629c9437fa026f7, points: 0, grade: false, solution: true -->
```python
# exercise_5 =
### BEGIN SOLUTION
exercise_5 = "A"
### END SOLUTION
```

<!-- grade_id: cell-9b757ef86b6b31a4, points: 2, grade: true, solution: false -->
```python
ex_5 = '5cfe2cddbb9940fb4d8505e25ea77e763a0077693dbb01b1a6aa94f2'
assert isinstance(ex_5, str), "The variable should be a string."
assert ex_5 == hashlib.sha224(bytes(exercise_5, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 6: Remove

<img src="./assets/tortoise.jpg" width="400"/>

Before you start questioning whether you have a weird fixation over tortoises, you might want to remove some directories and files.

What is the correct command to remove the `tortoisephilia` folder (which has some tortoise pictures)?

* Option A: `drop tortoisephilia`
* Option B: `rm tortoisephilia/*`
* Option C: `remove tortoisephilia`
* Option D: `rm -r tortoisephilia`

<!-- grade_id: cell-d629c9437fa126f7, points: 0, grade: false, solution: true -->
```python
# exercise_6 =
### BEGIN SOLUTION
exercise_6 = "D"
### END SOLUTION
```

<!-- grade_id: cell-9b757ef81b6b31a4, points: 2, grade: true, solution: false -->
```python
ex_6 = '7497e88b8c64062edfb2f30b9c5f845552b841d1fbae97ab3f224f74'
assert isinstance(ex_6, str), "The variable should be a string."
assert ex_6 == hashlib.sha224(bytes(exercise_6, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 7:
Our `tortoise` created a `tortoise_codes.py` file with some Python code (as all good tortoises do), now he needs to check what is stored in the file in a simple and easy way.

What is the correct command for that procedure?

- Option A: *cat tortoise_codes.py*
- Option B: *read tortoise_codes.py*
- Option C: *ls tortoise_codes.py*
- Option D: *touch tortoise_codes.py*

<!-- grade_id: cell-1792bbe6ebc8564d, points: 0, grade: false, solution: true -->
```python
# exercise_7 =
### BEGIN SOLUTION
exercise_7 = "A"
### END SOLUTION
```

<!-- grade_id: cell-d57ae086b71ffed4, points: 2, grade: true, solution: false -->
```python
ex_7 = 'ad14aaf25020bef2fd4e3eb5ec0c50272cdfd66074b0ed037c9a11254321aac0729985374beeaa5b80a504d048be1864'
assert isinstance(ex_7,str), "The variable should be a string."
assert ex_7 == hashlib.sha384(bytes(exercise_7, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 8: Visual Studio Code (VS Code)

Managing to install and set VS Code up is part of this SLU, so if you have any problems please do contact an instructor! 

It is very important that you get it running properly.

- Option A: Rock On! Everything is working fine :)
- Option B: NO, I am still struggling with it. It's this week's SLU instructor's fault!!!

<!-- grade_id: cell-82043c053fb1c39d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_8 = "A"
### END SOLUTION
# exercise_8 =
```

<!-- grade_id: cell-fa61be6aa4097dad, points: 4, grade: true, solution: false -->
```python
ex_8 = '559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd'
assert isinstance(ex_8,str), "The variable should be a string."
assert ex_8 == hashlib.sha256(bytes(exercise_8, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 9: Running basic Python programs

To finish this SLU in style:
1. create a folder in the location of your choice
1. open that folder in VS Code
1. in VS Code, create a Python file that prints your name, and 2 other files that print 2 other names
1. activate your virtual environment
1. run those Python files and check that the names are printed!

Were you able to do this?

- Option A: No, I've been stuck on this for more than 15 min! I need to ask for help on Slack!
- Option B: Yes, I did it!

<!-- grade_id: cell-8204dc053fb1c39d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_9 = "B"
### END SOLUTION
# exercise_9 =
```

<!-- grade_id: cell-fa61be6aa4d97dad, points: 4, grade: true, solution: false -->
```python
ex_9 = 'df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c'
assert isinstance(ex_9,str), "The variable should be a string."
assert ex_9 == hashlib.sha256(bytes(exercise_9, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-2d96ea28de941280, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
slack_id = "UTS63FC02"
### END SOLUTION
# slack_id =
```

<!-- grade_id: cell-cbb863bacf2a8bca, points: 0, grade: true, solution: false -->
```python
from submit import submit

assert slack_id is not None
submit(slack_id=slack_id, learning_unit=2)
```