# SLU05 - Git Intermediate

I hope you enjoyed the **Learning Notebook**! Like SLU03, the Exercise Notebook is multiple choice so it's really important that you have a good grasp of everything before starting it. 

**Good luck!**

```python
import hashlib
import requests
from submit import submit
```

## Exercise 1: Git remotes


What is the main purpose of a remote repository?
- Option **A**: To serve as a backup to your local files.
- Option **B**: Tracking and committing changes to your files.
- Option **C**: Sharing work with your teammates.
- Option **D**: See who committed what changes to a project and hunt them down.

<!-- grade_id: cell-14d1430f3cf9833a, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_1 the option letter that you think is correct.

### BEGIN SOLUTION
exercise_1 = "B" 
### END SOLUTION
# exercise_1 =
```

<!-- grade_id: cell-fccabbb7b3b73a9a, points: 2, grade: true, solution: false -->
```python
exercise_1_hash = "df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_1,str), "The variable should be a string."
assert exercise_1.isupper(), "Please use uppercase for your answer."
assert exercise_1.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_1_hash == hashlib.sha256(bytes(exercise_1, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 2: Git remotes

What is the function of `git push`?

- Option **A**: Updates remote refs along with associated objects.
- Option **B**: Push all changes to the remote branch
- Option **C**: Push all changes to the main branch
- Option **D**: Updates local refs along with associated objects

<!-- grade_id: cell-349ecbd85ce6985f, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_2 the option letter that you think is correct.

### BEGIN SOLUTION
exercise_2 = "A" 
### END SOLUTION
# exercise_2 =
```

<!-- grade_id: cell-dbaa29386c745914, points: 2, grade: true, solution: false -->
```python
exercise_2_hash = "559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_2,str), "The variable should be a string."
assert exercise_2.isupper(), "Please use uppercase for your answer."
assert exercise_2.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_2_hash == hashlib.sha256(bytes(exercise_2, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 3: Git remotes

What is the function of `git clone`?

- Option **A**: Creates a new Git repository on your local machine
- Option **B**: Creates a copy of an existing Git repository on your remote repository
- Option **C**: Creates a copy of the files in a Git repository, without any history 
- Option **D**: Creates a copy of an existing Git repository on your local machine

<!-- grade_id: cell-548374c265341b10, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_3 the option letter that you think is correct.

### BEGIN SOLUTION
exercise_3 = "D" 
### END SOLUTION
# exercise_3 =
```

<!-- grade_id: cell-9bb7f9bb7974e4d7, points: 4, grade: true, solution: false -->
```python
exercise_3_hash = "3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_3,str), "The variable should be a string."
assert exercise_3.isupper(), "Please use uppercase for your answer."
assert exercise_3.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_3_hash == hashlib.sha256(bytes(exercise_3, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 4: Git remotes

Did your review how to create and clone a repository [here](https://github.com/LDSSA/ds-prep-course-2021#12-setup-git-and-github)? Good!

Go to [Github](https://github.com/), and create a new repository called *SLU05* and then clone it to your machine.

**Important:** Make sure you set your repository to *public*!

It's the same thing you did for setting up the *ds-prep-workspace* repository, but with a different name:

1. You don't need to setup git and Github again, that part is already covered
2. Follow the instruction on [Setup your Workspace Repository](https://github.com/LDSSA/ds-prep-course-2021#13-setup-your-workspace-repository), but name your repository *SLU05*
3. Follow the instruction on [Cloning the workspace](https://github.com/LDSSA/ds-prep-course-2021#14-clone-the-ds-prep-course-2021-repository) for cloning the remote repository into your machine, but for the repository you've just created.

That's it for this exercise. We will use this new repository for the remaining exercises

<!-- grade_id: cell-f7886da512b156a0, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_4 the URL to your SLU05 repository on github, in between quotes

### BEGIN SOLUTION
exercise_4 = "https://github.com/LDSSA/ds-prep-course-2021" 
### END SOLUTION
# exercise_4 = ""
```

<!-- grade_id: cell-b99b518a1ed4e6d9, points: 4, grade: true, solution: false -->
```python
assert isinstance(exercise_4,str), "The variable should be a string."
request = requests.get(exercise_4)
assert request.status_code == 200, "Can't find the repository! Have you set it as public?"
```

## Exercise 5: Pushing changes into a remote repository

Do the following workflow:

1. On your **local SLU05 repository**, create a groceries.txt file and add some grocery items to it
2. Stage and commit your changes
3. You just forgot the butter! Add it to your grocery list (groceries.txt)
4. Stage and commit your changes (again!)
5. Verify, by using `git status` and `git log`, that you have 2 commits in the local repository that are not in the remote repository
6. Push the changes to the remote repository
7. Verify, by using `git status` and `git log`, that your local and remote repositories are in sync.


**Note:** You don't have to follow the exact workflow above. Feel free to experiment on this one :) The idea is to get used to the commands. You will use them a lot, so start getting that muscle memory now!

## Exercise 6: Pulling changes from remote

Now try to do the same workflow we did on the *learning notebook*:

1. Use your browser to go to your SLU repository on Github (probably https://github.com/"your_github_userid_here"/SLU05)
2. Edit your groceries.txt file on the GitHub website - Your mum wants eggs to make a pudding! ;) 
3. Commit your changes on the website
4. Check, on your machine, that `git status` and `git log` don't show any differences. *If you feel like it, disconnect your computer from the network and check that those commands keep working :)*
5. Try using `git remote` to check if the repositories are synced
6. Pull the changes you did in the remote into your local repository

## Exercise 7. Stash Changes

Follow the workflow:

1. Create a new file named winelist.txt.
2. Add a few wines you'd like to drink (*yey quarantine for making day drinking a thing!*).
3. Stash the changes (**hint:** remember that stash does not save untracked files).
4. Add a new item to the groceries.txt file (*you always need a massive amount of toilet paper in a global crisis*).
5. Commit those changes.
6. Pop the changes you have in the stash.
7. If you're feeling like it, add a nice bottle of Moscatel to your winelist. When you're ready, commit-
8. Check, using whatever method you like, that the local repository is 2 commits ahead of the remote repository
9. Push the commits to the remote repository

```python

```

## Exercise 8. Create and push branches¶
Try the following workflow:

1. Create a new branch based on main. Name it "Ubereats"
2. In the new branch create a new file restaurantes_to_try.txt and add some content to it
3. Commit the changes and push the branch to the remote
4. Check on GitHub that the branch is there

## Exercise 9. Merging branches
**Warning:** I'm leaving out a lot of step-by-step instructions in this exercise and the next one. Remember that the *learning notebook* is your friend

1. Merge the branch you created on exercise 7 into main. Verify that it was a **fast-forward** merge.
2. Do a 3-way merge of a branch into master. This might be a tough one. Go through the learning notebook carefully. Everything you need is in there.
3. Push master to the remote.
4. **Bonus:** Do a Google search on how to delete a branch and delete the branches you created, both on the local and remote repositories.

## Exercise 10: Collaborating with Git

What's the difference between a "pull request" and a "branch"?

**Select the false statement**

- Option **A**: A "branch" is just a separate version of the code.
- Option **B**: A "pull request" is when someone who has been working on a project branch wants to merge their work into another branch and requests a code review before doing it.  
- Option **C**: You can only do pull requests to the main branch, where you can branch from any other branch, including the main.
- Option **D**: A pull request lets the project maintainer know that an update is ready to be integrated

<!-- grade_id: cell-4fc49f4aad340d87, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_10 the option letter that you think is correct.

### BEGIN SOLUTION
exercise_10 = "C" 
### END SOLUTION
# exercise_10 =
```

<!-- grade_id: cell-813abdc5b03ac7ce, points: 4, grade: true, solution: false -->
```python
exercise_10_hash = "6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_10,str), "The variable should be a string."
assert exercise_10.isupper(), "Please use uppercase for your answer."
assert exercise_10.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_10_hash == hashlib.sha256(bytes(exercise_10, encoding='utf8')).hexdigest(), "Wrong answer"
```

## Exercise 11: Conflicts

How can conflict in git be resolved?

**Select the false statement**

- Option **A**: You can abort the merge and discard the alterations that are in conflict.
- Option **B**: You can dit the files to fix the conflicting changes and then add the resolved files by running `git add`, followed by `git commit`. 
- Option **C**: You can choose what version you want to keep of the two conflicting branches.
- Option **D**: After a conflict is created, you cannot go back to your previous state.

<!-- grade_id: cell-bb2837e739fc4c32, points: 0, grade: false, solution: true -->
```python
# Assign to the variable exercise_11 the option letter that you think is correct.

### BEGIN SOLUTION
exercise_11 = "D" 
### END SOLUTION
# exercise_11 =
```

<!-- grade_id: cell-9e36d4171ea208f0, points: 4, grade: true, solution: false -->
```python
exercise_11_hash = "3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43"
options = ["A", "B", "C", "D"]
assert isinstance(exercise_11,str), "The variable should be a string."
assert exercise_11.isupper(), "Please use uppercase for your answer."
assert exercise_11.upper() in options, "Please choose one of the options: %s" % ', '.join(str(opt) for opt in options)
assert exercise_11_hash == hashlib.sha256(bytes(exercise_11, encoding='utf8')).hexdigest(), "Wrong answer"
```

**Now you're ready to make the most out of Git!** You'll be using it throughout the rest of the Prep Course so it's important that you two become good friends!

Keep the git workflow in your heart, and use it whenever you find a typo in a Learning/Exercise Notebook! I'll give you a golden start for it! :D

## Submit your work!
To submit your work, get your slack id and fill it in the slack_id variable.

Example: slack_id = "UTS63FC02"

<!-- grade_id: cell-a3dc67590ead160d, points: 0, grade: false, solution: true -->
```python
#slack_id = ...  
### BEGIN SOLUTION
slack_id = "U01RTPQGEH0"
### END SOLUTION
```

```python
assert isinstance(slack_id,str)
submit(slack_id, 5)
```

Output:
```
Success
```