# SLU 09 - Exercise notebook 

Before you start:
>- This notebook is divided into 8 questions: 
    - Question 1 to 6 are easy, practice questions, expect that part to take about 2 hours. 
    - Questions 7 and 8 are hard question, expect to spend at least another 3 hours on them. They are close to the level of difficulty to pass the Academy test. 
    
> - If you are unsure about your answer, run the asserts (the tests). Even if it is a bit incomplete, run the asserts, and see if they give you clues. 
    - In the last 2 questions the asserts will be a bit less helpful, so that you can get some practice of dealing with genuine python errors. 
    - Read the instructions carefully. Everything that is there is critical
    - Don't worry if these concepts feel weird at first. Classes are not easy, you are now getting into advanced stuff! 
    
> - If you are stuck, remember to:
    - Play around! Open a new cell, and mess around with the objects. Instantiate them and see if they complain. See what they print, what they return, etc. Remember, in real life there are no asserts to guide you!  
    - Always read the error messages carefully. They should be good clues about what is wrong. 
    - As always, don't be afraid to ask questions! 


----

### Exercise 1 - Playing around with objects 

This exercise is very simple. We are just going to make a couple new fruits, like we did in the first Learning Notebook. 

Make the following two fruits: 
- pears, which is called "pears", has 4 units, costs 3 euros per unit, and expires in 14 days
- tangerines, which is called "tangerines", has 2 units, costs 6 euros per unit, and expires in 2 days

This exercise is very easy, just look up how it was done in the Learning notebook and do a couple of them here. 

Reminder: the Fruit class has the following attributes: 
- `price_per_unit`
- `name`
- `nr_units`
- `days_until_expired`

> Protip: Note that you can use the tab auto-complete to help you.  
> Try writing `Fruit(pr` and hitting tab. It should auto-complete to price! 


<!-- grade_id: cell-4640dc55a4c658cd, points: 0, grade: false, solution: true -->
```python
from utils import Fruit

# pears = ...
# tangerines = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f4f479ecb4b7c3ef, points: 1, grade: true, solution: false -->
```python
from utils import exercise_1_grading
assert (pears and tangerines), 'Did you forget to make one of the fruits?'
exercise_1_grading(pears, tangerines)
print('Great success! Exercise passed!')
```

----
### Exercise 2 - add up the price of fruits 

Great! For your second exercise you will add the price of all the fruits: 
- apples
- bananas 
- oranges 
- pears
- tangerines 

The price is the product of the price per unit and number of units. But remember, the Fruit object already comes with a `calculate_price` method! 

Create a function called `calculate_price_of_all_fruits`, which takes a list called `all_fruits` that contains fruit objects and returns the total price. We've already filled in some things for you. 

Remember: the function should take any list of fruits, don't hard code any particular list! 

<!-- grade_id: cell-d5a600cefd3c41c7, points: 0, grade: false, solution: true -->
```python
def calculate_price_of_all_fruits(all_fruits): 
    """
    Calculates the price for every fruit in the all_fruits list, 
        and returns the total price of fruits in the list
    Args: 
        all_fruits (list): a list with fruit objects 
        
    Returns: 
        total_price (int): the total price of the fruits in the list
    """
    
    # hint: this may require a for loop, and a variable to keep track of the total price 

    # YOUR CODE HERE
    raise NotImplementedError()
    
# this should return 24 
print(calculate_price_of_all_fruits([pears, tangerines]))
```

<!-- grade_id: cell-ee1a997f801ac0bb, points: 1, grade: true, solution: false -->
```python
from utils import exercise_2_grading
exercise_2_grading(calculate_price_of_all_fruits)
print('Great success! Exercise passed!')
```

---
### Exercise 3 - Using objects with other objects 

In this exercise you will do something very similar to the learning notebook. You will simply instantiate a basket, create some toilet paper, and then will add and remove stuff until you have the following: 

- Total value of the basket should be between 30 and 35 
- the basket must have at least 3 types of fruit
- the basket must not contain any items that will expire in the next 8 days 
- the basket must have at least some toilet paper. It should be ... 
    - named "Soft but cheap" 
    - cost 4 euros
    - have "Double leaf". 
    - The expiration date and number of units are up to you, they won't be tested

Reminders: 
```
- You already have some instance variables to work with: 
    - apples 
    - oranges
    - bananas 
    - pears 
    - tangerines 
- You only have the Class variable Toiletpaper to work with. 
- You'll need to create some toilet paper before you can use it. 
    - Call the variable toilet_paper.
    - Give it nr_units = 1 
```

_Note: obviously you can figure out exactly what you need to add ahead of time without trial and error. However, it's good practice to just use the objects. Try adding, checking the price, removing, etc until you have an answer that meets all the conditions._

<!-- grade_id: cell-cffd2f4f2528de89, points: 0, grade: false, solution: true -->
```python
from utils import Toiletpaper, Basket, get_fruits
apples, bananas, oranges = get_fruits()  # you already have some fruits

# my_basket = ... 
# toilet_paper = ...
# my_basket.<do something> 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-277ef4e6f1a47d0d, points: 2, grade: true, solution: false -->
```python
from utils import exercise_3_grading 
exercise_3_grading(my_basket, toilet_paper)
print('Great success! Exercise passed!')
```

----
### Exercise 4: Making a basic object 

Ok, here is what is going ot happen. We're gonna move house. And we're going to have to pack all our stuff. 

We will take a bunch of `Item`s in our house, and put them into `Boxes`. 

Let's start with creating some `Item`s. Here is what you need to do: 

Create an `Item` class, which will have 3 attributes: 
- name 
- weight 
- fragile (this will be boolean)
- broken (this will be boolean, always starts out as False) 

Create the following items: 
- `chair` which will weigh 15, be called "kitchen chair", and not be fragile 
- `vase` which will weigh 3, be called "ming vase", and be fragile 
- `pillow` which will weigh 1, be called "bed pillow", and not be fragile  

<!-- grade_id: cell-2e5ef92f86968bc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-61f3abe376b0ed6d, points: 2, grade: true, solution: false -->
```python
from utils import exercise_4_grading 
assert Item, "please create the Item class!"
assert all([vase, chair, pillow])
exercise_4_grading(Item, chair, vase, pillow)
print('Great success! Section passed!')
```

---
### Exercise 5: working with methods 

For this exercise we will create a `BasicBox` class (note: we'll create a `Box` class in the next exercise). This box will allow you to add and remove `Item`s, and will have a weight limit.

First step, let's create the `BasicBox` object. It will have the following attributes: 
- `contents`: a list that will start empty. Will will use it to keep track of items. 
- `weight_limit`: a number specific to each new box we create (different `Box`es will have different weight limits)
- `name`: a string name 
- `broken`: boolean, whether the box has been broken by having too much weight. Starts as False. 
- `weight`: defaults to zero (the weight of the box is negligible) 

The `BasicBox` should also have a method called `add_item` that takes an `item` and adds it to the `contents`. 

<!-- grade_id: cell-a9722b9debaeb161, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a9853f7b4e57fe87, points: 2, grade: true, solution: false -->
```python
from utils import exercise_5_grading 
assert BasicBox
exercise_5_grading(BasicBox)
```

---
### Exercise 6: calling methods with other methods 
Ok, our box is a bit too basic. Let's make a proper box (feel free to copy your code from the previous exercise as this one is an extension of it) 

This time we'll create a `Box` object, that allows us to add items, but has a few more methods: 

- `calculate_total_weight`: calcualtes the weight of all `items` in the `contents`, and updates the `weight` attribute. 

- `check_if_weight_limit_exceeded`: this method will call `calculate_total_weight`, then check if the total weight is over the box's `weight_limit`. If it has been exceeded, the box should change its `broken` attribute to `True`, and `break_all_fragile_items`

- `break_all_fragile_items`: change the `broken` attribute of all `fragile` items in the `contents` to `True` (don't break the non-fragile items!) 

_Note 1: don't return the total_weight, update the `self.weight` attribute of the box itself._  
_Note 2: you will probably want to use some of your Items from exercise 4 to test whether your Box does what is expected._

<!-- grade_id: cell-ef56f4c9a022e15a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9eae48817b9d3b7f, points: 4, grade: true, solution: false -->
```python
from utils import exercise_6_grading
assert Box, "Implement Box!"
exercise_6_grading(Box)
print('Great success! Section passed!')
```

----
> Before you continue:
- _The exercises up to here were good for practice. The last two are hard ones._ 
- The last 2 exercises have no new material. They will just have less hand-holding and be more demanding. If you get in trouble, see your own solutions above for clues!
- To give you a more realistic idea of debugging in the real the world we won't tell you where your mistakes are, but simply whether you are passing the tests (whether it does what is expected). You will have to use the errors that python provides, as will be the case in the real world. 
- You can of course ask for help at any time, but try suffering through these two on your own if you can, it's very good practice for the real world :)  
- _Expect to take quite a while (probably over 2 hours) to get exercise 7 and 8 right._
- _It's often useful to draft the "skeleton" of your code before jumping in. What are the steps? What are the methods of each class you will need? Sometimes writing on a piece of paper is a good support._ 
- _This section will test you on what you learned, but also includes functions, and other things you have learned before. You may need to check a few past SLUs for this one_

- _After you create a class, make an instance, see if it behaves the way you expect. It's crucial to gain the habit of testing your stuff as you go, to see if it is working._
    - it's your responsibility to create instances, find the actual error, and try to fix it!_ 
- _This is the last block in this SLU! If you got this, you are getting pretty good at objects, and are probably ready to deal with the Academy entry exercise that has this SLU's material._

---
### Exercise 7: let's make a SportsBall league! 

You know those games where there are professional sportsball leagues, with players, teams and competitions? Yeah, neither do I, and it will show in this exercise. 

Our objective is to create two `Competition`s, each of which with 5 `Team`s, each of which have 10 `Player`s. 

```python
# remember to tell them that the attempt to score actually needs to return stuff.
# test for their correct choice of defaults (everywhere)
```

<!-- grade_id: cell-05a39eba6e7e96c1, points: 0, grade: false, solution: true -->
```python
import numpy as np  # we will need this for attempts to score 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e429403c7f5eff78, points: 4, grade: true, solution: false -->
```python
from utils import exercise_7_validate_Player, exercise_7_validate_Team
import numpy as np 
assert Player, team
exercise_7_validate_Player(Player)
exercise_7_validate_Team(Team, Player)
print('Great success! Section passed!')
```

---
### Exercise 8: putting your classes to use! 

This one will just be a bit of fun. You will create a few teams, and then will use your classes to play some games, and in the end will print the results. 

<!-- grade_id: cell-a43a4e57177eca9d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1d7bb58d538228da, points: 4, grade: true, solution: false -->
```python
from utils import exercise_7_validate_League
exercise_7_validate_League(League, Team, Player)
print("Great Success! Section passed!")
```

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-555cf377c715246c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# slack_id =
```

<!-- grade_id: cell-52a14f9f098dd4f6, points: 0, grade: true, solution: false -->
```python
from submit import submit

assert slack_id is not None
submit(slack_id=slack_id, learning_unit=9)
```