# SLU14 - Modules and Packages 

#### Start by importing the following module

```python
#used for evaluation
import utils
```

### Q1 What is a module in python?

a) A set of attributes;   
b) A set of files with the `.py` extension;   
c) A file with the `.py` extension;   
d) A file with the `.ipynb` extension;

<!-- grade_id: cell-ec299e14ecce1f27, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-60bb3a2487965c64, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q1, "Wrong answer."
```

### Q2 How should we organize code in a module?

a) In attributes: classes, functions and variables;      

b) With no special organization, but with all code inside `if __name__=="__main__"`;   

c) With no special organization at all;   

d) In attributes: classes, functions and variables. If we also want to run it as a script, the entry point should be inside  `if __name__=="__main__"`;

<!-- grade_id: cell-2f1a5dfccc2eee9b, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-18368bb83c4a41b4, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q2, "Wrong answer."
```

### Q3 Why should modules be organized that way? 

a) Just because it is better organized and easier to debug;   
b) In that way we can use the attributes outside of the file and still run it as a script.

<!-- grade_id: cell-7a2c49000d940397, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cd2427e89b7707cb, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q3, "Wrong answer."
```

### Q4 What is NOT an advantage of organizing our projects by modules? 

a) It enables us to separate code by features and avoid duplicated code;   
b) It makes the code easier to debug, maintain and scale;   
c) It allows us to run Python scripts;   
d) It allows us to package our modules together and use them in other projects;  

<!-- grade_id: cell-a4658c8feaa4f078, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9a98124ab22debe2, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q4, "Wrong answer."
```

### Q5 What should we write inside our module, so that a part of the code only gets executed if the module is run as a script?

a) Nothing;   
b) `if __name__=="__main__"`  
c) `if __main__=="__name__"`  
d) `if main:`   

<!-- grade_id: cell-aea5631b25c21caa, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-99cbfb671d0969b3, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q5, "Wrong answer."
```

### Q6 Inside directory `shapes`, create a module named `triangle`  with the following attributes:

- A `class` named `Triangle` with two input variables, `base` and `height`, that are attributes of the class. This class should have also a method named `get_area` that calculates and returns the area in the following way: `(base * height) / 2`.    
- A function named `is_equilateral` that accepts three variables as an input. If they have all the same value, return True, otherwise return False.   
- A variable named `description` with the following string value `this is a module named triangle`. 

To evaluate your module, create a variable named done with the value True.

<!-- grade_id: cell-8d65782fc8519012, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ce448ffff172a0f2, points: 0, grade: true, solution: false -->
```python
if done:
    from shapes import triangle
    assert triangle, "Did you store a module named triangle inside shapes folder?"
    this_triangle = triangle.Triangle(10, 30)
    assert this_triangle.base == 10, "Did you set base as an attribute of class Triangle?"
    assert this_triangle.height == 30, "Did you set height as an attribute of class Triangle?"
    assert int(this_triangle.get_area())==150, "Did you set get_area correctly as a method of class triangle?"
    assert triangle.is_equilateral, "Did you create a function named is_equilateral?"
    assert triangle.is_equilateral(1, 1, 1), "is_equilateral is not returning True when all the inputs are equal"
    assert triangle.is_equilateral(0, 1, 1)==False, "is_equilateral is not returning False when all the inputs are not equal"
    assert triangle.description == "this is a module named triangle", "Did you set description well?"
```

### Q7 Import the class Book, from module book, and create a new object from that.

The `library` folder is in the current directory. Explore its contents.
Import the class Book , that belongs to module book. Do this import in a way that the Book class name has no prefix. This doesn't mean that we cannot have dots in our import statement.
Afterwards, create a new book, named this_book, with name=prep-course, author=ldssa and year=2020. Hint: Use from ___ import ___ statement.

<!-- grade_id: cell-8dafeb32a322444f, points: 0, grade: false, solution: true -->
```python
# from _ import _
# YOUR CODE HERE
raise NotImplementedError()
#this_book = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ac799d1b55a4cba, points: 0, grade: true, solution: false -->
```python
assert Book, "Book class was not imported with success"
assert isinstance(this_book, Book), "this_book is not of type Book"
assert this_book.name == "prep-course", "this_book has the wrong name"
assert this_book.author == "ldssa", "this_book has the wrong author"
assert this_book.year == 2020, "this_book has the wrong year"
```

### Q8 What is a package?

a) A folder;   
b) A set of modules that share the same directory;

<!-- grade_id: cell-17731565ade4c8d8, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed761f5a351a7094, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q8, "Wrong answer."
```

### Q9 Exercise

Build a python package called `program` with the following directory structure:

```
program
|____main.py
|____A
| |____module_a.py
|____B
| |____module_b.py
```

`module_a.py` should define the following function:

```python
def get_one():
    return 1
```

`module_b.py` should define the following function:

```python
def get_two():
    return get_one() + get_one()
```

`main.py` should define the following statement:

```python
print(get_two())
```

You should then be able to run the program in a terminal and get the output:

```
> python program/main.py
2
```

Beware:
* the directory diagram above may be missing something... (check the learning notebook)
* the code for each of the files may be missing something... (check the learning notebook)

Were you able to complete this exercise? (if not, ask for help!)

a) Yes   
b) No

<!-- grade_id: cell-8df2491f6b0042a0, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-839cdbbdbfebce43, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q9, "Wrong answer."
```

### Q10 What is the mandatory file in a directory with modules for it to be considered a package?

a) `setup.py`   
b) `README.md`   
c) `LICENSE`   
d) `__init__.py`

<!-- grade_id: cell-471e5327dcc8e7b7, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9d78f1dc248bb974, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q10, "Wrong answer."
```

### Q11 What is the unix command used to manage packages in a python environment?

a) cat   
b) pip   
c) python   
d) tree   

<!-- grade_id: cell-6d4d775081c8be95, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-342a8848d6c5e6fe, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q11, "Wrong answer."
```

# Submit your work!

To submit your work, [get your slack id](https://moshfeu.medium.com/how-to-find-my-member-id-in-slack-workspace-d4bba942e38c) and fill it in the `slack_id` variable.

Example: `slack_id = "UTS63FC02"`

<!-- grade_id: cell-ac5c752ffe206fec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# slack_id =
```

<!-- grade_id: cell-d0d1b1294247a44e, points: 0, grade: true, solution: false -->
```python
from submit import submit

assert slack_id is not None
submit(slack_id=slack_id, learning_unit=14)
```